SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: kalish250 has SAT parents the problem 
 RESULT: kalish250_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: kalish250 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents kalish250 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a cth of kalish250 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_3 c_0_4 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_5 is a thm of c_0_3 c_0_4 
WARNING: c_0_6 has UNS parents c_0_3 c_0_5 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_3 c_0_5 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 5.95
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN411+1/E---3.2.0.THM-CRf.s
fof(kalish250,conjecture,
    ( ! [X1,X2,X3] : f(X1,X2,X3)
  <=> ~ ? [X4,X5,X6] : ~ f(X4,X5,X6) ),
    file('/export/starexec/sandbox2/tmp/tmp.AMADGAUWL3/E---3.1_6157.p',kalish250),
    [verified_file(kalish250_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ! [X1,X2,X3] : f(X1,X2,X3)
    <=> ~ ? [X4,X5,X6] : ~ f(X4,X5,X6) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[kalish250])]),
    [nnpp(kalish250),verified_file(c_0_1_thm),verified(cth),nnpp(kalish250),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X13,X14,X15,X16,X17,X18] :
      ( ( ~ f(esk1_0,esk2_0,esk3_0)
        | ~ f(esk4_0,esk5_0,esk6_0) )
      & ( f(X13,X14,X15)
        | f(X16,X17,X18) ) ),
    inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])]),
    [nnpp(kalish250),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ( ~ f(esk1_0,esk2_0,esk3_0)
    | ~ f(esk4_0,esk5_0,esk6_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish250),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X6,X5,X4,X3,X2,X1] :
      ( f(X1,X2,X3)
      | f(X4,X5,X6) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish250),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X3,X2,X1] : f(X1,X2,X3),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_3,c_0_4]),c_0_4]),
    [nnpp(kalish250),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3,c_0_5]),c_0_5])]),
    [proof,nnpp(kalish250),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN411+1/E---3.2.0.THM-CRf.s
