SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel29 has SAT parents the problem 
 RESULT: pel29_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: pel29 is a thm of the problem 
SUCCESS: pel29_2 has SAT parents the problem 
 RESULT: pel29_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel29_2 is a thm of the problem 
SUCCESS: pel29_1 has SAT parents the problem 
 RESULT: pel29_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: pel29_1 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_3 has SAT parents pel29 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a cth of pel29 
SUCCESS: c_0_4 has SAT parents c_0_3 (Theorem esa)
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.01
FAILURE: c_0_4 fails to be a thm of c_0_3 (Theorem esa)
SUCCESS: c_0_3 has SAT parents c_0_4 (Inverted esa)
 RESULT: c_0_4_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_3 is a thm of c_0_4 (Inverted esa)
FAILURE: c_0_4 fails to be a esa of c_0_3
 ASSUME: c_0_4 is a esa of c_0_3
SUCCESS: c_0_5 has SAT parents c_0_4 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_4 
SUCCESS: c_0_6 has SAT parents pel29_2 (Theorem esa)
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.04
FAILURE: c_0_6 fails to be a thm of pel29_2 (Theorem esa)
SUCCESS: pel29_2 has SAT parents c_0_6 (Inverted esa)
 RESULT: c_0_6_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel29_2 is a thm of c_0_6 (Inverted esa)
FAILURE: c_0_6 fails to be a esa of pel29_2
 ASSUME: c_0_6 is a esa of pel29_2
SUCCESS: c_0_7 has SAT parents c_0_5 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of c_0_5 
SUCCESS: c_0_8 has SAT parents c_0_6 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_8 is a thm of c_0_6 
SUCCESS: c_0_9 has SAT parents c_0_4 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_9 is a thm of c_0_4 
SUCCESS: c_0_10 has SAT parents c_0_4 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_4 
SUCCESS: c_0_11 has SAT parents c_0_7 c_0_8 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_11 is a thm of c_0_7 c_0_8 
SUCCESS: c_0_12 has SAT parents c_0_4 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_4 
SUCCESS: c_0_13 has SAT parents c_0_4 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_13 is a thm of c_0_4 
SUCCESS: c_0_14 has SAT parents c_0_4 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_4 
SUCCESS: c_0_15 has SAT parents c_0_4 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_15 is a thm of c_0_4 
SUCCESS: c_0_16 has SAT parents c_0_4 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_16 is a thm of c_0_4 
SUCCESS: c_0_17 has SAT parents c_0_9 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_17 is a thm of c_0_9 
SUCCESS: c_0_18 has SAT parents c_0_10 c_0_11 c_0_12 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_18 is a thm of c_0_10 c_0_11 c_0_12 
SUCCESS: c_0_19 has SAT parents c_0_13 c_0_11 c_0_14 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_19 is a thm of c_0_13 c_0_11 c_0_14 
SUCCESS: c_0_20 has SAT parents pel29_1 (Theorem esa)
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.04
FAILURE: c_0_20 fails to be a thm of pel29_1 (Theorem esa)
SUCCESS: pel29_1 has SAT parents c_0_20 (Inverted esa)
 RESULT: c_0_20_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel29_1 is a thm of c_0_20 (Inverted esa)
FAILURE: c_0_20 fails to be a esa of pel29_1
 ASSUME: c_0_20 is a esa of pel29_1
SUCCESS: c_0_21 has SAT parents c_0_4 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_21 is a thm of c_0_4 
SUCCESS: c_0_22 has SAT parents c_0_15 c_0_11 c_0_16 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_22 is a thm of c_0_15 c_0_11 c_0_16 
SUCCESS: c_0_23 has SAT parents c_0_17 c_0_18 c_0_19 
 RESULT: c_0_23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_23 is a thm of c_0_17 c_0_18 c_0_19 
SUCCESS: c_0_24 has SAT parents c_0_20 
 RESULT: c_0_24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_24 is a thm of c_0_20 
SUCCESS: c_0_25 has SAT parents c_0_21 c_0_11 c_0_18 
 RESULT: c_0_25_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_25 is a thm of c_0_21 c_0_11 c_0_18 
SUCCESS: c_0_26 has SAT parents c_0_17 c_0_22 
 RESULT: c_0_26_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_26 is a thm of c_0_17 c_0_22 
SUCCESS: c_0_27 has SAT parents c_0_23 c_0_24 
 RESULT: c_0_27_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_27 is a thm of c_0_23 c_0_24 
SUCCESS: c_0_28 has SAT parents c_0_25 c_0_11 
 RESULT: c_0_28_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_28 is a thm of c_0_25 c_0_11 
SUCCESS: c_0_29 has SAT parents c_0_26 c_0_27 
 RESULT: c_0_29_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_29 is a thm of c_0_26 c_0_27 
SUCCESS: c_0_30 has SAT parents c_0_4 
 RESULT: c_0_30_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_30 is a thm of c_0_4 
SUCCESS: c_0_31 has SAT parents c_0_28 c_0_29 
 RESULT: c_0_31_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_31 is a thm of c_0_28 c_0_29 
SUCCESS: c_0_32 has SAT parents c_0_30 c_0_29 c_0_31 
 RESULT: c_0_32_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_32 is a thm of c_0_30 c_0_29 c_0_31 
SUCCESS: c_0_33 has SAT parents c_0_4 
 RESULT: c_0_33_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_33 is a thm of c_0_4 
SUCCESS: c_0_34 has SAT parents c_0_32 c_0_11 c_0_27 
 RESULT: c_0_34_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_34 is a thm of c_0_32 c_0_11 c_0_27 
SUCCESS: c_0_35 has SAT parents c_0_33 c_0_11 c_0_13 
 RESULT: c_0_35_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_35 is a thm of c_0_33 c_0_11 c_0_13 
SUCCESS: c_0_36 has SAT parents c_0_4 
 RESULT: c_0_36_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_36 is a thm of c_0_4 
SUCCESS: c_0_37 has SAT parents c_0_4 
 RESULT: c_0_37_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_37 is a thm of c_0_4 
SUCCESS: c_0_38 has SAT parents c_0_17 c_0_34 
 RESULT: c_0_38_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_38 is a thm of c_0_17 c_0_34 
SUCCESS: c_0_39 has SAT parents c_0_4 
 RESULT: c_0_39_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_39 is a thm of c_0_4 
SUCCESS: c_0_40 has SAT parents c_0_35 c_0_11 
 RESULT: c_0_40_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_40 is a thm of c_0_35 c_0_11 
SUCCESS: c_0_41 has SAT parents c_0_36 c_0_11 c_0_37 
 RESULT: c_0_41_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_41 is a thm of c_0_36 c_0_11 c_0_37 
SUCCESS: c_0_42 has SAT parents c_0_38 c_0_27 
 RESULT: c_0_42_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_42 is a thm of c_0_38 c_0_27 
SUCCESS: c_0_43 has SAT parents c_0_39 c_0_11 c_0_14 c_0_40 
 RESULT: c_0_43_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_43 is a thm of c_0_39 c_0_11 c_0_14 c_0_40 
SUCCESS: c_0_44 has SAT parents c_0_41 c_0_42 
 RESULT: c_0_44_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_44 is a thm of c_0_41 c_0_42 
SUCCESS: c_0_45 has SAT parents c_0_43 c_0_42 
 RESULT: c_0_45_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_45 is a thm of c_0_43 c_0_42 
SUCCESS: c_0_46 has SAT parents c_0_17 c_0_44 c_0_45 
 RESULT: c_0_46_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_46 is a thm of c_0_17 c_0_44 c_0_45 
WARNING: c_0_47 has UNS parents c_0_24 c_0_46 
 RESULT: c_0_47_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_47 is a thm of c_0_24 c_0_46 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 29.02
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN059+1/E---3.2.0.THM-CRf.s
fof(pel29,conjecture,
    ( ( ! [X1] :
          ( big_f(X1)
         => big_h(X1) )
      & ! [X3] :
          ( big_g(X3)
         => big_j(X3) ) )
  <=> ! [X4,X2] :
        ( ( big_f(X4)
          & big_g(X2) )
       => ( big_h(X4)
          & big_j(X2) ) ) ),
    file('/export/starexec/sandbox/tmp/tmp.lqyUH3z2iK/E---3.1_12163.p',pel29),
    [verified_file(pel29_thm),verified(leaf),verified(lambdapi)] ).
fof(pel29_2,axiom,
    ? [X2] : big_g(X2),
    file('/export/starexec/sandbox/tmp/tmp.lqyUH3z2iK/E---3.1_12163.p',pel29_2),
    [verified_file(pel29_2_thm),verified(leaf),nnpp(pel29),verified(lambdapi)] ).
fof(pel29_1,axiom,
    ? [X1] : big_f(X1),
    file('/export/starexec/sandbox/tmp/tmp.lqyUH3z2iK/E---3.1_12163.p',pel29_1),
    [verified_file(pel29_1_thm),verified(leaf),nnpp(pel29),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ~ ( ( ! [X1] :
            ( big_f(X1)
           => big_h(X1) )
        & ! [X3] :
            ( big_g(X3)
           => big_j(X3) ) )
    <=> ! [X4,X2] :
          ( ( big_f(X4)
            & big_g(X2) )
         => ( big_h(X4)
            & big_j(X2) ) ) ),
    inference(assume_negation,[status(cth)],[pel29]),
    [nnpp(pel29),verified_file(c_0_3_thm),verified(cth),nnpp(pel29),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X11,X12,X13,X14] :
      ( ( big_f(esk5_0)
        | big_g(esk4_0)
        | big_f(esk3_0) )
      & ( big_g(esk6_0)
        | big_g(esk4_0)
        | big_f(esk3_0) )
      & ( ~ big_h(esk5_0)
        | ~ big_j(esk6_0)
        | big_g(esk4_0)
        | big_f(esk3_0) )
      & ( big_f(esk5_0)
        | ~ big_j(esk4_0)
        | big_f(esk3_0) )
      & ( big_g(esk6_0)
        | ~ big_j(esk4_0)
        | big_f(esk3_0) )
      & ( ~ big_h(esk5_0)
        | ~ big_j(esk6_0)
        | ~ big_j(esk4_0)
        | big_f(esk3_0) )
      & ( big_f(esk5_0)
        | big_g(esk4_0)
        | ~ big_h(esk3_0) )
      & ( big_g(esk6_0)
        | big_g(esk4_0)
        | ~ big_h(esk3_0) )
      & ( ~ big_h(esk5_0)
        | ~ big_j(esk6_0)
        | big_g(esk4_0)
        | ~ big_h(esk3_0) )
      & ( big_f(esk5_0)
        | ~ big_j(esk4_0)
        | ~ big_h(esk3_0) )
      & ( big_g(esk6_0)
        | ~ big_j(esk4_0)
        | ~ big_h(esk3_0) )
      & ( ~ big_h(esk5_0)
        | ~ big_j(esk6_0)
        | ~ big_j(esk4_0)
        | ~ big_h(esk3_0) )
      & ( big_h(X13)
        | ~ big_f(X13)
        | ~ big_g(X14)
        | ~ big_f(X11)
        | big_h(X11) )
      & ( big_j(X14)
        | ~ big_f(X13)
        | ~ big_g(X14)
        | ~ big_f(X11)
        | big_h(X11) )
      & ( big_h(X13)
        | ~ big_f(X13)
        | ~ big_g(X14)
        | ~ big_g(X12)
        | big_j(X12) )
      & ( big_j(X14)
        | ~ big_f(X13)
        | ~ big_g(X14)
        | ~ big_g(X12)
        | big_j(X12) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])]),
    [nnpp(pel29),verified_file(c_0_4_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X3,X1] :
      ( big_h(X1)
      | big_h(X3)
      | ~ big_f(X1)
      | ~ big_g(X2)
      | ~ big_f(X3) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,plain,
    big_g(esk2_0),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[pel29_2])]),
    [nnpp(pel29),verified_file(c_0_6_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X2,X1] :
      ( big_h(X1)
      | ~ big_g(X2)
      | ~ big_f(X1) ),
    inference(ef,[status(thm)],[c_0_5]),
    [nnpp(pel29),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,plain,
    big_g(esk2_0),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel29),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X2,X3,X1] :
      ( big_j(X1)
      | big_j(X3)
      | ~ big_f(X2)
      | ~ big_g(X1)
      | ~ big_g(X3) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ( big_f(esk5_0)
    | big_g(esk4_0)
    | ~ big_h(esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ! [X1] :
      ( big_h(X1)
      | ~ big_f(X1) ),
    inference(spm,[status(thm)],[c_0_7,c_0_8]),
    [nnpp(pel29),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ( big_f(esk5_0)
    | big_g(esk4_0)
    | big_f(esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ( big_f(esk5_0)
    | ~ big_j(esk4_0)
    | ~ big_h(esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ( big_f(esk5_0)
    | big_f(esk3_0)
    | ~ big_j(esk4_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ( big_g(esk6_0)
    | big_g(esk4_0)
    | ~ big_h(esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ( big_g(esk6_0)
    | big_g(esk4_0)
    | big_f(esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ! [X2,X1] :
      ( big_j(X1)
      | ~ big_g(X1)
      | ~ big_f(X2) ),
    inference(ef,[status(thm)],[c_0_9]),
    [nnpp(pel29),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    ( big_g(esk4_0)
    | big_f(esk5_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10,c_0_11]),c_0_12]),
    [nnpp(pel29),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,negated_conjecture,
    ( big_f(esk5_0)
    | ~ big_j(esk4_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13,c_0_11]),c_0_14]),
    [nnpp(pel29),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,plain,
    big_f(esk1_0),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[pel29_1])]),
    [nnpp(pel29),verified_file(c_0_20_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_21,negated_conjecture,
    ( big_g(esk4_0)
    | ~ big_h(esk5_0)
    | ~ big_j(esk6_0)
    | ~ big_h(esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,negated_conjecture,
    ( big_g(esk4_0)
    | big_g(esk6_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15,c_0_11]),c_0_16]),
    [nnpp(pel29),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_23,negated_conjecture,
    ! [X1] :
      ( big_f(esk5_0)
      | ~ big_f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17,c_0_18]),c_0_19]),
    [nnpp(pel29),verified_file(c_0_23_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_24,plain,
    big_f(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_20]),
    [nnpp(pel29),verified_file(c_0_24_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_25,negated_conjecture,
    ( big_g(esk4_0)
    | ~ big_j(esk6_0)
    | ~ big_h(esk3_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21,c_0_11]),c_0_18]),
    [nnpp(pel29),verified_file(c_0_25_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_26,negated_conjecture,
    ! [X1] :
      ( big_j(esk6_0)
      | big_g(esk4_0)
      | ~ big_f(X1) ),
    inference(spm,[status(thm)],[c_0_17,c_0_22]),
    [nnpp(pel29),verified_file(c_0_26_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_27,negated_conjecture,
    big_f(esk5_0),
    inference(spm,[status(thm)],[c_0_23,c_0_24]),
    [nnpp(pel29),verified_file(c_0_27_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_28,negated_conjecture,
    ( big_g(esk4_0)
    | ~ big_j(esk6_0)
    | ~ big_f(esk3_0) ),
    inference(spm,[status(thm)],[c_0_25,c_0_11]),
    [nnpp(pel29),verified_file(c_0_28_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_29,negated_conjecture,
    ( big_j(esk6_0)
    | big_g(esk4_0) ),
    inference(spm,[status(thm)],[c_0_26,c_0_27]),
    [nnpp(pel29),verified_file(c_0_29_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_30,negated_conjecture,
    ( big_g(esk4_0)
    | big_f(esk3_0)
    | ~ big_h(esk5_0)
    | ~ big_j(esk6_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_30_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_31,negated_conjecture,
    ( big_g(esk4_0)
    | ~ big_f(esk3_0) ),
    inference(spm,[status(thm)],[c_0_28,c_0_29]),
    [nnpp(pel29),verified_file(c_0_31_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_32,negated_conjecture,
    ( big_g(esk4_0)
    | ~ big_h(esk5_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30,c_0_29]),c_0_31]),
    [nnpp(pel29),verified_file(c_0_32_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_33,negated_conjecture,
    ( ~ big_h(esk5_0)
    | ~ big_j(esk6_0)
    | ~ big_j(esk4_0)
    | ~ big_h(esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_33_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_34,negated_conjecture,
    big_g(esk4_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32,c_0_11]),c_0_27])]),
    [nnpp(pel29),verified_file(c_0_34_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_35,negated_conjecture,
    ( ~ big_j(esk4_0)
    | ~ big_j(esk6_0)
    | ~ big_h(esk3_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33,c_0_11]),c_0_13]),
    [nnpp(pel29),verified_file(c_0_35_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_36,negated_conjecture,
    ( big_g(esk6_0)
    | ~ big_j(esk4_0)
    | ~ big_h(esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_36_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_37,negated_conjecture,
    ( big_g(esk6_0)
    | big_f(esk3_0)
    | ~ big_j(esk4_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_37_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_38,negated_conjecture,
    ! [X1] :
      ( big_j(esk4_0)
      | ~ big_f(X1) ),
    inference(spm,[status(thm)],[c_0_17,c_0_34]),
    [nnpp(pel29),verified_file(c_0_38_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_39,negated_conjecture,
    ( big_f(esk3_0)
    | ~ big_h(esk5_0)
    | ~ big_j(esk6_0)
    | ~ big_j(esk4_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel29),verified_file(c_0_39_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_40,negated_conjecture,
    ( ~ big_j(esk4_0)
    | ~ big_j(esk6_0)
    | ~ big_f(esk3_0) ),
    inference(spm,[status(thm)],[c_0_35,c_0_11]),
    [nnpp(pel29),verified_file(c_0_40_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_41,negated_conjecture,
    ( big_g(esk6_0)
    | ~ big_j(esk4_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36,c_0_11]),c_0_37]),
    [nnpp(pel29),verified_file(c_0_41_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_42,negated_conjecture,
    big_j(esk4_0),
    inference(spm,[status(thm)],[c_0_38,c_0_27]),
    [nnpp(pel29),verified_file(c_0_42_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_43,negated_conjecture,
    ( ~ big_j(esk4_0)
    | ~ big_j(esk6_0) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_39,c_0_11]),c_0_14]),c_0_40]),
    [nnpp(pel29),verified_file(c_0_43_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_44,negated_conjecture,
    big_g(esk6_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_41,c_0_42])]),
    [nnpp(pel29),verified_file(c_0_44_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_45,negated_conjecture,
    ~ big_j(esk6_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43,c_0_42])]),
    [nnpp(pel29),verified_file(c_0_45_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_46,negated_conjecture,
    ! [X1] : ~ big_f(X1),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17,c_0_44]),c_0_45]),
    [nnpp(pel29),verified_file(c_0_46_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_47,plain,
    $false,
    inference(sr,[status(thm)],[c_0_24,c_0_46]),
    [proof,nnpp(pel29),verified_file(c_0_47_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN059+1/E---3.2.0.THM-CRf.s
