SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel9 has SAT parents the problem 
 RESULT: pel9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel9 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents pel9 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_1 is a cth of pel9 
WARNING: c_0_2 has UNS parents c_0_1 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_2 is a thm of c_0_1 
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_2 
SUCCESS: c_0_8 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_8 is a thm of c_0_5 c_0_6 
WARNING: c_0_9 has UNS parents c_0_7 c_0_6 c_0_8 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_7 c_0_6 c_0_8 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 6.17
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN391+1/E---3.2.0.THM-CRf.s
fof(pel9,conjecture,
    ( ( ( p1
        | p2 )
      & ( ~ p1
        | p2 )
      & ( p1
        | ~ p2 ) )
   => ~ ( ~ p1
        | ~ p2 ) ),
    file('/export/starexec/sandbox/tmp/tmp.mWWahIHjUk/E---3.1_5565.p',pel9),
    [verified_file(pel9_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( ( p1
          | p2 )
        & ( ~ p1
          | p2 )
        & ( p1
          | ~ p2 ) )
     => ~ ( ~ p1
          | ~ p2 ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[pel9])]),
    [nnpp(pel9),verified_file(c_0_1_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ( ( p1
      | p2 )
    & ( ~ p1
      | p2 )
    & ( p1
      | ~ p2 )
    & ( ~ p1
      | ~ p2 ) ),
    inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])]),
    [nnpp(pel9),verified_file(c_0_2_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ( p2
    | ~ p1 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel9),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( p1
    | p2 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel9),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ( ~ p1
    | ~ p2 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel9),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    p2,
    inference(csr,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(pel9),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ( p1
    | ~ p2 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel9),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ~ p1,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5,c_0_6])]),
    [nnpp(pel9),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7,c_0_6])]),c_0_8]),
    [proof,nnpp(pel9),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN391+1/E---3.2.0.THM-CRf.s
