SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: church_46_18_3 has SAT parents the problem 
 RESULT: church_46_18_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: church_46_18_3 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents church_46_18_3 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_1 is a cth of church_46_18_3 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_3 c_0_4 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_3 c_0_4 
WARNING: c_0_6 has UNS parents c_0_5 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.07
SUCCESS: c_0_6 is a thm of c_0_5 c_0_4 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 193.46
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN351+1/E---3.2.0.THM-CRf.s
fof(church_46_18_3,conjecture,
    ! [X1,X2] :
    ? [X3,X4] :
    ! [X5] :
      ( big_f(X1,X4,X1,X5)
     => ( ( big_f(X1,X3,X1,X4)
        <=> big_f(X3,X2,X3,X4) )
       => ( big_f(X1,X3,X1,X4)
         => ( ( big_f(X1,X4,X3,X4)
             => big_f(X1,X5,X3,X5) )
            & ( big_f(X1,X5,X3,X5)
             => ( big_f(X1,X3,X1,X4)
              <=> big_f(X1,X4,X3,X4) ) ) ) ) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.RBlJTETlc2/E---3.1_22513.p',church_46_18_3),
    [verified_file(church_46_18_3_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ! [X1,X2] :
      ? [X3,X4] :
      ! [X5] :
        ( big_f(X1,X4,X1,X5)
       => ( ( big_f(X1,X3,X1,X4)
          <=> big_f(X3,X2,X3,X4) )
         => ( big_f(X1,X3,X1,X4)
           => ( ( big_f(X1,X4,X3,X4)
               => big_f(X1,X5,X3,X5) )
              & ( big_f(X1,X5,X3,X5)
               => ( big_f(X1,X3,X1,X4)
                <=> big_f(X1,X4,X3,X4) ) ) ) ) ) ),
    inference(assume_negation,[status(cth)],[church_46_18_3]),
    [nnpp(church_46_18_3),verified_file(c_0_1_thm),verified(cth),nnpp(church_46_18_3),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X8,X9] :
      ( big_f(esk1_0,X9,esk1_0,esk3_2(X8,X9))
      & ( ~ big_f(esk1_0,X8,esk1_0,X9)
        | big_f(X8,esk2_0,X8,X9) )
      & ( ~ big_f(X8,esk2_0,X8,X9)
        | big_f(esk1_0,X8,esk1_0,X9) )
      & big_f(esk1_0,X8,esk1_0,X9)
      & ( big_f(esk1_0,esk3_2(X8,X9),X8,esk3_2(X8,X9))
        | big_f(esk1_0,X9,X8,X9) )
      & ( ~ big_f(esk1_0,X8,esk1_0,X9)
        | ~ big_f(esk1_0,X9,X8,X9)
        | big_f(esk1_0,X9,X8,X9) )
      & ( big_f(esk1_0,X8,esk1_0,X9)
        | big_f(esk1_0,X9,X8,X9)
        | big_f(esk1_0,X9,X8,X9) )
      & ( big_f(esk1_0,esk3_2(X8,X9),X8,esk3_2(X8,X9))
        | ~ big_f(esk1_0,esk3_2(X8,X9),X8,esk3_2(X8,X9)) )
      & ( ~ big_f(esk1_0,X8,esk1_0,X9)
        | ~ big_f(esk1_0,X9,X8,X9)
        | ~ big_f(esk1_0,esk3_2(X8,X9),X8,esk3_2(X8,X9)) )
      & ( big_f(esk1_0,X8,esk1_0,X9)
        | big_f(esk1_0,X9,X8,X9)
        | ~ big_f(esk1_0,esk3_2(X8,X9),X8,esk3_2(X8,X9)) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])]),
    [nnpp(church_46_18_3),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( ~ big_f(esk1_0,X1,esk1_0,X2)
      | ~ big_f(esk1_0,X2,X1,X2)
      | ~ big_f(esk1_0,esk3_2(X1,X2),X1,esk3_2(X1,X2)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_18_3),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X2,X1] : big_f(esk1_0,X1,esk1_0,X2),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_18_3),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X1] :
      ( ~ big_f(esk1_0,esk3_2(X1,X2),X1,esk3_2(X1,X2))
      | ~ big_f(esk1_0,X2,X1,X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_3,c_0_4])]),
    [nnpp(church_46_18_3),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5,c_0_4]),c_0_4])]),
    [proof,nnpp(church_46_18_3),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN351+1/E---3.2.0.THM-CRf.s
