SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: kalish256 has SAT parents the problem 
 RESULT: kalish256_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: kalish256 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents kalish256 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a cth of kalish256 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
WARNING: c_0_7 has UNS parents c_0_5 c_0_6 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.07
SUCCESS: c_0_7 is a thm of c_0_5 c_0_6 
WARNING: c_0_8 has UNS parents c_0_5 c_0_7 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of c_0_5 c_0_7 
WARNING: c_0_9 has UNS parents c_0_8 c_0_7 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_9 is a thm of c_0_8 c_0_7 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 256.87
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN413+1/E---3.2.0.THM-CRf.s
fof(kalish256,conjecture,
    ( ! [X1] :
      ? [X2] :
      ! [X3] :
        ( f(X3,X2)
      <=> ( f(X3,X1)
          & ~ f(X3,X3) ) )
   => ~ ? [X4] :
        ! [X5] : f(X5,X4) ),
    file('/export/starexec/sandbox/tmp/tmp.uzlP3ktE0z/E---3.1_17827.p',kalish256),
    [verified_file(kalish256_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ! [X1] :
        ? [X2] :
        ! [X3] :
          ( f(X3,X2)
        <=> ( f(X3,X1)
            & ~ f(X3,X3) ) )
     => ~ ? [X4] :
          ! [X5] : f(X5,X4) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[kalish256])]),
    [nnpp(kalish256),verified_file(c_0_1_thm),verified(cth),nnpp(kalish256),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X6,X8,X9,X11] :
      ( ( f(X8,X6)
        | ~ f(X8,esk1_1(X6)) )
      & ( ~ f(X8,X8)
        | ~ f(X8,esk1_1(X6)) )
      & ( ~ f(X9,X6)
        | f(X9,X9)
        | f(X9,esk1_1(X6)) )
      & f(X11,esk2_0) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])])]),
    [nnpp(kalish256),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( f(X1,X1)
      | f(X1,esk1_1(X2))
      | ~ f(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish256),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X1] : f(X1,esk2_0),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish256),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X1] :
      ( ~ f(X1,X1)
      | ~ f(X1,esk1_1(X2)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish256),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X1] :
      ( f(X1,esk1_1(esk2_0))
      | f(X1,X1) ),
    inference(spm,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(kalish256),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1] : f(esk1_1(X1),esk1_1(esk2_0)),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_5,c_0_6]),c_0_6]),
    [nnpp(kalish256),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X1] : ~ f(esk1_1(X1),esk1_1(X1)),
    inference(spm,[status(thm)],[c_0_5,c_0_7]),
    [nnpp(kalish256),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    $false,
    inference(spm,[status(thm)],[c_0_8,c_0_7]),
    [proof,nnpp(kalish256),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN413+1/E---3.2.0.THM-CRf.s
