SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: kalish234 has SAT parents the problem 
 RESULT: kalish234_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: kalish234 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents kalish234 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_1 is a cth of kalish234 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_2 
WARNING: c_0_8 has UNS parents c_0_5 c_0_6 c_0_7 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_5 c_0_6 c_0_7 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 6.48
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN403+1/E---3.2.0.THM-CRf.s
fof(kalish234,conjecture,
    ! [X1] :
      ( ( ( f(X1)
         => g(X1) )
        & ( g(X1)
         => h(X1) ) )
     => ( f(X1)
       => h(X1) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.yyMmEhffaL/E---3.1_1949.p',kalish234),
    [verified_file(kalish234_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ! [X1] :
        ( ( ( f(X1)
           => g(X1) )
          & ( g(X1)
           => h(X1) ) )
       => ( f(X1)
         => h(X1) ) ),
    inference(assume_negation,[status(cth)],[kalish234]),
    [nnpp(kalish234),verified_file(c_0_1_thm),verified(cth),nnpp(kalish234),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ( ( ~ f(esk1_0)
      | g(esk1_0) )
    & ( ~ g(esk1_0)
      | h(esk1_0) )
    & f(esk1_0)
    & ~ h(esk1_0) ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])]),
    [nnpp(kalish234),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ( g(esk1_0)
    | ~ f(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish234),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    f(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish234),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ( h(esk1_0)
    | ~ g(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish234),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    g(esk1_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_3,c_0_4])]),
    [nnpp(kalish234),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ~ h(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish234),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5,c_0_6])]),c_0_7]),
    [proof,nnpp(kalish234),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN403+1/E---3.2.0.THM-CRf.s
