SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: church_46_4_2 has SAT parents the problem 
 RESULT: church_46_4_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: church_46_4_2 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents church_46_4_2 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_1 is a cth of church_46_4_2 
WARNING: c_0_2 has UNS parents c_0_1 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_2 is a thm of c_0_1 
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_7 is a thm of c_0_2 
SUCCESS: c_0_8 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_8 is a thm of c_0_5 c_0_6 
WARNING: c_0_9 has UNS parents c_0_7 c_0_8 c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_9 is a thm of c_0_7 c_0_8 c_0_6 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 6.31
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN323+1/E---3.2.0.THM-CRf.s
fof(church_46_4_2,conjecture,
    ? [X1] :
      ( ! [X2] :
          ( ( big_f(X2,X1)
           => big_f(X1,X2) )
         => big_g(X2,X1) )
     => ~ ! [X2] :
            ( ( big_f(X2,X1)
             => big_f(X1,X2) )
           => ~ big_g(X2,X1) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.VralGIlztZ/E---3.1_17550.p',church_46_4_2),
    [verified_file(church_46_4_2_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ? [X1] :
        ( ! [X2] :
            ( ( big_f(X2,X1)
             => big_f(X1,X2) )
           => big_g(X2,X1) )
       => ~ ! [X2] :
              ( ( big_f(X2,X1)
               => big_f(X1,X2) )
             => ~ big_g(X2,X1) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[church_46_4_2])]),
    [nnpp(church_46_4_2),verified_file(c_0_1_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X3,X4,X5] :
      ( ( big_f(X4,X3)
        | big_g(X4,X3) )
      & ( ~ big_f(X3,X4)
        | big_g(X4,X3) )
      & ( big_f(X5,X3)
        | ~ big_g(X5,X3) )
      & ( ~ big_f(X3,X5)
        | ~ big_g(X5,X3) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])]),
    [nnpp(church_46_4_2),verified_file(c_0_2_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X1,X2] :
      ( big_g(X2,X1)
      | ~ big_f(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_4_2),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X2,X1] :
      ( ~ big_f(X1,X2)
      | ~ big_g(X2,X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_4_2),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X1] :
      ( big_f(X1,X2)
      | big_g(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_4_2),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X2,X1] : ~ big_f(X1,X2),
    inference(csr,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(church_46_4_2),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X2,X1] :
      ( big_f(X1,X2)
      | ~ big_g(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_4_2),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X2,X1] : big_g(X1,X2),
    inference(sr,[status(thm)],[c_0_5,c_0_6]),
    [nnpp(church_46_4_2),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7,c_0_8])]),c_0_6]),
    [proof,nnpp(church_46_4_2),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN323+1/E---3.2.0.THM-CRf.s
