SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: hyp2 has SAT parents the problem 
 RESULT: hyp2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: hyp2 is a thm of the problem 
SUCCESS: hyp1 has SAT parents the problem 
 RESULT: hyp1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: hyp1 is a thm of the problem 
SUCCESS: ck has SAT parents the problem 
 RESULT: ck_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: ck is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_3 has SAT parents hyp2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_3 is a thm of hyp2 
SUCCESS: c_0_4 has SAT parents hyp1 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of hyp1 
SUCCESS: c_0_5 has SAT parents c_0_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_5 is a thm of c_0_3 
SUCCESS: c_0_6 has SAT parents c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_6 is a thm of c_0_4 
SUCCESS: c_0_7 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_5 c_0_6 
SUCCESS: c_0_8 has SAT parents c_0_5 c_0_7 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_8 is a thm of c_0_5 c_0_7 
SUCCESS: c_0_9 has SAT parents c_0_8 c_0_7 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_8 c_0_7 
SUCCESS: c_0_10 has SAT parents c_0_9 c_0_6 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_10 is a thm of c_0_9 c_0_6 
SUCCESS: c_0_11 has SAT parents ck 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_11 is a cth of ck 
SUCCESS: c_0_12 has SAT parents c_0_5 c_0_10 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_12 is a thm of c_0_5 c_0_10 
SUCCESS: c_0_13 has SAT parents c_0_11 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_13 is a thm of c_0_11 
SUCCESS: c_0_14 has SAT parents c_0_12 c_0_10 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_12 c_0_10 
SUCCESS: c_0_15 has SAT parents c_0_13 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_15 is a thm of c_0_13 
SUCCESS: c_0_16 has SAT parents c_0_14 c_0_6 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_16 is a thm of c_0_14 c_0_6 
SUCCESS: c_0_17 has SAT parents c_0_15 c_0_6 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_17 is a thm of c_0_15 c_0_6 
SUCCESS: c_0_18 has SAT parents c_0_5 c_0_16 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_18 is a thm of c_0_5 c_0_16 
SUCCESS: c_0_19 has SAT parents c_0_17 c_0_7 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_19 is a thm of c_0_17 c_0_7 
SUCCESS: c_0_20 has SAT parents c_0_18 c_0_16 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_20 is a thm of c_0_18 c_0_16 
SUCCESS: c_0_21 has SAT parents c_0_19 c_0_10 c_0_6 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_21 is a thm of c_0_19 c_0_10 c_0_6 
SUCCESS: c_0_22 has SAT parents c_0_20 c_0_6 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_22 is a thm of c_0_20 c_0_6 
WARNING: c_0_23 has UNS parents c_0_21 c_0_22 c_0_6 
 RESULT: c_0_23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_23 is a thm of c_0_21 c_0_22 c_0_6 
WARNING: c_0_24 has UNS parents c_0_23 c_0_6 
 RESULT: c_0_24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_24 is a thm of c_0_23 c_0_6 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 71.75
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN986+1.003/E---3.2.0.THM-CRf.s
fof(hyp2,axiom,
    ! [X1,X2,X3,X4] :
      ( r(X1,X2,X3)
     => ( r(X3,X2,X4)
       => r(X1,succ(X2),X4) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.jKFaraOOxk/E---3.1_8959.p',hyp2),
    [verified_file(hyp2_thm),verified(leaf),verified(lambdapi)] ).
fof(hyp1,axiom,
    ! [X1] : r(X1,zero,succ(X1)),
    file('/export/starexec/sandbox2/tmp/tmp.jKFaraOOxk/E---3.1_8959.p',hyp1),
    [verified_file(hyp1_thm),verified(leaf),verified(lambdapi)] ).
fof(ck,conjecture,
    ? [X5,X6,X4,X7] :
      ( r(zero,zero,X5)
      & r(zero,X5,X6)
      & r(zero,X6,X4)
      & r(zero,X4,X7) ),
    file('/export/starexec/sandbox2/tmp/tmp.jKFaraOOxk/E---3.1_8959.p',ck),
    [verified_file(ck_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_3,plain,
    ! [X13,X14,X15,X16] :
      ( ~ r(X13,X14,X15)
      | ~ r(X15,X14,X16)
      | r(X13,succ(X14),X16) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[hyp2])])]),
    [nnpp(ck),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,plain,
    ! [X12] : r(X12,zero,succ(X12)),
    inference(variable_rename,[status(thm)],[hyp1]),
    [nnpp(ck),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,plain,
    ! [X3,X4,X2,X1] :
      ( r(X1,succ(X2),X4)
      | ~ r(X1,X2,X3)
      | ~ r(X3,X2,X4) ),
    inference(split_conjunct,[status(thm)],[c_0_3]),
    [nnpp(ck),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,plain,
    ! [X1] : r(X1,zero,succ(X1)),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(ck),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,plain,
    ! [X2,X1] :
      ( r(X1,succ(zero),succ(X2))
      | ~ r(X1,zero,X2) ),
    inference(spm,[status(thm)],[c_0_5,c_0_6]),
    [nnpp(ck),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,plain,
    ! [X3,X2,X1] :
      ( r(X1,succ(succ(zero)),succ(X2))
      | ~ r(X1,succ(zero),X3)
      | ~ r(X3,zero,X2) ),
    inference(spm,[status(thm)],[c_0_5,c_0_7]),
    [nnpp(ck),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,plain,
    ! [X3,X2,X1] :
      ( r(X1,succ(succ(zero)),succ(X2))
      | ~ r(succ(X3),zero,X2)
      | ~ r(X1,zero,X3) ),
    inference(spm,[status(thm)],[c_0_8,c_0_7]),
    [nnpp(ck),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,plain,
    ! [X2,X1] :
      ( r(X1,succ(succ(zero)),succ(succ(succ(X2))))
      | ~ r(X1,zero,X2) ),
    inference(spm,[status(thm)],[c_0_9,c_0_6]),
    [nnpp(ck),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ~ ? [X5,X6,X4,X7] :
        ( r(zero,zero,X5)
        & r(zero,X5,X6)
        & r(zero,X6,X4)
        & r(zero,X4,X7) ),
    inference(assume_negation,[status(cth)],[ck]),
    [nnpp(ck),verified_file(c_0_11_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_12,plain,
    ! [X3,X2,X1] :
      ( r(X1,succ(succ(succ(zero))),succ(succ(succ(X2))))
      | ~ r(X1,succ(succ(zero)),X3)
      | ~ r(X3,zero,X2) ),
    inference(spm,[status(thm)],[c_0_5,c_0_10]),
    [nnpp(ck),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ! [X8,X9,X10,X11] :
      ( ~ r(zero,zero,X8)
      | ~ r(zero,X8,X9)
      | ~ r(zero,X9,X10)
      | ~ r(zero,X10,X11) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])]),
    [nnpp(ck),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,plain,
    ! [X3,X2,X1] :
      ( r(X1,succ(succ(succ(zero))),succ(succ(succ(X2))))
      | ~ r(succ(succ(succ(X3))),zero,X2)
      | ~ r(X1,zero,X3) ),
    inference(spm,[status(thm)],[c_0_12,c_0_10]),
    [nnpp(ck),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ! [X4,X3,X2,X1] :
      ( ~ r(zero,zero,X1)
      | ~ r(zero,X1,X2)
      | ~ r(zero,X2,X3)
      | ~ r(zero,X3,X4) ),
    inference(split_conjunct,[status(thm)],[c_0_13]),
    [nnpp(ck),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,plain,
    ! [X2,X1] :
      ( r(X1,succ(succ(succ(zero))),succ(succ(succ(succ(succ(succ(succ(X2))))))))
      | ~ r(X1,zero,X2) ),
    inference(spm,[status(thm)],[c_0_14,c_0_6]),
    [nnpp(ck),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ! [X3,X2,X1] :
      ( ~ r(zero,succ(zero),X1)
      | ~ r(zero,X2,X3)
      | ~ r(zero,X1,X2) ),
    inference(spm,[status(thm)],[c_0_15,c_0_6]),
    [nnpp(ck),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,plain,
    ! [X3,X2,X1] :
      ( r(X1,succ(succ(succ(succ(zero)))),succ(succ(succ(succ(succ(succ(succ(X2))))))))
      | ~ r(X1,succ(succ(succ(zero))),X3)
      | ~ r(X3,zero,X2) ),
    inference(spm,[status(thm)],[c_0_5,c_0_16]),
    [nnpp(ck),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,negated_conjecture,
    ! [X3,X2,X1] :
      ( ~ r(zero,succ(X1),X2)
      | ~ r(zero,zero,X1)
      | ~ r(zero,X2,X3) ),
    inference(spm,[status(thm)],[c_0_17,c_0_7]),
    [nnpp(ck),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,plain,
    ! [X3,X2,X1] :
      ( r(X1,succ(succ(succ(succ(zero)))),succ(succ(succ(succ(succ(succ(succ(X2))))))))
      | ~ r(succ(succ(succ(succ(succ(succ(succ(X3))))))),zero,X2)
      | ~ r(X1,zero,X3) ),
    inference(spm,[status(thm)],[c_0_18,c_0_16]),
    [nnpp(ck),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,negated_conjecture,
    ! [X2,X1] :
      ( ~ r(zero,succ(succ(succ(X1))),X2)
      | ~ r(zero,zero,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19,c_0_10]),c_0_6])]),
    [nnpp(ck),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,plain,
    ! [X2,X1] :
      ( r(X1,succ(succ(succ(succ(zero)))),succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(succ(X2))))))))))))))))
      | ~ r(X1,zero,X2) ),
    inference(spm,[status(thm)],[c_0_20,c_0_6]),
    [nnpp(ck),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_23,negated_conjecture,
    ! [X1] : ~ r(zero,zero,X1),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21,c_0_22]),c_0_6])]),
    [nnpp(ck),verified_file(c_0_23_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_24,negated_conjecture,
    $false,
    inference(spm,[status(thm)],[c_0_23,c_0_6]),
    [proof,nnpp(ck),verified_file(c_0_24_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN986+1.003/E---3.2.0.THM-CRf.s
