SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
 RESULT: problem_axioms_sat_model - Paradox---4.0 says Satisfiable - CPU = 0.00
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel44 has SAT parents the problem 
 RESULT: pel44_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel44 is a thm of the problem 
SUCCESS: pel44_1 has SAT parents the problem 
 RESULT: pel44_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel44_1 is a thm of the problem 
SUCCESS: pel44_2 has SAT parents the problem 
 RESULT: pel44_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel44_2 is a thm of the problem 
SUCCESS: Leaves are verified
 RESULT: leaf_axioms_sat_model - Paradox---4.0 says Satisfiable - CPU = 0.00
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_3 has SAT parents pel44 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_3 is a cth of pel44 
SUCCESS: c_0_4 has SAT parents pel44_1 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_4 is a thm of pel44_1 
SUCCESS: c_0_5 has SAT parents c_0_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_3 
SUCCESS: c_0_6 has SAT parents pel44_2 (Theorem esa)
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.01
FAILURE: c_0_6 fails to be a thm of pel44_2 (Theorem esa)
SUCCESS: pel44_2 has SAT parents c_0_6 (Inverted esa)
 RESULT: c_0_6_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel44_2 is a thm of c_0_6 (Inverted esa)
FAILURE: c_0_6 fails to be a esa of pel44_2
 ASSUME: c_0_6 is a esa of pel44_2
SUCCESS: c_0_7 has SAT parents c_0_4 (Theorem esa)
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.04
FAILURE: c_0_7 fails to be a thm of c_0_4 (Theorem esa)
SUCCESS: c_0_4 has SAT parents c_0_7 (Inverted esa)
 RESULT: c_0_7_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_7 (Inverted esa)
FAILURE: c_0_7 fails to be a esa of c_0_4
 ASSUME: c_0_7 is a esa of c_0_4
SUCCESS: c_0_8 has SAT parents c_0_5 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_8 is a thm of c_0_5 
SUCCESS: c_0_9 has SAT parents c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_9 is a thm of c_0_6 
SUCCESS: c_0_10 has SAT parents c_0_7 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_10 is a thm of c_0_7 
SUCCESS: c_0_11 has SAT parents c_0_6 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_11 is a thm of c_0_6 
SUCCESS: c_0_12 has SAT parents c_0_8 c_0_9 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_12 is a thm of c_0_8 c_0_9 
SUCCESS: c_0_13 has SAT parents c_0_10 c_0_11 c_0_12 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_13 is a thm of c_0_10 c_0_11 c_0_12 
SUCCESS: c_0_14 has SAT parents c_0_7 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_7 
WARNING: c_0_15 has UNS parents c_0_13 c_0_14 c_0_12 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_15 is a thm of c_0_13 c_0_14 c_0_12 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 38.03
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN068+1/E---3.2.0.THM-CRf.s
fof(pel44,conjecture,
    ? [X1] :
      ( big_j(X1)
      & ~ big_f(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.lmwj7DqSWX/E---3.1_14673.p',pel44),
    [verified_file(pel44_thm),verified(leaf),verified(lambdapi)] ).
fof(pel44_1,axiom,
    ! [X1] :
      ( big_f(X1)
     => ( ? [X2] :
            ( big_g(X2)
            & big_h(X1,X2) )
        & ? [X3] :
            ( big_g(X3)
            & ~ big_h(X1,X3) ) ) ),
    file('/export/starexec/sandbox/tmp/tmp.lmwj7DqSWX/E---3.1_14673.p',pel44_1),
    [verified_file(pel44_1_thm),verified(leaf),verified(lambdapi)] ).
fof(pel44_2,axiom,
    ? [X1] :
      ( big_j(X1)
      & ! [X2] :
          ( big_g(X2)
         => big_h(X1,X2) ) ),
    file('/export/starexec/sandbox/tmp/tmp.lmwj7DqSWX/E---3.1_14673.p',pel44_2),
    [verified_file(pel44_2_thm),verified(leaf),nnpp(pel44),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ~ ? [X1] :
        ( big_j(X1)
        & ~ big_f(X1) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[pel44])]),
    [nnpp(pel44),verified_file(c_0_3_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_4,plain,
    ! [X1] :
      ( big_f(X1)
     => ( ? [X2] :
            ( big_g(X2)
            & big_h(X1,X2) )
        & ? [X3] :
            ( big_g(X3)
            & ~ big_h(X1,X3) ) ) ),
    inference(fof_simplification,[status(thm)],[pel44_1]),
    [nnpp(pel44),verified_file(c_0_4_thm),verified(thm),nnpp(pel44),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X4] :
      ( ~ big_j(X4)
      | big_f(X4) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])]),
    [nnpp(pel44),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,plain,
    ! [X9] :
      ( big_j(esk3_0)
      & ( ~ big_g(X9)
        | big_h(esk3_0,X9) ) ),
    inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel44_2])])])])]),
    [nnpp(pel44),verified_file(c_0_6_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_7,plain,
    ! [X5] :
      ( ( big_g(esk1_1(X5))
        | ~ big_f(X5) )
      & ( big_h(X5,esk1_1(X5))
        | ~ big_f(X5) )
      & ( big_g(esk2_1(X5))
        | ~ big_f(X5) )
      & ( ~ big_h(X5,esk2_1(X5))
        | ~ big_f(X5) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])]),
    [nnpp(pel44),verified_file(c_0_7_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X1] :
      ( big_f(X1)
      | ~ big_j(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel44),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,plain,
    big_j(esk3_0),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel44),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,plain,
    ! [X1] :
      ( ~ big_h(X1,esk2_1(X1))
      | ~ big_f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel44),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ! [X1] :
      ( big_h(esk3_0,X1)
      | ~ big_g(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel44),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    big_f(esk3_0),
    inference(spm,[status(thm)],[c_0_8,c_0_9]),
    [nnpp(pel44),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    ~ big_g(esk2_1(esk3_0)),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10,c_0_11]),c_0_12])]),
    [nnpp(pel44),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,plain,
    ! [X1] :
      ( big_g(esk2_1(X1))
      | ~ big_f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel44),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,plain,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13,c_0_14]),c_0_12])]),
    [proof,nnpp(pel44),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN068+1/E---3.2.0.THM-CRf.s
