SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: x2107 has SAT parents the problem 
 RESULT: x2107_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: x2107 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents x2107 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_1 is a cth of x2107 
WARNING: c_0_2 has UNS parents c_0_1 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_2 is a thm of c_0_1 
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_2 
WARNING: c_0_6 has UNS parents c_0_2 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_2 
SUCCESS: c_0_7 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of c_0_3 c_0_4 
SUCCESS: c_0_8 has SAT parents c_0_5 c_0_4 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_8 is a thm of c_0_5 c_0_4 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_6 c_0_7 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_6 c_0_7 
SUCCESS: c_0_11 has SAT parents c_0_3 c_0_8 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_11 is a thm of c_0_3 c_0_8 
SUCCESS: c_0_12 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_12 is a thm of c_0_9 c_0_10 
WARNING: c_0_13 has UNS parents c_0_6 c_0_11 c_0_12 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_13 is a thm of c_0_6 c_0_11 c_0_12 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 7.96
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN356+1/E---3.2.0.THM-CRf.s
fof(x2107,conjecture,
    ( ( big_r(a,b)
      & ! [X1,X2] :
          ( big_r(X1,X2)
         => ( big_r(X2,X1)
            & big_q(X1,X2) ) )
      & ! [X3,X4] :
          ( big_q(X3,X4)
         => big_q(X3,X3) ) )
   => ( big_q(a,a)
      & big_q(b,b) ) ),
    file('/export/starexec/sandbox/tmp/tmp.K0Dq1O30ct/E---3.1_16635.p',x2107),
    [verified_file(x2107_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( big_r(a,b)
        & ! [X1,X2] :
            ( big_r(X1,X2)
           => ( big_r(X2,X1)
              & big_q(X1,X2) ) )
        & ! [X3,X4] :
            ( big_q(X3,X4)
           => big_q(X3,X3) ) )
     => ( big_q(a,a)
        & big_q(b,b) ) ),
    inference(assume_negation,[status(cth)],[x2107]),
    [nnpp(x2107),verified_file(c_0_1_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X5,X6,X7,X8] :
      ( big_r(a,b)
      & ( big_r(X6,X5)
        | ~ big_r(X5,X6) )
      & ( big_q(X5,X6)
        | ~ big_r(X5,X6) )
      & ( ~ big_q(X7,X8)
        | big_q(X7,X7) )
      & ( ~ big_q(a,a)
        | ~ big_q(b,b) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])]),
    [nnpp(x2107),verified_file(c_0_2_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( big_q(X1,X2)
      | ~ big_r(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2107),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    big_r(a,b),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2107),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X1] :
      ( big_r(X1,X2)
      | ~ big_r(X2,X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2107),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X2,X1] :
      ( big_q(X1,X1)
      | ~ big_q(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2107),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    big_q(a,b),
    inference(spm,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(x2107),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    big_r(b,a),
    inference(spm,[status(thm)],[c_0_5,c_0_4]),
    [nnpp(x2107),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ( ~ big_q(a,a)
    | ~ big_q(b,b) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2107),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    big_q(a,a),
    inference(spm,[status(thm)],[c_0_6,c_0_7]),
    [nnpp(x2107),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    big_q(b,a),
    inference(spm,[status(thm)],[c_0_3,c_0_8]),
    [nnpp(x2107),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ~ big_q(b,b),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9,c_0_10])]),
    [nnpp(x2107),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_6,c_0_11]),c_0_12]),
    [proof,nnpp(x2107),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN356+1/E---3.2.0.THM-CRf.s
