SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: thm31 has SAT parents the problem 
 RESULT: thm31_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: thm31 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents thm31 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a cth of thm31 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_6 is a thm of c_0_3 
SUCCESS: c_0_7 has SAT parents c_0_4 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_7 is a thm of c_0_4 
SUCCESS: c_0_8 has SAT parents c_0_5 c_0_6 c_0_7 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_8 is a thm of c_0_5 c_0_6 c_0_7 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.07
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_7 c_0_8 c_0_6 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_10 is a thm of c_0_7 c_0_8 c_0_6 
SUCCESS: c_0_11 has SAT parents c_0_9 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_11 is a thm of c_0_9 
SUCCESS: c_0_12 has SAT parents c_0_5 c_0_10 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_12 is a thm of c_0_5 c_0_10 
WARNING: c_0_13 has UNS parents c_0_2 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_13 is a thm of c_0_2 
SUCCESS: c_0_14 has SAT parents c_0_11 c_0_6 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_14 is a thm of c_0_11 c_0_6 
SUCCESS: c_0_15 has SAT parents c_0_12 c_0_10 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.07
SUCCESS: c_0_15 is a thm of c_0_12 c_0_10 
SUCCESS: c_0_16 has SAT parents c_0_13 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_16 is a thm of c_0_13 
SUCCESS: c_0_17 has SAT parents c_0_14 c_0_15 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_17 is a thm of c_0_14 c_0_15 
SUCCESS: c_0_18 has SAT parents c_0_16 c_0_7 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_18 is a thm of c_0_16 c_0_7 
SUCCESS: c_0_19 has SAT parents c_0_5 c_0_17 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_19 is a thm of c_0_5 c_0_17 
SUCCESS: c_0_20 has SAT parents c_0_18 c_0_15 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_20 is a thm of c_0_18 c_0_15 
WARNING: c_0_21 has UNS parents c_0_19 c_0_17 c_0_20 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_21 is a thm of c_0_19 c_0_17 c_0_20 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 11.41
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN724+1/E---3.2.0.THM-CRf.s
fof(thm31,conjecture,
    ( ! [X1] :
        ( r(X1)
       => s(X1) )
  <=> ! [X1] :
        ( ( r(X1)
          & s(X1) )
      <=> r(X1) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.MnHXQjTGiv/E---3.1_29576.p',thm31),
    [verified_file(thm31_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ! [X1] :
          ( r(X1)
         => s(X1) )
    <=> ! [X1] :
          ( ( r(X1)
            & s(X1) )
        <=> r(X1) ) ),
    inference(assume_negation,[status(cth)],[thm31]),
    [nnpp(thm31),verified_file(c_0_1_thm),verified(cth),nnpp(thm31),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X4,X5] :
      ( ( ~ r(esk2_0)
        | ~ s(esk2_0)
        | ~ r(esk2_0)
        | r(esk1_0) )
      & ( r(esk2_0)
        | r(esk2_0)
        | r(esk1_0) )
      & ( s(esk2_0)
        | r(esk2_0)
        | r(esk1_0) )
      & ( ~ r(esk2_0)
        | ~ s(esk2_0)
        | ~ r(esk2_0)
        | ~ s(esk1_0) )
      & ( r(esk2_0)
        | r(esk2_0)
        | ~ s(esk1_0) )
      & ( s(esk2_0)
        | r(esk2_0)
        | ~ s(esk1_0) )
      & ( ~ r(X5)
        | ~ s(X5)
        | r(X5)
        | ~ r(X4)
        | s(X4) )
      & ( r(X5)
        | ~ r(X5)
        | ~ r(X4)
        | s(X4) )
      & ( s(X5)
        | ~ r(X5)
        | ~ r(X4)
        | s(X4) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])]),
    [nnpp(thm31),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ( r(esk2_0)
    | r(esk2_0)
    | r(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm31),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( r(esk2_0)
    | r(esk2_0)
    | ~ s(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm31),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X1] :
      ( s(X1)
      | s(X2)
      | ~ r(X1)
      | ~ r(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm31),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ( r(esk1_0)
    | r(esk2_0) ),
    inference(cn,[status(thm)],[c_0_3]),
    [nnpp(thm31),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ( r(esk2_0)
    | ~ s(esk1_0) ),
    inference(cn,[status(thm)],[c_0_4]),
    [nnpp(thm31),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X1] :
      ( s(X1)
      | r(esk2_0)
      | ~ r(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_5,c_0_6]),c_0_7]),
    [nnpp(thm31),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ( r(esk1_0)
    | ~ r(esk2_0)
    | ~ s(esk2_0)
    | ~ r(esk2_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm31),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    r(esk2_0),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7,c_0_8]),c_0_6]),
    [nnpp(thm31),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ( r(esk1_0)
    | ~ r(esk2_0)
    | ~ s(esk2_0) ),
    inference(cn,[status(thm)],[c_0_9]),
    [nnpp(thm31),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X1] :
      ( s(esk2_0)
      | s(X1)
      | ~ r(X1) ),
    inference(spm,[status(thm)],[c_0_5,c_0_10]),
    [nnpp(thm31),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ( ~ r(esk2_0)
    | ~ s(esk2_0)
    | ~ r(esk2_0)
    | ~ s(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm31),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ( r(esk1_0)
    | ~ s(esk2_0) ),
    inference(csr,[status(thm)],[c_0_11,c_0_6]),
    [nnpp(thm31),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    s(esk2_0),
    inference(spm,[status(thm)],[c_0_12,c_0_10]),
    [nnpp(thm31),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ( ~ r(esk2_0)
    | ~ s(esk1_0)
    | ~ s(esk2_0) ),
    inference(cn,[status(thm)],[c_0_13]),
    [nnpp(thm31),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    r(esk1_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14,c_0_15])]),
    [nnpp(thm31),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    ( ~ s(esk1_0)
    | ~ s(esk2_0) ),
    inference(csr,[status(thm)],[c_0_16,c_0_7]),
    [nnpp(thm31),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,negated_conjecture,
    ! [X1] :
      ( s(esk1_0)
      | s(X1)
      | ~ r(X1) ),
    inference(spm,[status(thm)],[c_0_5,c_0_17]),
    [nnpp(thm31),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,negated_conjecture,
    ~ s(esk1_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18,c_0_15])]),
    [nnpp(thm31),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19,c_0_17]),c_0_20]),
    [proof,nnpp(thm31),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN724+1/E---3.2.0.THM-CRf.s
