SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel26 has SAT parents the problem 
 RESULT: pel26_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel26 is a thm of the problem 
SUCCESS: pel26_1 has SAT parents the problem 
 RESULT: pel26_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: pel26_1 is a thm of the problem 
SUCCESS: pel26_2 has SAT parents the problem 
 RESULT: pel26_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel26_2 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_3 has SAT parents pel26 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_3 is a cth of pel26 
SUCCESS: c_0_4 has SAT parents pel26_1 (Theorem esa)
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.01
FAILURE: c_0_4 fails to be a thm of pel26_1 (Theorem esa)
SUCCESS: pel26_1 has SAT parents c_0_4 (Inverted esa)
 RESULT: c_0_4_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel26_1 is a thm of c_0_4 (Inverted esa)
FAILURE: c_0_4 fails to be a esa of pel26_1
 ASSUME: c_0_4 is a esa of pel26_1
SUCCESS: c_0_5 has SAT parents c_0_3 (Theorem esa)
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.00
FAILURE: c_0_5 fails to be a thm of c_0_3 (Theorem esa)
SUCCESS: c_0_3 has SAT parents c_0_5 (Inverted esa)
 RESULT: c_0_5_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_3 is a thm of c_0_5 (Inverted esa)
FAILURE: c_0_5 fails to be a esa of c_0_3
 ASSUME: c_0_5 is a esa of c_0_3
SUCCESS: c_0_6 has SAT parents pel26_2 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of pel26_2 
SUCCESS: c_0_7 has SAT parents c_0_4 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_7 is a thm of c_0_4 
SUCCESS: c_0_8 has SAT parents c_0_5 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_5 
SUCCESS: c_0_9 has SAT parents c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_6 
SUCCESS: c_0_10 has SAT parents c_0_5 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_10 is a thm of c_0_5 
SUCCESS: c_0_11 has SAT parents c_0_4 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_11 is a thm of c_0_4 
SUCCESS: c_0_12 has SAT parents c_0_7 c_0_8 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_12 is a thm of c_0_7 c_0_8 
SUCCESS: c_0_13 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_13 is a thm of c_0_9 c_0_10 
SUCCESS: c_0_14 has SAT parents c_0_11 c_0_12 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_11 c_0_12 
SUCCESS: c_0_15 has SAT parents c_0_6 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_15 is a thm of c_0_6 
SUCCESS: c_0_16 has SAT parents c_0_5 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_16 is a thm of c_0_5 
SUCCESS: c_0_17 has SAT parents c_0_13 c_0_14 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_17 is a thm of c_0_13 c_0_14 
SUCCESS: c_0_18 has SAT parents c_0_5 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_18 is a thm of c_0_5 
SUCCESS: c_0_19 has SAT parents c_0_15 c_0_13 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_19 is a thm of c_0_15 c_0_13 
SUCCESS: c_0_20 has SAT parents c_0_7 c_0_14 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_20 is a thm of c_0_7 c_0_14 
SUCCESS: c_0_21 has SAT parents c_0_5 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_21 is a thm of c_0_5 
SUCCESS: c_0_22 has SAT parents c_0_16 c_0_17 c_0_18 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_22 is a thm of c_0_16 c_0_17 c_0_18 
SUCCESS: c_0_23 has SAT parents c_0_19 c_0_20 
 RESULT: c_0_23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_23 is a thm of c_0_19 c_0_20 
SUCCESS: c_0_24 has SAT parents c_0_21 c_0_17 c_0_8 
 RESULT: c_0_24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_24 is a thm of c_0_21 c_0_17 c_0_8 
WARNING: c_0_25 has UNS parents c_0_22 c_0_23 c_0_24 
 RESULT: c_0_25_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_25 is a thm of c_0_22 c_0_23 c_0_24 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 10.03
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN056+1/E---3.2.0.THM-CRf.s
fof(pel26,conjecture,
    ( ! [X1] :
        ( big_p(X1)
       => big_r(X1) )
  <=> ! [X2] :
        ( big_q(X2)
       => big_s(X2) ) ),
    file('/export/starexec/sandbox/tmp/tmp.NRoZJMDlFC/E---3.1_3494.p',pel26),
    [verified_file(pel26_thm),verified(leaf),verified(lambdapi)] ).
fof(pel26_1,axiom,
    ( ? [X1] : big_p(X1)
  <=> ? [X2] : big_q(X2) ),
    file('/export/starexec/sandbox/tmp/tmp.NRoZJMDlFC/E---3.1_3494.p',pel26_1),
    [verified_file(pel26_1_thm),verified(leaf),nnpp(pel26),verified(lambdapi)] ).
fof(pel26_2,axiom,
    ! [X1,X2] :
      ( ( big_p(X1)
        & big_q(X2) )
     => ( big_r(X1)
      <=> big_s(X2) ) ),
    file('/export/starexec/sandbox/tmp/tmp.NRoZJMDlFC/E---3.1_3494.p',pel26_2),
    [verified_file(pel26_2_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ~ ( ! [X1] :
          ( big_p(X1)
         => big_r(X1) )
    <=> ! [X2] :
          ( big_q(X2)
         => big_s(X2) ) ),
    inference(assume_negation,[status(cth)],[pel26]),
    [nnpp(pel26),verified_file(c_0_3_thm),verified(cth),nnpp(pel26),verified(lambdapi)] ).
fof(c_0_4,plain,
    ! [X9,X11] :
      ( ( ~ big_p(X9)
        | big_q(esk3_0) )
      & ( ~ big_q(X11)
        | big_p(esk4_0) ) ),
    inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel26_1])])])])]),
    [nnpp(pel26),verified_file(c_0_4_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X5,X6] :
      ( ( big_q(esk2_0)
        | big_p(esk1_0) )
      & ( ~ big_s(esk2_0)
        | big_p(esk1_0) )
      & ( big_q(esk2_0)
        | ~ big_r(esk1_0) )
      & ( ~ big_s(esk2_0)
        | ~ big_r(esk1_0) )
      & ( ~ big_p(X5)
        | big_r(X5)
        | ~ big_q(X6)
        | big_s(X6) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])]),
    [nnpp(pel26),verified_file(c_0_5_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_6,plain,
    ! [X7,X8] :
      ( ( ~ big_r(X7)
        | big_s(X8)
        | ~ big_p(X7)
        | ~ big_q(X8) )
      & ( ~ big_s(X8)
        | big_r(X7)
        | ~ big_p(X7)
        | ~ big_q(X8) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel26_2])])])]),
    [nnpp(pel26),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,plain,
    ! [X1] :
      ( big_p(esk4_0)
      | ~ big_q(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel26),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ( big_q(esk2_0)
    | big_p(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel26),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,plain,
    ! [X1,X2] :
      ( big_r(X2)
      | ~ big_s(X1)
      | ~ big_p(X2)
      | ~ big_q(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel26),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ! [X2,X1] :
      ( big_r(X1)
      | big_s(X2)
      | ~ big_p(X1)
      | ~ big_q(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel26),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ! [X1] :
      ( big_q(esk3_0)
      | ~ big_p(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel26),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ( big_p(esk1_0)
    | big_p(esk4_0) ),
    inference(spm,[status(thm)],[c_0_7,c_0_8]),
    [nnpp(pel26),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    ! [X2,X1] :
      ( big_r(X1)
      | ~ big_q(X2)
      | ~ big_p(X1) ),
    inference(csr,[status(thm)],[c_0_9,c_0_10]),
    [nnpp(pel26),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    big_q(esk3_0),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11,c_0_12]),c_0_11]),
    [nnpp(pel26),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,plain,
    ! [X1,X2] :
      ( big_s(X2)
      | ~ big_r(X1)
      | ~ big_p(X1)
      | ~ big_q(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel26),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ( ~ big_s(esk2_0)
    | ~ big_r(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel26),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ! [X1] :
      ( big_r(X1)
      | ~ big_p(X1) ),
    inference(spm,[status(thm)],[c_0_13,c_0_14]),
    [nnpp(pel26),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    ( big_p(esk1_0)
    | ~ big_s(esk2_0) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel26),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,plain,
    ! [X2,X1] :
      ( big_s(X1)
      | ~ big_q(X1)
      | ~ big_p(X2) ),
    inference(csr,[status(thm)],[c_0_15,c_0_13]),
    [nnpp(pel26),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,negated_conjecture,
    big_p(esk4_0),
    inference(spm,[status(thm)],[c_0_7,c_0_14]),
    [nnpp(pel26),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,negated_conjecture,
    ( big_q(esk2_0)
    | ~ big_r(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel26),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,negated_conjecture,
    ~ big_s(esk2_0),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16,c_0_17]),c_0_18]),
    [nnpp(pel26),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_23,negated_conjecture,
    ! [X1] :
      ( big_s(X1)
      | ~ big_q(X1) ),
    inference(spm,[status(thm)],[c_0_19,c_0_20]),
    [nnpp(pel26),verified_file(c_0_23_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_24,negated_conjecture,
    big_q(esk2_0),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21,c_0_17]),c_0_8]),
    [nnpp(pel26),verified_file(c_0_24_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_25,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22,c_0_23]),c_0_24])]),
    [proof,nnpp(pel26),verified_file(c_0_25_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN056+1/E---3.2.0.THM-CRf.s
