SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel23 has SAT parents the problem 
 RESULT: pel23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: pel23 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents pel23 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a cth of pel23 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
SUCCESS: c_0_5 has SAT parents c_0_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_3 
SUCCESS: c_0_6 has SAT parents c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_6 is a thm of c_0_4 
SUCCESS: c_0_7 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_5 c_0_6 
WARNING: c_0_8 has UNS parents c_0_2 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_8 is a thm of c_0_2 
SUCCESS: c_0_9 has SAT parents c_0_7 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_9 is a thm of c_0_7 
WARNING: c_0_10 has UNS parents c_0_8 c_0_9 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_8 c_0_9 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 6.83
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN053+1/E---3.2.0.THM-CRf.s
fof(pel23,conjecture,
    ( ! [X1] :
        ( p
        | big_f(X1) )
  <=> ( p
      | ! [X2] : big_f(X2) ) ),
    file('/export/starexec/sandbox/tmp/tmp.es2BfKOJ4k/E---3.1_22369.p',pel23),
    [verified_file(pel23_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ! [X1] :
          ( p
          | big_f(X1) )
    <=> ( p
        | ! [X2] : big_f(X2) ) ),
    inference(assume_negation,[status(cth)],[pel23]),
    [nnpp(pel23),verified_file(c_0_1_thm),verified(cth),nnpp(pel23),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X5,X6] :
      ( ( ~ p
        | ~ p )
      & ( ~ big_f(esk2_0)
        | ~ p )
      & ( ~ p
        | ~ big_f(esk1_0) )
      & ( ~ big_f(esk2_0)
        | ~ big_f(esk1_0) )
      & ( p
        | big_f(X5)
        | p
        | big_f(X6) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])]),
    [nnpp(pel23),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( p
      | big_f(X1)
      | p
      | big_f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel23),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( ~ p
    | ~ p ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel23),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X1,X2] :
      ( p
      | big_f(X2)
      | big_f(X1) ),
    inference(cn,[status(thm)],[c_0_3]),
    [nnpp(pel23),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ~ p,
    inference(cn,[status(thm)],[c_0_4]),
    [nnpp(pel23),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X2,X1] :
      ( big_f(X1)
      | big_f(X2) ),
    inference(sr,[status(thm)],[c_0_5,c_0_6]),
    [nnpp(pel23),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ( ~ big_f(esk2_0)
    | ~ big_f(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel23),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] : big_f(X1),
    inference(ef,[status(thm)],[c_0_7]),
    [nnpp(pel23),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_8,c_0_9]),c_0_9])]),
    [proof,nnpp(pel23),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN053+1/E---3.2.0.THM-CRf.s
