SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel10 has SAT parents the problem 
 RESULT: pel10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel10 is a thm of the problem 
SUCCESS: pel10_3 has SAT parents the problem 
 RESULT: pel10_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel10_3 is a thm of the problem 
SUCCESS: pel10_1 has SAT parents the problem 
 RESULT: pel10_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: pel10_1 is a thm of the problem 
SUCCESS: pel10_2 has SAT parents the problem 
 RESULT: pel10_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel10_2 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_4 has SAT parents pel10 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a cth of pel10 
SUCCESS: c_0_5 has SAT parents pel10_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_5 is a thm of pel10_3 
SUCCESS: c_0_6 has SAT parents c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_4 
SUCCESS: c_0_7 has SAT parents pel10_1 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of pel10_1 
SUCCESS: c_0_8 has SAT parents pel10_2 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of pel10_2 
SUCCESS: c_0_9 has SAT parents c_0_5 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_5 
SUCCESS: c_0_10 has SAT parents c_0_6 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_10 is a thm of c_0_6 
SUCCESS: c_0_11 has SAT parents c_0_7 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_11 is a thm of c_0_7 
SUCCESS: c_0_12 has SAT parents c_0_8 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_8 
SUCCESS: c_0_13 has SAT parents c_0_9 c_0_10 c_0_11 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_13 is a thm of c_0_9 c_0_10 c_0_11 
SUCCESS: c_0_14 has SAT parents c_0_6 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_14 is a thm of c_0_6 
SUCCESS: c_0_15 has SAT parents c_0_12 c_0_13 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_15 is a thm of c_0_12 c_0_13 
SUCCESS: c_0_16 has SAT parents c_0_8 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_16 is a thm of c_0_8 
SUCCESS: c_0_17 has SAT parents c_0_14 c_0_15 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_17 is a thm of c_0_14 c_0_15 
WARNING: c_0_18 has UNS parents c_0_16 c_0_13 c_0_17 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_18 is a thm of c_0_16 c_0_13 c_0_17 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 6.59
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN044+1/E---3.2.0.THM-CRf.s
fof(pel10,conjecture,
    ( p
  <=> q ),
    file('/export/starexec/sandbox/tmp/tmp.zb7NlAcgLk/E---3.1_31143.p',pel10),
    [verified_file(pel10_thm),verified(leaf),verified(lambdapi)] ).
fof(pel10_3,axiom,
    ( p
   => ( q
      | r ) ),
    file('/export/starexec/sandbox/tmp/tmp.zb7NlAcgLk/E---3.1_31143.p',pel10_3),
    [verified_file(pel10_3_thm),verified(leaf),verified(lambdapi)] ).
fof(pel10_1,axiom,
    ( q
   => r ),
    file('/export/starexec/sandbox/tmp/tmp.zb7NlAcgLk/E---3.1_31143.p',pel10_1),
    [verified_file(pel10_1_thm),verified(leaf),verified(lambdapi)] ).
fof(pel10_2,axiom,
    ( r
   => ( p
      & q ) ),
    file('/export/starexec/sandbox/tmp/tmp.zb7NlAcgLk/E---3.1_31143.p',pel10_2),
    [verified_file(pel10_2_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ~ ( p
    <=> q ),
    inference(assume_negation,[status(cth)],[pel10]),
    [nnpp(pel10),verified_file(c_0_4_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_5,plain,
    ( ~ p
    | q
    | r ),
    inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[pel10_3])]),
    [nnpp(pel10),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ( ( ~ p
      | ~ q )
    & ( p
      | q ) ),
    inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])]),
    [nnpp(pel10),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,plain,
    ( ~ q
    | r ),
    inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[pel10_1])]),
    [nnpp(pel10),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,plain,
    ( ( p
      | ~ r )
    & ( q
      | ~ r ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[pel10_2])])]),
    [nnpp(pel10),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,plain,
    ( q
    | r
    | ~ p ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel10),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ( p
    | q ),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel10),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ( r
    | ~ q ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel10),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,plain,
    ( p
    | ~ r ),
    inference(split_conjunct,[status(thm)],[c_0_8]),
    [nnpp(pel10),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    r,
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_9,c_0_10]),c_0_11]),
    [nnpp(pel10),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ( ~ p
    | ~ q ),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel10),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,plain,
    p,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12,c_0_13])]),
    [nnpp(pel10),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,plain,
    ( q
    | ~ r ),
    inference(split_conjunct,[status(thm)],[c_0_8]),
    [nnpp(pel10),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ~ q,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14,c_0_15])]),
    [nnpp(pel10),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,plain,
    $false,
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16,c_0_13])]),c_0_17]),
    [proof,nnpp(pel10),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN044+1/E---3.2.0.THM-CRf.s
