SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: x2125 has SAT parents the problem 
 RESULT: x2125_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: x2125 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents x2125 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_1 is a cth of x2125 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_4 is a thm of c_0_2 
SUCCESS: c_0_5 has SAT parents c_0_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_3 
SUCCESS: c_0_6 has SAT parents c_0_4 c_0_5 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_4 c_0_5 
SUCCESS: c_0_7 has SAT parents c_0_6 c_0_5 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of c_0_6 c_0_5 
SUCCESS: c_0_8 has SAT parents c_0_4 c_0_7 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_4 c_0_7 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_8 c_0_7 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_10 is a thm of c_0_8 c_0_7 
WARNING: c_0_11 has UNS parents c_0_2 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_11 is a thm of c_0_2 
SUCCESS: c_0_12 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_12 is a thm of c_0_9 c_0_10 
SUCCESS: c_0_13 has SAT parents c_0_11 c_0_12 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_13 is a thm of c_0_11 c_0_12 
SUCCESS: c_0_14 has SAT parents c_0_13 c_0_12 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_13 c_0_12 
SUCCESS: c_0_15 has SAT parents c_0_14 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_15 is a thm of c_0_14 
WARNING: c_0_16 has UNS parents c_0_9 c_0_15 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_16 is a thm of c_0_9 c_0_15 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 228.30
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN374+1/E---3.2.0.THM-CRf.s
fof(x2125,conjecture,
    ( ? [X1] :
      ! [X2] :
        ( big_p(X1)
      <=> big_p(X2) )
  <=> ( ? [X1] : big_p(X1)
    <=> ! [X2] : big_p(X2) ) ),
    file('/export/starexec/sandbox/tmp/tmp.Ttdmg3Etvb/E---3.1_12247.p',x2125),
    [verified_file(x2125_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ? [X1] :
        ! [X2] :
          ( big_p(X1)
        <=> big_p(X2) )
    <=> ( ? [X1] : big_p(X1)
      <=> ! [X2] : big_p(X2) ) ),
    inference(assume_negation,[status(cth)],[x2125]),
    [nnpp(x2125),verified_file(c_0_1_thm),verified(cth),nnpp(x2125),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X3,X5,X8,X10,X11,X12] :
      ( ( ~ big_p(X5)
        | ~ big_p(esk2_0)
        | ~ big_p(X3)
        | ~ big_p(esk1_1(X3)) )
      & ( big_p(esk3_0)
        | big_p(X8)
        | ~ big_p(X3)
        | ~ big_p(esk1_1(X3)) )
      & ( ~ big_p(X5)
        | ~ big_p(esk2_0)
        | big_p(X3)
        | big_p(esk1_1(X3)) )
      & ( big_p(esk3_0)
        | big_p(X8)
        | big_p(X3)
        | big_p(esk1_1(X3)) )
      & ( ~ big_p(X11)
        | big_p(X12)
        | ~ big_p(esk4_0)
        | big_p(X10) )
      & ( ~ big_p(esk5_0)
        | big_p(esk6_0)
        | ~ big_p(esk4_0)
        | big_p(X10) )
      & ( ~ big_p(X11)
        | big_p(X12)
        | ~ big_p(X10)
        | big_p(esk4_0) )
      & ( ~ big_p(esk5_0)
        | big_p(esk6_0)
        | ~ big_p(X10)
        | big_p(esk4_0) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])]),
    [nnpp(x2125),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( big_p(esk3_0)
      | big_p(X1)
      | big_p(X2)
      | big_p(esk1_1(X2)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2125),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X3,X1,X2] :
      ( big_p(X2)
      | big_p(esk4_0)
      | ~ big_p(X1)
      | ~ big_p(X3) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2125),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X1] :
      ( big_p(esk1_1(X1))
      | big_p(esk3_0)
      | big_p(X1) ),
    inference(ef,[status(thm)],[c_0_3]),
    [nnpp(x2125),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X1] :
      ( big_p(esk3_0)
      | big_p(esk4_0)
      | ~ big_p(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_4,c_0_5]),c_0_4]),c_0_4]),
    [nnpp(x2125),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ( big_p(esk4_0)
    | big_p(esk3_0) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_6,c_0_5]),c_0_6]),
    [nnpp(x2125),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X2,X1] :
      ( big_p(esk4_0)
      | big_p(X1)
      | ~ big_p(X2) ),
    inference(spm,[status(thm)],[c_0_4,c_0_7]),
    [nnpp(x2125),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X2,X1] :
      ( ~ big_p(X1)
      | ~ big_p(esk2_0)
      | ~ big_p(X2)
      | ~ big_p(esk1_1(X2)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2125),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ! [X1] :
      ( big_p(esk4_0)
      | big_p(X1) ),
    inference(spm,[status(thm)],[c_0_8,c_0_7]),
    [nnpp(x2125),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ! [X1,X3,X2] :
      ( big_p(X2)
      | big_p(X3)
      | ~ big_p(X1)
      | ~ big_p(esk4_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2125),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    big_p(esk4_0),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9,c_0_10]),c_0_10]),c_0_10]),c_0_10]),
    [nnpp(x2125),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ! [X3,X2,X1] :
      ( big_p(X1)
      | big_p(X2)
      | ~ big_p(X3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11,c_0_12])]),
    [nnpp(x2125),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ! [X2,X1] :
      ( big_p(X1)
      | big_p(X2) ),
    inference(spm,[status(thm)],[c_0_13,c_0_12]),
    [nnpp(x2125),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ! [X1] : big_p(X1),
    inference(ef,[status(thm)],[c_0_14]),
    [nnpp(x2125),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_9,c_0_15]),c_0_15]),c_0_15]),c_0_15])]),
    [proof,nnpp(x2125),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN374+1/E---3.2.0.THM-CRf.s
