SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
 RESULT: problem_axioms_sat_model - Paradox---4.0 says Satisfiable - CPU = 0.00
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel24_4 has SAT parents the problem 
 RESULT: pel24_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel24_4 is a thm of the problem 
SUCCESS: pel24_1 has SAT parents the problem 
 RESULT: pel24_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel24_1 is a thm of the problem 
SUCCESS: pel24 has SAT parents the problem 
 RESULT: pel24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel24 is a thm of the problem 
SUCCESS: pel24_3 has SAT parents the problem 
 RESULT: pel24_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel24_3 is a thm of the problem 
SUCCESS: pel24_2 has SAT parents the problem 
 RESULT: pel24_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel24_2 is a thm of the problem 
SUCCESS: Leaves are verified
 RESULT: leaf_axioms_sat_model - Paradox---4.0 says Satisfiable - CPU = 0.00
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_5 has SAT parents pel24_4 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_5 is a thm of pel24_4 
SUCCESS: c_0_6 has SAT parents pel24_1 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_6 is a thm of pel24_1 
SUCCESS: c_0_7 has SAT parents pel24 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a cth of pel24 
SUCCESS: c_0_8 has SAT parents pel24_3 (Theorem esa)
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.01
FAILURE: c_0_8 fails to be a thm of pel24_3 (Theorem esa)
SUCCESS: pel24_3 has SAT parents c_0_8 (Inverted esa)
 RESULT: c_0_8_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel24_3 is a thm of c_0_8 (Inverted esa)
FAILURE: c_0_8 fails to be a esa of pel24_3
 ASSUME: c_0_8 is a esa of pel24_3
SUCCESS: c_0_9 has SAT parents c_0_5 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_9 is a thm of c_0_5 
SUCCESS: c_0_10 has SAT parents c_0_6 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_10 is a thm of c_0_6 
SUCCESS: c_0_11 has SAT parents pel24_2 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_11 is a thm of pel24_2 
SUCCESS: c_0_12 has SAT parents c_0_7 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_12 is a thm of c_0_7 
SUCCESS: c_0_13 has SAT parents c_0_8 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_13 is a thm of c_0_8 
SUCCESS: c_0_14 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_14 is a thm of c_0_9 c_0_10 
SUCCESS: c_0_15 has SAT parents c_0_11 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_15 is a thm of c_0_11 
SUCCESS: c_0_16 has SAT parents c_0_12 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_16 is a thm of c_0_12 
SUCCESS: c_0_17 has SAT parents c_0_13 c_0_14 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_17 is a thm of c_0_13 c_0_14 
SUCCESS: c_0_18 has SAT parents c_0_15 c_0_14 c_0_16 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_18 is a thm of c_0_15 c_0_14 c_0_16 
WARNING: c_0_19 has UNS parents c_0_17 c_0_18 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_19 is a thm of c_0_17 c_0_18 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 7.87
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN054+1/E---3.2.0.THM-CRf.s
fof(pel24_4,axiom,
    ! [X1] :
      ( ( big_q(X1)
        | big_r(X1) )
     => big_s(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.SUDtrE3WdA/E---3.1_14596.p',pel24_4),
    [verified_file(pel24_4_thm),verified(leaf),verified(lambdapi)] ).
fof(pel24_1,axiom,
    ~ ? [X1] :
        ( big_s(X1)
        & big_q(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.SUDtrE3WdA/E---3.1_14596.p',pel24_1),
    [verified_file(pel24_1_thm),verified(leaf),verified(lambdapi)] ).
fof(pel24,conjecture,
    ? [X1] :
      ( big_p(X1)
      & big_r(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.SUDtrE3WdA/E---3.1_14596.p',pel24),
    [verified_file(pel24_thm),verified(leaf),verified(lambdapi)] ).
fof(pel24_3,axiom,
    ( ~ ? [X1] : big_p(X1)
   => ? [X2] : big_q(X2) ),
    file('/export/starexec/sandbox/tmp/tmp.SUDtrE3WdA/E---3.1_14596.p',pel24_3),
    [verified_file(pel24_3_thm),verified(leaf),nnpp(pel24),verified(lambdapi)] ).
fof(pel24_2,axiom,
    ! [X1] :
      ( big_p(X1)
     => ( big_q(X1)
        | big_r(X1) ) ),
    file('/export/starexec/sandbox/tmp/tmp.SUDtrE3WdA/E---3.1_14596.p',pel24_2),
    [verified_file(pel24_2_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_5,plain,
    ! [X5] :
      ( ( ~ big_q(X5)
        | big_s(X5) )
      & ( ~ big_r(X5)
        | big_s(X5) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel24_4])])])]),
    [nnpp(pel24),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,plain,
    ! [X8] :
      ( ~ big_s(X8)
      | ~ big_q(X8) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel24_1])])]),
    [nnpp(pel24),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ~ ? [X1] :
        ( big_p(X1)
        & big_r(X1) ),
    inference(assume_negation,[status(cth)],[pel24]),
    [nnpp(pel24),verified_file(c_0_7_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_8,plain,
    ( big_p(esk1_0)
    | big_q(esk2_0) ),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel24_3])])]),
    [nnpp(pel24),verified_file(c_0_8_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_9,plain,
    ! [X1] :
      ( big_s(X1)
      | ~ big_q(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel24),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,plain,
    ! [X1] :
      ( ~ big_s(X1)
      | ~ big_q(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel24),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ! [X4] :
      ( ~ big_p(X4)
      | big_q(X4)
      | big_r(X4) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel24_2])])]),
    [nnpp(pel24),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X3] :
      ( ~ big_p(X3)
      | ~ big_r(X3) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])]),
    [nnpp(pel24),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    ( big_p(esk1_0)
    | big_q(esk2_0) ),
    inference(split_conjunct,[status(thm)],[c_0_8]),
    [nnpp(pel24),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,plain,
    ! [X1] : ~ big_q(X1),
    inference(csr,[status(thm)],[c_0_9,c_0_10]),
    [nnpp(pel24),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,plain,
    ! [X1] :
      ( big_q(X1)
      | big_r(X1)
      | ~ big_p(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_11]),
    [nnpp(pel24),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ! [X1] :
      ( ~ big_p(X1)
      | ~ big_r(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_12]),
    [nnpp(pel24),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,plain,
    big_p(esk1_0),
    inference(sr,[status(thm)],[c_0_13,c_0_14]),
    [nnpp(pel24),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,plain,
    ! [X1] : ~ big_p(X1),
    inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_15,c_0_14]),c_0_16]),
    [nnpp(pel24),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,plain,
    $false,
    inference(sr,[status(thm)],[c_0_17,c_0_18]),
    [proof,nnpp(pel24),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN054+1/E---3.2.0.THM-CRf.s
