SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
 RESULT: problem_axioms_sat_model - Paradox---4.0 says Satisfiable - CPU = 0.00
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel31 has SAT parents the problem 
 RESULT: pel31_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel31 is a thm of the problem 
SUCCESS: pel31_3 has SAT parents the problem 
 RESULT: pel31_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: pel31_3 is a thm of the problem 
SUCCESS: pel31_1 has SAT parents the problem 
 RESULT: pel31_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel31_1 is a thm of the problem 
SUCCESS: pel31_2 has SAT parents the problem 
 RESULT: pel31_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel31_2 is a thm of the problem 
SUCCESS: Leaves are verified
 RESULT: leaf_axioms_sat_model - Paradox---4.0 says Satisfiable - CPU = 0.00
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_4 has SAT parents pel31 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a cth of pel31 
SUCCESS: c_0_5 has SAT parents pel31_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of pel31_3 
SUCCESS: c_0_6 has SAT parents c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_6 is a thm of c_0_4 
SUCCESS: c_0_7 has SAT parents c_0_5 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_5 
SUCCESS: c_0_8 has SAT parents pel31_1 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_8 is a thm of pel31_1 
SUCCESS: c_0_9 has SAT parents c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_6 
SUCCESS: c_0_10 has SAT parents c_0_7 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_10 is a thm of c_0_7 
SUCCESS: c_0_11 has SAT parents c_0_8 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_11 is a thm of c_0_8 
SUCCESS: c_0_12 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_9 c_0_10 
SUCCESS: c_0_13 has SAT parents pel31_2 (Theorem esa)
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.01
FAILURE: c_0_13 fails to be a thm of pel31_2 (Theorem esa)
SUCCESS: pel31_2 has SAT parents c_0_13 (Inverted esa)
 RESULT: c_0_13_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel31_2 is a thm of c_0_13 (Inverted esa)
FAILURE: c_0_13 fails to be a esa of pel31_2
 ASSUME: c_0_13 is a esa of pel31_2
SUCCESS: c_0_14 has SAT parents c_0_11 c_0_12 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_14 is a thm of c_0_11 c_0_12 
SUCCESS: c_0_15 has SAT parents c_0_13 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_15 is a thm of c_0_13 
SUCCESS: c_0_16 has SAT parents c_0_13 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_16 is a thm of c_0_13 
WARNING: c_0_17 has UNS parents c_0_14 c_0_15 c_0_16 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_17 is a thm of c_0_14 c_0_15 c_0_16 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 7.13
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN061+1/E---3.2.0.THM-CRf.s
fof(pel31,conjecture,
    ? [X1] :
      ( big_i(X1)
      & big_j(X1) ),
    file('/export/starexec/sandbox2/tmp/tmp.a06oxE6XQM/E---3.1_24135.p',pel31),
    [verified_file(pel31_thm),verified(leaf),verified(lambdapi)] ).
fof(pel31_3,axiom,
    ! [X1] :
      ( ~ big_h(X1)
     => big_j(X1) ),
    file('/export/starexec/sandbox2/tmp/tmp.a06oxE6XQM/E---3.1_24135.p',pel31_3),
    [verified_file(pel31_3_thm),verified(leaf),verified(lambdapi)] ).
fof(pel31_1,axiom,
    ~ ? [X1] :
        ( big_f(X1)
        & ( big_g(X1)
          | big_h(X1) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.a06oxE6XQM/E---3.1_24135.p',pel31_1),
    [verified_file(pel31_1_thm),verified(leaf),verified(lambdapi)] ).
fof(pel31_2,axiom,
    ? [X1] :
      ( big_i(X1)
      & big_f(X1) ),
    file('/export/starexec/sandbox2/tmp/tmp.a06oxE6XQM/E---3.1_24135.p',pel31_2),
    [verified_file(pel31_2_thm),verified(leaf),nnpp(pel31),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ~ ? [X1] :
        ( big_i(X1)
        & big_j(X1) ),
    inference(assume_negation,[status(cth)],[pel31]),
    [nnpp(pel31),verified_file(c_0_4_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_5,plain,
    ! [X1] :
      ( ~ big_h(X1)
     => big_j(X1) ),
    inference(fof_simplification,[status(thm)],[pel31_3]),
    [nnpp(pel31),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X5] :
      ( ~ big_i(X5)
      | ~ big_j(X5) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])]),
    [nnpp(pel31),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,plain,
    ! [X4] :
      ( big_h(X4)
      | big_j(X4) ),
    inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])]),
    [nnpp(pel31),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,plain,
    ! [X2] :
      ( ( ~ big_g(X2)
        | ~ big_f(X2) )
      & ( ~ big_h(X2)
        | ~ big_f(X2) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel31_1])])])]),
    [nnpp(pel31),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] :
      ( ~ big_i(X1)
      | ~ big_j(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel31),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,plain,
    ! [X1] :
      ( big_h(X1)
      | big_j(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel31),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ! [X1] :
      ( ~ big_h(X1)
      | ~ big_f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_8]),
    [nnpp(pel31),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X1] :
      ( big_h(X1)
      | ~ big_i(X1) ),
    inference(spm,[status(thm)],[c_0_9,c_0_10]),
    [nnpp(pel31),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    ( big_i(esk1_0)
    & big_f(esk1_0) ),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[pel31_2])]),
    [nnpp(pel31),verified_file(c_0_13_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ! [X1] :
      ( ~ big_i(X1)
      | ~ big_f(X1) ),
    inference(spm,[status(thm)],[c_0_11,c_0_12]),
    [nnpp(pel31),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,plain,
    big_i(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_13]),
    [nnpp(pel31),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,plain,
    big_f(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_13]),
    [nnpp(pel31),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14,c_0_15]),c_0_16])]),
    [proof,nnpp(pel31),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN061+1/E---3.2.0.THM-CRf.s
