SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel62 has SAT parents the problem 
 RESULT: pel62_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: pel62 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents pel62 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a cth of pel62 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_5 is a thm of c_0_2 
WARNING: c_0_6 has UNS parents c_0_2 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_6 is a thm of c_0_2 
SUCCESS: c_0_7 has SAT parents c_0_3 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_3 
SUCCESS: c_0_8 has SAT parents c_0_4 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_8 is a thm of c_0_4 
SUCCESS: c_0_9 has SAT parents c_0_5 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_5 
SUCCESS: c_0_10 has SAT parents c_0_6 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_10 is a thm of c_0_6 
SUCCESS: c_0_11 has SAT parents c_0_7 c_0_8 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_11 is a thm of c_0_7 c_0_8 
WARNING: c_0_12 has UNS parents c_0_2 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_2 
SUCCESS: c_0_13 has SAT parents c_0_9 c_0_8 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_13 is a thm of c_0_9 c_0_8 
SUCCESS: c_0_14 has SAT parents c_0_10 c_0_11 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_14 is a thm of c_0_10 c_0_11 
SUCCESS: c_0_15 has SAT parents c_0_12 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_15 is a thm of c_0_12 
SUCCESS: c_0_16 has SAT parents c_0_10 c_0_13 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_16 is a thm of c_0_10 c_0_13 
SUCCESS: c_0_17 has SAT parents c_0_14 c_0_15 c_0_16 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_17 is a thm of c_0_14 c_0_15 c_0_16 
WARNING: c_0_18 has UNS parents c_0_2 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_18 is a thm of c_0_2 
SUCCESS: c_0_19 has SAT parents c_0_17 c_0_13 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_19 is a thm of c_0_17 c_0_13 
SUCCESS: c_0_20 has SAT parents c_0_17 c_0_11 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_20 is a thm of c_0_17 c_0_11 
SUCCESS: c_0_21 has SAT parents c_0_18 c_0_13 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_21 is a thm of c_0_18 c_0_13 
SUCCESS: c_0_22 has SAT parents c_0_18 c_0_19 c_0_20 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_22 is a thm of c_0_18 c_0_19 c_0_20 
WARNING: c_0_23 has UNS parents c_0_2 
 RESULT: c_0_23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_23 is a thm of c_0_2 
SUCCESS: c_0_24 has SAT parents c_0_21 c_0_19 c_0_20 
 RESULT: c_0_24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_24 is a thm of c_0_21 c_0_19 c_0_20 
SUCCESS: c_0_25 has SAT parents c_0_22 c_0_11 
 RESULT: c_0_25_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_25 is a thm of c_0_22 c_0_11 
WARNING: c_0_26 has UNS parents c_0_23 c_0_19 c_0_24 c_0_20 c_0_25 
 RESULT: c_0_26_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_26 is a thm of c_0_23 c_0_19 c_0_24 c_0_20 c_0_25 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 303.44
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN084+1/E---3.2.0.THM-CRf.s
fof(pel62,conjecture,
    ( ! [X1] :
        ( ( big_p(a)
          & ( big_p(X1)
           => big_p(f(X1)) ) )
       => big_p(f(f(X1))) )
  <=> ! [X2] :
        ( ( ~ big_p(a)
          | big_p(X2)
          | big_p(f(f(X2))) )
        & ( ~ big_p(a)
          | ~ big_p(f(X2))
          | big_p(f(f(X2))) ) ) ),
    file('/export/starexec/sandbox/tmp/tmp.mGFZF8Has9/E---3.1_19547.p',pel62),
    [verified_file(pel62_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ! [X1] :
          ( ( big_p(a)
            & ( big_p(X1)
             => big_p(f(X1)) ) )
         => big_p(f(f(X1))) )
    <=> ! [X2] :
          ( ( ~ big_p(a)
            | big_p(X2)
            | big_p(f(f(X2))) )
          & ( ~ big_p(a)
            | ~ big_p(f(X2))
            | big_p(f(f(X2))) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[pel62])]),
    [nnpp(pel62),verified_file(c_0_1_thm),verified(cth),nnpp(pel62),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X5,X6] :
      ( ( big_p(a)
        | big_p(a)
        | big_p(a) )
      & ( big_p(f(esk2_0))
        | big_p(a)
        | big_p(a) )
      & ( ~ big_p(f(f(esk2_0)))
        | big_p(a)
        | big_p(a) )
      & ( big_p(a)
        | ~ big_p(esk2_0)
        | big_p(a) )
      & ( big_p(f(esk2_0))
        | ~ big_p(esk2_0)
        | big_p(a) )
      & ( ~ big_p(f(f(esk2_0)))
        | ~ big_p(esk2_0)
        | big_p(a) )
      & ( big_p(a)
        | ~ big_p(f(f(esk2_0)))
        | big_p(a) )
      & ( big_p(f(esk2_0))
        | ~ big_p(f(f(esk2_0)))
        | big_p(a) )
      & ( ~ big_p(f(f(esk2_0)))
        | ~ big_p(f(f(esk2_0)))
        | big_p(a) )
      & ( big_p(a)
        | big_p(a)
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( big_p(f(esk2_0))
        | big_p(a)
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( ~ big_p(f(f(esk2_0)))
        | big_p(a)
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( big_p(a)
        | ~ big_p(esk2_0)
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( big_p(f(esk2_0))
        | ~ big_p(esk2_0)
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( ~ big_p(f(f(esk2_0)))
        | ~ big_p(esk2_0)
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( big_p(a)
        | ~ big_p(f(f(esk2_0)))
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( big_p(f(esk2_0))
        | ~ big_p(f(f(esk2_0)))
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( ~ big_p(f(f(esk2_0)))
        | ~ big_p(f(f(esk2_0)))
        | ~ big_p(esk1_0)
        | big_p(f(esk1_0)) )
      & ( big_p(a)
        | big_p(a)
        | ~ big_p(f(f(esk1_0))) )
      & ( big_p(f(esk2_0))
        | big_p(a)
        | ~ big_p(f(f(esk1_0))) )
      & ( ~ big_p(f(f(esk2_0)))
        | big_p(a)
        | ~ big_p(f(f(esk1_0))) )
      & ( big_p(a)
        | ~ big_p(esk2_0)
        | ~ big_p(f(f(esk1_0))) )
      & ( big_p(f(esk2_0))
        | ~ big_p(esk2_0)
        | ~ big_p(f(f(esk1_0))) )
      & ( ~ big_p(f(f(esk2_0)))
        | ~ big_p(esk2_0)
        | ~ big_p(f(f(esk1_0))) )
      & ( big_p(a)
        | ~ big_p(f(f(esk2_0)))
        | ~ big_p(f(f(esk1_0))) )
      & ( big_p(f(esk2_0))
        | ~ big_p(f(f(esk2_0)))
        | ~ big_p(f(f(esk1_0))) )
      & ( ~ big_p(f(f(esk2_0)))
        | ~ big_p(f(f(esk2_0)))
        | ~ big_p(f(f(esk1_0))) )
      & ( ~ big_p(a)
        | big_p(X6)
        | big_p(f(f(X6)))
        | big_p(X5)
        | ~ big_p(a)
        | big_p(f(f(X5))) )
      & ( ~ big_p(a)
        | ~ big_p(f(X6))
        | big_p(f(f(X6)))
        | big_p(X5)
        | ~ big_p(a)
        | big_p(f(f(X5))) )
      & ( ~ big_p(a)
        | big_p(X6)
        | big_p(f(f(X6)))
        | ~ big_p(f(X5))
        | ~ big_p(a)
        | big_p(f(f(X5))) )
      & ( ~ big_p(a)
        | ~ big_p(f(X6))
        | big_p(f(f(X6)))
        | ~ big_p(f(X5))
        | ~ big_p(a)
        | big_p(f(f(X5))) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])]),
    [nnpp(pel62),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( big_p(f(f(X1)))
      | big_p(f(f(X2)))
      | ~ big_p(a)
      | ~ big_p(f(X1))
      | ~ big_p(f(X2))
      | ~ big_p(a) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel62),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( big_p(a)
    | big_p(a)
    | big_p(a) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel62),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X1] :
      ( big_p(X1)
      | big_p(f(f(X1)))
      | big_p(X2)
      | big_p(f(f(X2)))
      | ~ big_p(a)
      | ~ big_p(a) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel62),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ( ~ big_p(f(f(esk2_0)))
    | ~ big_p(f(f(esk2_0)))
    | ~ big_p(f(f(esk1_0))) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel62),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1,X2] :
      ( big_p(f(f(X2)))
      | big_p(f(f(X1)))
      | ~ big_p(a)
      | ~ big_p(f(X2))
      | ~ big_p(f(X1)) ),
    inference(cn,[status(thm)],[c_0_3]),
    [nnpp(pel62),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    big_p(a),
    inference(cn,[status(thm)],[c_0_4]),
    [nnpp(pel62),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1,X2] :
      ( big_p(X2)
      | big_p(X1)
      | big_p(f(f(X2)))
      | big_p(f(f(X1)))
      | ~ big_p(a) ),
    inference(cn,[status(thm)],[c_0_5]),
    [nnpp(pel62),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ( ~ big_p(f(f(esk1_0)))
    | ~ big_p(f(f(esk2_0))) ),
    inference(cn,[status(thm)],[c_0_6]),
    [nnpp(pel62),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ! [X1] :
      ( big_p(f(f(X1)))
      | ~ big_p(f(X1)) ),
    inference(condense,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7,c_0_8])])]),
    [nnpp(pel62),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ( big_p(f(esk1_0))
    | ~ big_p(f(f(esk2_0)))
    | ~ big_p(f(f(esk2_0)))
    | ~ big_p(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel62),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ! [X1] :
      ( big_p(f(f(X1)))
      | big_p(X1) ),
    inference(condense,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9,c_0_8])])]),
    [nnpp(pel62),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ( ~ big_p(f(f(esk2_0)))
    | ~ big_p(f(esk1_0)) ),
    inference(spm,[status(thm)],[c_0_10,c_0_11]),
    [nnpp(pel62),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ( big_p(f(esk1_0))
    | ~ big_p(esk1_0)
    | ~ big_p(f(f(esk2_0))) ),
    inference(cn,[status(thm)],[c_0_12]),
    [nnpp(pel62),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ( big_p(esk1_0)
    | ~ big_p(f(f(esk2_0))) ),
    inference(spm,[status(thm)],[c_0_10,c_0_13]),
    [nnpp(pel62),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ~ big_p(f(f(esk2_0))),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14,c_0_15]),c_0_16]),
    [nnpp(pel62),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    ( big_p(f(esk2_0))
    | ~ big_p(esk2_0)
    | ~ big_p(f(f(esk1_0))) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel62),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,negated_conjecture,
    big_p(esk2_0),
    inference(spm,[status(thm)],[c_0_17,c_0_13]),
    [nnpp(pel62),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,negated_conjecture,
    ~ big_p(f(esk2_0)),
    inference(spm,[status(thm)],[c_0_17,c_0_11]),
    [nnpp(pel62),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,negated_conjecture,
    ( big_p(f(esk2_0))
    | big_p(esk1_0)
    | ~ big_p(esk2_0) ),
    inference(spm,[status(thm)],[c_0_18,c_0_13]),
    [nnpp(pel62),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,negated_conjecture,
    ~ big_p(f(f(esk1_0))),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18,c_0_19])]),c_0_20]),
    [nnpp(pel62),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_23,negated_conjecture,
    ( big_p(f(esk2_0))
    | big_p(f(esk1_0))
    | ~ big_p(esk2_0)
    | ~ big_p(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel62),verified_file(c_0_23_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_24,negated_conjecture,
    big_p(esk1_0),
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21,c_0_19])]),c_0_20]),
    [nnpp(pel62),verified_file(c_0_24_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_25,negated_conjecture,
    ~ big_p(f(esk1_0)),
    inference(spm,[status(thm)],[c_0_22,c_0_11]),
    [nnpp(pel62),verified_file(c_0_25_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_26,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23,c_0_19])]),c_0_24])]),c_0_20]),c_0_25]),
    [proof,nnpp(pel62),verified_file(c_0_26_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN084+1/E---3.2.0.THM-CRf.s
