SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel12 has SAT parents the problem 
 RESULT: pel12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel12 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents pel12 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a cth of pel12 
WARNING: c_0_2 has UNS parents c_0_1 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_2 is a thm of c_0_1 
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_3 
SUCCESS: c_0_7 has SAT parents c_0_4 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_7 is a thm of c_0_4 
WARNING: c_0_8 has UNS parents c_0_2 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of c_0_2 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_9 is a thm of c_0_2 
WARNING: c_0_10 has UNS parents c_0_2 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_10 is a thm of c_0_2 
SUCCESS: c_0_11 has SAT parents c_0_5 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_11 is a thm of c_0_5 
SUCCESS: c_0_12 has SAT parents c_0_6 c_0_7 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_6 c_0_7 
WARNING: c_0_13 has UNS parents c_0_2 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_13 is a thm of c_0_2 
SUCCESS: c_0_14 has SAT parents c_0_8 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_14 is a thm of c_0_8 
SUCCESS: c_0_15 has SAT parents c_0_9 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_15 is a thm of c_0_9 
SUCCESS: c_0_16 has SAT parents c_0_10 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_16 is a thm of c_0_10 
SUCCESS: c_0_17 has SAT parents c_0_11 c_0_12 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_17 is a thm of c_0_11 c_0_12 
SUCCESS: c_0_18 has SAT parents c_0_13 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_18 is a thm of c_0_13 
SUCCESS: c_0_19 has SAT parents c_0_14 c_0_15 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_19 is a thm of c_0_14 c_0_15 
SUCCESS: c_0_20 has SAT parents c_0_16 c_0_17 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_20 is a thm of c_0_16 c_0_17 
WARNING: c_0_21 has UNS parents c_0_2 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_21 is a thm of c_0_2 
SUCCESS: c_0_22 has SAT parents c_0_18 c_0_19 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_22 is a thm of c_0_18 c_0_19 
SUCCESS: c_0_23 has SAT parents c_0_12 c_0_19 c_0_20 
 RESULT: c_0_23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_23 is a thm of c_0_12 c_0_19 c_0_20 
SUCCESS: c_0_24 has SAT parents c_0_21 
 RESULT: c_0_24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_24 is a thm of c_0_21 
SUCCESS: c_0_25 has SAT parents c_0_22 c_0_23 
 RESULT: c_0_25_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_25 is a thm of c_0_22 c_0_23 
SUCCESS: c_0_26 has SAT parents c_0_24 c_0_19 
 RESULT: c_0_26_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_26 is a thm of c_0_24 c_0_19 
SUCCESS: c_0_27 has SAT parents c_0_17 c_0_25 
 RESULT: c_0_27_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_27 is a thm of c_0_17 c_0_25 
WARNING: c_0_28 has UNS parents c_0_26 c_0_23 c_0_27 
 RESULT: c_0_28_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_28 is a thm of c_0_26 c_0_23 c_0_27 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 18.92
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN393+1.003/E---3.2.0.THM-CRf.s
fof(pel12,conjecture,
    ( ( ( p1
      <=> p2 )
    <=> p3 )
  <=> ( p1
    <=> ( p2
      <=> p3 ) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.UUHdbspJtw/E---3.1_21514.p',pel12),
    [verified_file(pel12_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( ( p1
        <=> p2 )
      <=> p3 )
    <=> ( p1
      <=> ( p2
        <=> p3 ) ) ),
    inference(assume_negation,[status(cth)],[pel12]),
    [nnpp(pel12),verified_file(c_0_1_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ( ( ~ p2
      | ~ p3
      | ~ p1
      | ~ p1
      | ~ p2
      | ~ p3 )
    & ( p2
      | p3
      | ~ p1
      | ~ p1
      | ~ p2
      | ~ p3 )
    & ( ~ p2
      | p3
      | p1
      | ~ p1
      | ~ p2
      | ~ p3 )
    & ( ~ p3
      | p2
      | p1
      | ~ p1
      | ~ p2
      | ~ p3 )
    & ( ~ p2
      | ~ p3
      | ~ p1
      | p1
      | p2
      | ~ p3 )
    & ( p2
      | p3
      | ~ p1
      | p1
      | p2
      | ~ p3 )
    & ( ~ p2
      | p3
      | p1
      | p1
      | p2
      | ~ p3 )
    & ( ~ p3
      | p2
      | p1
      | p1
      | p2
      | ~ p3 )
    & ( ~ p2
      | ~ p3
      | ~ p1
      | ~ p1
      | p2
      | p3 )
    & ( p2
      | p3
      | ~ p1
      | ~ p1
      | p2
      | p3 )
    & ( ~ p2
      | p3
      | p1
      | ~ p1
      | p2
      | p3 )
    & ( ~ p3
      | p2
      | p1
      | ~ p1
      | p2
      | p3 )
    & ( ~ p2
      | ~ p3
      | ~ p1
      | ~ p2
      | p1
      | p3 )
    & ( p2
      | p3
      | ~ p1
      | ~ p2
      | p1
      | p3 )
    & ( ~ p2
      | p3
      | p1
      | ~ p2
      | p1
      | p3 )
    & ( ~ p3
      | p2
      | p1
      | ~ p2
      | p1
      | p3 )
    & ( ~ p2
      | p3
      | ~ p1
      | ~ p1
      | ~ p2
      | p3 )
    & ( ~ p3
      | p2
      | ~ p1
      | ~ p1
      | ~ p2
      | p3 )
    & ( ~ p2
      | ~ p3
      | p1
      | ~ p1
      | ~ p2
      | p3 )
    & ( p2
      | p3
      | p1
      | ~ p1
      | ~ p2
      | p3 )
    & ( ~ p2
      | p3
      | ~ p1
      | p1
      | p2
      | p3 )
    & ( ~ p3
      | p2
      | ~ p1
      | p1
      | p2
      | p3 )
    & ( ~ p2
      | ~ p3
      | p1
      | p1
      | p2
      | p3 )
    & ( p2
      | p3
      | p1
      | p1
      | p2
      | p3 )
    & ( ~ p2
      | p3
      | ~ p1
      | ~ p1
      | p2
      | ~ p3 )
    & ( ~ p3
      | p2
      | ~ p1
      | ~ p1
      | p2
      | ~ p3 )
    & ( ~ p2
      | ~ p3
      | p1
      | ~ p1
      | p2
      | ~ p3 )
    & ( p2
      | p3
      | p1
      | ~ p1
      | p2
      | ~ p3 )
    & ( ~ p2
      | p3
      | ~ p1
      | ~ p2
      | p1
      | ~ p3 )
    & ( ~ p3
      | p2
      | ~ p1
      | ~ p2
      | p1
      | ~ p3 )
    & ( ~ p2
      | ~ p3
      | p1
      | ~ p2
      | p1
      | ~ p3 )
    & ( p2
      | p3
      | p1
      | ~ p2
      | p1
      | ~ p3 ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])]),
    [nnpp(pel12),verified_file(c_0_2_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ( p2
    | p2
    | ~ p3
    | ~ p1
    | ~ p1
    | ~ p3 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel12),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( ~ p2
    | ~ p3
    | ~ p1
    | ~ p1
    | ~ p2
    | ~ p3 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel12),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ( p3
    | p3
    | ~ p2
    | ~ p1
    | ~ p1
    | ~ p2 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel12),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ( p2
    | ~ p1
    | ~ p3 ),
    inference(cn,[status(thm)],[c_0_3]),
    [nnpp(pel12),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ( ~ p1
    | ~ p2
    | ~ p3 ),
    inference(cn,[status(thm)],[c_0_4]),
    [nnpp(pel12),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ( p2
    | p1
    | p1
    | p2
    | ~ p3
    | ~ p3 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel12),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ( p2
    | p3
    | p1
    | p1
    | p2
    | p3 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel12),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ( p1
    | p1
    | ~ p2
    | ~ p3
    | ~ p2
    | ~ p3 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel12),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ( p3
    | ~ p1
    | ~ p2 ),
    inference(cn,[status(thm)],[c_0_5]),
    [nnpp(pel12),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ( ~ p3
    | ~ p1 ),
    inference(csr,[status(thm)],[c_0_6,c_0_7]),
    [nnpp(pel12),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ( p3
    | p1
    | p1
    | p3
    | ~ p2
    | ~ p2 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel12),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ( p1
    | p2
    | ~ p3 ),
    inference(cn,[status(thm)],[c_0_8]),
    [nnpp(pel12),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ( p1
    | p2
    | p3 ),
    inference(cn,[status(thm)],[c_0_9]),
    [nnpp(pel12),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ( p1
    | ~ p2
    | ~ p3 ),
    inference(cn,[status(thm)],[c_0_10]),
    [nnpp(pel12),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ( ~ p2
    | ~ p1 ),
    inference(csr,[status(thm)],[c_0_11,c_0_12]),
    [nnpp(pel12),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    ( p1
    | p3
    | ~ p2 ),
    inference(cn,[status(thm)],[c_0_13]),
    [nnpp(pel12),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,negated_conjecture,
    ( p2
    | p1 ),
    inference(csr,[status(thm)],[c_0_14,c_0_15]),
    [nnpp(pel12),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,negated_conjecture,
    ( ~ p3
    | ~ p2 ),
    inference(csr,[status(thm)],[c_0_16,c_0_17]),
    [nnpp(pel12),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,negated_conjecture,
    ( p2
    | p3
    | p2
    | p3
    | ~ p1
    | ~ p1 ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel12),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,negated_conjecture,
    ( p3
    | p1 ),
    inference(csr,[status(thm)],[c_0_18,c_0_19]),
    [nnpp(pel12),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_23,negated_conjecture,
    ~ p3,
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12,c_0_19]),c_0_20]),
    [nnpp(pel12),verified_file(c_0_23_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_24,negated_conjecture,
    ( p2
    | p3
    | ~ p1 ),
    inference(cn,[status(thm)],[c_0_21]),
    [nnpp(pel12),verified_file(c_0_24_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_25,negated_conjecture,
    p1,
    inference(sr,[status(thm)],[c_0_22,c_0_23]),
    [nnpp(pel12),verified_file(c_0_25_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_26,negated_conjecture,
    ( p3
    | p2 ),
    inference(csr,[status(thm)],[c_0_24,c_0_19]),
    [nnpp(pel12),verified_file(c_0_26_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_27,negated_conjecture,
    ~ p2,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_17,c_0_25])]),
    [nnpp(pel12),verified_file(c_0_27_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_28,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_26,c_0_23]),c_0_27]),
    [proof,nnpp(pel12),verified_file(c_0_28_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN393+1.003/E---3.2.0.THM-CRf.s
