SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: thm69 has SAT parents the problem 
 RESULT: thm69_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: thm69 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents thm69 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_1 is a cth of thm69 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.07
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_7 is a thm of c_0_2 
SUCCESS: c_0_8 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of c_0_5 c_0_6 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_7 c_0_8 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_10 is a thm of c_0_7 c_0_8 
SUCCESS: c_0_11 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_11 is a thm of c_0_9 c_0_10 
SUCCESS: c_0_12 has SAT parents c_0_5 c_0_11 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_5 c_0_11 
SUCCESS: c_0_13 has SAT parents c_0_3 c_0_11 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_13 is a thm of c_0_3 c_0_11 
WARNING: c_0_14 has UNS parents c_0_12 c_0_13 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_12 c_0_13 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 228.21
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN728+1/E---3.2.0.THM-CRf.s
fof(thm69,conjecture,
    ( ( ! [X1] :
          ( ? [X2] : p(X1,X2)
         => ! [X3] : p(X3,X3) )
      & ! [X4] :
        ? [X5] :
          ( p(X4,X5)
          | ( m(X4)
            & q(f(X4,X5)) ) )
      & ! [X6] :
          ( q(X6)
         => ~ m(g(X6)) ) )
   => ! [X7] :
      ? [X8] :
        ( p(g(X7),X8)
        & p(X7,X7) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.x4qlh7ZaAJ/E---3.1_8580.p',thm69),
    [verified_file(thm69_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( ! [X1] :
            ( ? [X2] : p(X1,X2)
           => ! [X3] : p(X3,X3) )
        & ! [X4] :
          ? [X5] :
            ( p(X4,X5)
            | ( m(X4)
              & q(f(X4,X5)) ) )
        & ! [X6] :
            ( q(X6)
           => ~ m(g(X6)) ) )
     => ! [X7] :
        ? [X8] :
          ( p(g(X7),X8)
          & p(X7,X7) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[thm69])]),
    [nnpp(thm69),verified_file(c_0_1_thm),verified(cth),nnpp(thm69),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X9,X10,X11,X12,X15,X17] :
      ( ( ~ p(X9,X10)
        | p(X11,X11) )
      & ( m(X12)
        | p(X12,esk1_1(X12)) )
      & ( q(f(X12,esk2_1(X12)))
        | p(X12,esk1_1(X12)) )
      & ( ~ q(X15)
        | ~ m(g(X15)) )
      & ( ~ p(g(esk3_0),X17)
        | ~ p(esk3_0,esk3_0) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])])]),
    [nnpp(thm69),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1,X3] :
      ( p(X3,X3)
      | ~ p(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm69),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X1] :
      ( m(X1)
      | p(X1,esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm69),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X1] :
      ( ~ p(g(esk3_0),X1)
      | ~ p(esk3_0,esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm69),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X2,X1] :
      ( m(X1)
      | p(X2,X2) ),
    inference(spm,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(thm69),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1] :
      ( ~ q(X1)
      | ~ m(g(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm69),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X1] : m(X1),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_5,c_0_6]),c_0_6]),
    [nnpp(thm69),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] :
      ( q(f(X1,esk2_1(X1)))
      | p(X1,esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(thm69),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ! [X1] : ~ q(X1),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7,c_0_8])]),
    [nnpp(thm69),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ! [X1] : p(X1,esk1_1(X1)),
    inference(sr,[status(thm)],[c_0_9,c_0_10]),
    [nnpp(thm69),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ~ p(esk3_0,esk3_0),
    inference(spm,[status(thm)],[c_0_5,c_0_11]),
    [nnpp(thm69),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ! [X1] : p(X1,X1),
    inference(spm,[status(thm)],[c_0_3,c_0_11]),
    [nnpp(thm69),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12,c_0_13])]),
    [proof,nnpp(thm69),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN728+1/E---3.2.0.THM-CRf.s
