SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel46 has SAT parents the problem 
 RESULT: pel46_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel46 is a thm of the problem 
SUCCESS: pel46_2 has SAT parents the problem 
 RESULT: pel46_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: pel46_2 is a thm of the problem 
SUCCESS: pel46_3 has SAT parents the problem 
 RESULT: pel46_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel46_3 is a thm of the problem 
SUCCESS: pel46_1 has SAT parents the problem 
 RESULT: pel46_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel46_1 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_4 has SAT parents pel46 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_4 is a cth of pel46 
SUCCESS: c_0_5 has SAT parents pel46_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_5 is a thm of pel46_2 
SUCCESS: c_0_6 has SAT parents pel46_3 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of pel46_3 
SUCCESS: c_0_7 has SAT parents pel46_1 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of pel46_1 
SUCCESS: c_0_8 has SAT parents c_0_4 (Theorem esa)
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.03
FAILURE: c_0_8 fails to be a thm of c_0_4 (Theorem esa)
SUCCESS: c_0_4 has SAT parents c_0_8 (Inverted esa)
 RESULT: c_0_8_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_4 is a thm of c_0_8 (Inverted esa)
FAILURE: c_0_8 fails to be a esa of c_0_4
 ASSUME: c_0_8 is a esa of c_0_4
SUCCESS: c_0_9 has SAT parents c_0_5 (Theorem esa)
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.00
FAILURE: c_0_9 fails to be a thm of c_0_5 (Theorem esa)
SUCCESS: c_0_5 has SAT parents c_0_9 (Inverted esa)
 RESULT: c_0_9_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_9 (Inverted esa)
FAILURE: c_0_9 fails to be a esa of c_0_5
 ASSUME: c_0_9 is a esa of c_0_5
SUCCESS: c_0_10 has SAT parents c_0_6 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_10 is a thm of c_0_6 
SUCCESS: c_0_11 has SAT parents c_0_7 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_11 is a thm of c_0_7 
SUCCESS: c_0_12 has SAT parents c_0_8 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_12 is a thm of c_0_8 
SUCCESS: c_0_13 has SAT parents c_0_8 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_13 is a thm of c_0_8 
SUCCESS: c_0_14 has SAT parents c_0_9 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_9 
SUCCESS: c_0_15 has SAT parents c_0_7 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_15 is a thm of c_0_7 
SUCCESS: c_0_16 has SAT parents c_0_10 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_16 is a thm of c_0_10 
SUCCESS: c_0_17 has SAT parents c_0_11 c_0_12 c_0_13 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_17 is a thm of c_0_11 c_0_12 c_0_13 
SUCCESS: c_0_18 has SAT parents c_0_14 c_0_11 c_0_15 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_18 is a thm of c_0_14 c_0_11 c_0_15 
SUCCESS: c_0_19 has SAT parents c_0_16 c_0_17 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_19 is a thm of c_0_16 c_0_17 
SUCCESS: c_0_20 has SAT parents c_0_18 c_0_17 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_20 is a thm of c_0_18 c_0_17 
SUCCESS: c_0_21 has SAT parents c_0_7 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_21 is a thm of c_0_7 
SUCCESS: c_0_22 has SAT parents c_0_19 c_0_20 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_22 is a thm of c_0_19 c_0_20 
SUCCESS: c_0_23 has SAT parents c_0_21 c_0_17 
 RESULT: c_0_23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_23 is a thm of c_0_21 c_0_17 
SUCCESS: c_0_24 has SAT parents c_0_15 c_0_17 
 RESULT: c_0_24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_24 is a thm of c_0_15 c_0_17 
SUCCESS: c_0_25 has SAT parents c_0_22 c_0_23 c_0_24 
 RESULT: c_0_25_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_25 is a thm of c_0_22 c_0_23 c_0_24 
SUCCESS: c_0_26 has SAT parents c_0_24 c_0_25 
 RESULT: c_0_26_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_26 is a thm of c_0_24 c_0_25 
WARNING: c_0_27 has UNS parents c_0_13 c_0_26 
 RESULT: c_0_27_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_27 is a thm of c_0_13 c_0_26 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 10.51
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN070+1/E---3.2.0.THM-CRf.s
fof(pel46,conjecture,
    ! [X1] :
      ( big_f(X1)
     => big_g(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.FRjZDiCoqr/E---3.1_26674.p',pel46),
    [verified_file(pel46_thm),verified(leaf),verified(lambdapi)] ).
fof(pel46_2,axiom,
    ( ? [X1] :
        ( big_f(X1)
        & ~ big_g(X1) )
   => ? [X3] :
        ( big_f(X3)
        & ~ big_g(X3)
        & ! [X2] :
            ( ( big_f(X2)
              & ~ big_g(X2) )
           => big_j(X3,X2) ) ) ),
    file('/export/starexec/sandbox/tmp/tmp.FRjZDiCoqr/E---3.1_26674.p',pel46_2),
    [verified_file(pel46_2_thm),verified(leaf),verified(lambdapi)] ).
fof(pel46_3,axiom,
    ! [X1,X2] :
      ( ( big_f(X1)
        & big_f(X2)
        & big_h(X1,X2) )
     => ~ big_j(X2,X1) ),
    file('/export/starexec/sandbox/tmp/tmp.FRjZDiCoqr/E---3.1_26674.p',pel46_3),
    [verified_file(pel46_3_thm),verified(leaf),verified(lambdapi)] ).
fof(pel46_1,axiom,
    ! [X1,X2] :
      ( ( big_f(X1)
        & ( ( big_f(X2)
            & big_h(X2,X1) )
         => big_g(X2) ) )
     => big_g(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.FRjZDiCoqr/E---3.1_26674.p',pel46_1),
    [verified_file(pel46_1_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ~ ! [X1] :
        ( big_f(X1)
       => big_g(X1) ),
    inference(assume_negation,[status(cth)],[pel46]),
    [nnpp(pel46),verified_file(c_0_4_thm),verified(cth),nnpp(pel46),verified(lambdapi)] ).
fof(c_0_5,plain,
    ( ? [X1] :
        ( big_f(X1)
        & ~ big_g(X1) )
   => ? [X3] :
        ( big_f(X3)
        & ~ big_g(X3)
        & ! [X2] :
            ( ( big_f(X2)
              & ~ big_g(X2) )
           => big_j(X3,X2) ) ) ),
    inference(fof_simplification,[status(thm)],[pel46_2]),
    [nnpp(pel46),verified_file(c_0_5_thm),verified(thm),nnpp(pel46),verified(lambdapi)] ).
fof(c_0_6,plain,
    ! [X1,X2] :
      ( ( big_f(X1)
        & big_f(X2)
        & big_h(X1,X2) )
     => ~ big_j(X2,X1) ),
    inference(fof_simplification,[status(thm)],[pel46_3]),
    [nnpp(pel46),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,plain,
    ! [X4,X5] :
      ( ( big_f(X5)
        | ~ big_f(X4)
        | big_g(X4) )
      & ( big_h(X5,X4)
        | ~ big_f(X4)
        | big_g(X4) )
      & ( ~ big_g(X5)
        | ~ big_f(X4)
        | big_g(X4) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel46_1])])])]),
    [nnpp(pel46),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ( big_f(esk2_0)
    & ~ big_g(esk2_0) ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])]),
    [nnpp(pel46),verified_file(c_0_8_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_9,plain,
    ! [X6,X8] :
      ( ( big_f(esk1_0)
        | ~ big_f(X6)
        | big_g(X6) )
      & ( ~ big_g(esk1_0)
        | ~ big_f(X6)
        | big_g(X6) )
      & ( ~ big_f(X8)
        | big_g(X8)
        | big_j(esk1_0,X8)
        | ~ big_f(X6)
        | big_g(X6) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])])]),
    [nnpp(pel46),verified_file(c_0_9_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_10,plain,
    ! [X9,X10] :
      ( ~ big_f(X9)
      | ~ big_f(X10)
      | ~ big_h(X9,X10)
      | ~ big_j(X10,X9) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])]),
    [nnpp(pel46),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ! [X2,X1] :
      ( big_f(X1)
      | big_g(X2)
      | ~ big_f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel46),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    big_f(esk2_0),
    inference(split_conjunct,[status(thm)],[c_0_8]),
    [nnpp(pel46),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ~ big_g(esk2_0),
    inference(split_conjunct,[status(thm)],[c_0_8]),
    [nnpp(pel46),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,plain,
    ! [X2,X1] :
      ( big_g(X1)
      | big_j(esk1_0,X1)
      | big_g(X2)
      | ~ big_f(X1)
      | ~ big_f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_9]),
    [nnpp(pel46),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,plain,
    ! [X1,X2] :
      ( big_g(X2)
      | ~ big_g(X1)
      | ~ big_f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel46),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,plain,
    ! [X2,X1] :
      ( ~ big_f(X1)
      | ~ big_f(X2)
      | ~ big_h(X1,X2)
      | ~ big_j(X2,X1) ),
    inference(split_conjunct,[status(thm)],[c_0_10]),
    [nnpp(pel46),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ! [X1] : big_f(X1),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_11,c_0_12]),c_0_13]),
    [nnpp(pel46),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,plain,
    ! [X2,X1] :
      ( big_j(esk1_0,X1)
      | big_g(X2)
      | ~ big_f(X2) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_14,c_0_11]),c_0_15]),
    [nnpp(pel46),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,plain,
    ! [X2,X1] :
      ( ~ big_j(X1,X2)
      | ~ big_h(X2,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16,c_0_17]),c_0_17])]),
    [nnpp(pel46),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,plain,
    ! [X2,X1] :
      ( big_j(esk1_0,X1)
      | big_g(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18,c_0_17])]),
    [nnpp(pel46),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,plain,
    ! [X2,X1] :
      ( big_h(X1,X2)
      | big_g(X2)
      | ~ big_f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel46),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,plain,
    ! [X2,X1] :
      ( big_g(X1)
      | ~ big_h(X2,esk1_0) ),
    inference(spm,[status(thm)],[c_0_19,c_0_20]),
    [nnpp(pel46),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_23,plain,
    ! [X2,X1] :
      ( big_g(X1)
      | big_h(X2,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21,c_0_17])]),
    [nnpp(pel46),verified_file(c_0_23_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_24,plain,
    ! [X2,X1] :
      ( big_g(X1)
      | ~ big_g(X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_15,c_0_17])]),
    [nnpp(pel46),verified_file(c_0_24_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_25,plain,
    big_g(esk1_0),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22,c_0_23]),c_0_24]),
    [nnpp(pel46),verified_file(c_0_25_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_26,plain,
    ! [X1] : big_g(X1),
    inference(spm,[status(thm)],[c_0_24,c_0_25]),
    [nnpp(pel46),verified_file(c_0_26_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_27,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_13,c_0_26])]),
    [proof,nnpp(pel46),verified_file(c_0_27_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN070+1/E---3.2.0.THM-CRf.s
