SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: lx1 has SAT parents the problem 
 RESULT: lx1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: lx1 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents lx1 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_1 is a cth of lx1 
WARNING: c_0_2 has UNS parents c_0_1 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_2 is a thm of c_0_1 
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_2 
SUCCESS: c_0_8 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_5 c_0_6 
SUCCESS: c_0_9 has SAT parents c_0_7 c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_7 c_0_6 
WARNING: c_0_10 has UNS parents c_0_8 c_0_9 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_10 is a thm of c_0_8 c_0_9 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 6.73
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN721+1/E---3.2.0.THM-CRf.s
fof(lx1,conjecture,
    ( ( r(a,b)
      & ! [X1] :
          ( ? [X2] : r(X1,X2)
         => q(X1,X1) )
      & ! [X3,X4] :
          ( q(X3,X4)
         => ! [X5] : r(X5,X4) ) )
   => ? [X6] :
        ( r(b,X6)
        & q(X6,a) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.Mmb4JAASZ0/E---3.1_800.p',lx1),
    [verified_file(lx1_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( r(a,b)
        & ! [X1] :
            ( ? [X2] : r(X1,X2)
           => q(X1,X1) )
        & ! [X3,X4] :
            ( q(X3,X4)
           => ! [X5] : r(X5,X4) ) )
     => ? [X6] :
          ( r(b,X6)
          & q(X6,a) ) ),
    inference(assume_negation,[status(cth)],[lx1]),
    [nnpp(lx1),verified_file(c_0_1_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X7,X8,X9,X10,X11,X12] :
      ( r(a,b)
      & ( ~ r(X7,X8)
        | q(X7,X7) )
      & ( ~ q(X9,X10)
        | r(X11,X10) )
      & ( ~ r(b,X12)
        | ~ q(X12,a) ) ),
    inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])]),
    [nnpp(lx1),verified_file(c_0_2_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( q(X1,X1)
      | ~ r(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(lx1),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    r(a,b),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(lx1),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X1] :
      ( ~ r(b,X1)
      | ~ q(X1,a) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(lx1),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    q(a,a),
    inference(spm,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(lx1),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1,X2,X3] :
      ( r(X3,X2)
      | ~ q(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(lx1),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ~ r(b,a),
    inference(spm,[status(thm)],[c_0_5,c_0_6]),
    [nnpp(lx1),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] : r(X1,a),
    inference(spm,[status(thm)],[c_0_7,c_0_6]),
    [nnpp(lx1),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8,c_0_9])]),
    [proof,nnpp(lx1),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN721+1/E---3.2.0.THM-CRf.s
