SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel56 has SAT parents the problem 
 RESULT: pel56_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel56 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents pel56 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a cth of pel56 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_3 is a thm of c_0_2 
SUCCESS: c_0_4 has SAT parents c_0_3 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.07
SUCCESS: c_0_4 is a thm of c_0_3 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_6 is a thm of c_0_4 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_2 
WARNING: c_0_8 has UNS parents c_0_2 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_8 is a thm of c_0_2 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_5 c_0_6 c_0_7 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_10 is a thm of c_0_5 c_0_6 c_0_7 
SUCCESS: c_0_11 has SAT parents c_0_8 c_0_6 c_0_9 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_11 is a thm of c_0_8 c_0_6 c_0_9 
WARNING: c_0_12 has UNS parents c_0_2 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_2 
SUCCESS: c_0_13 has SAT parents c_0_6 c_0_10 c_0_11 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_13 is a thm of c_0_6 c_0_10 c_0_11 
SUCCESS: c_0_14 has SAT parents c_0_12 c_0_13 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_14 is a thm of c_0_12 c_0_13 
WARNING: c_0_15 has UNS parents c_0_2 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_15 is a thm of c_0_2 
SUCCESS: c_0_16 has SAT parents c_0_14 c_0_6 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_16 is a thm of c_0_14 c_0_6 
WARNING: c_0_17 has UNS parents c_0_15 c_0_13 c_0_16 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_17 is a thm of c_0_15 c_0_13 c_0_16 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 291.25
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN078+1/E---3.2.0.THM-CRf.s
fof(pel56,conjecture,
    ( ! [X1] :
        ( ? [X2] :
            ( big_p(X2)
            & X1 = f(X2) )
       => big_p(X1) )
  <=> ! [X3] :
        ( big_p(X3)
       => big_p(f(X3)) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.H9Mw404N6D/E---3.1_14400.p',pel56),
    [verified_file(pel56_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ! [X1] :
          ( ? [X2] :
              ( big_p(X2)
              & X1 = f(X2) )
         => big_p(X1) )
    <=> ! [X3] :
          ( big_p(X3)
         => big_p(f(X3)) ) ),
    inference(assume_negation,[status(cth)],[pel56]),
    [nnpp(pel56),verified_file(c_0_1_thm),verified(cth),nnpp(pel56),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X7,X8,X9] :
      ( ( big_p(esk3_0)
        | big_p(esk2_0) )
      & ( ~ big_p(f(esk3_0))
        | big_p(esk2_0) )
      & ( big_p(esk3_0)
        | esk1_0 = f(esk2_0) )
      & ( ~ big_p(f(esk3_0))
        | esk1_0 = f(esk2_0) )
      & ( big_p(esk3_0)
        | ~ big_p(esk1_0) )
      & ( ~ big_p(f(esk3_0))
        | ~ big_p(esk1_0) )
      & ( ~ big_p(X8)
        | X7 != f(X8)
        | big_p(X7)
        | ~ big_p(X9)
        | big_p(f(X9)) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])]),
    [nnpp(pel56),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X1,X3,X2] :
      ( big_p(X2)
      | big_p(f(X3))
      | ~ big_p(X1)
      | X2 != f(X1)
      | ~ big_p(X3) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel56),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X2,X1] :
      ( big_p(f(X1))
      | big_p(f(X2))
      | ~ big_p(X2)
      | ~ big_p(X1) ),
    inference(er,[status(thm)],[c_0_3]),
    [nnpp(pel56),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ( esk1_0 = f(esk2_0)
    | ~ big_p(f(esk3_0)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel56),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X1] :
      ( big_p(f(X1))
      | ~ big_p(X1) ),
    inference(ef,[status(thm)],[c_0_4]),
    [nnpp(pel56),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ( big_p(esk3_0)
    | esk1_0 = f(esk2_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel56),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ( big_p(esk2_0)
    | ~ big_p(f(esk3_0)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel56),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ( big_p(esk3_0)
    | big_p(esk2_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel56),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    f(esk2_0) = esk1_0,
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_5,c_0_6]),c_0_7]),
    [nnpp(pel56),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    big_p(esk2_0),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_8,c_0_6]),c_0_9]),
    [nnpp(pel56),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ( ~ big_p(f(esk3_0))
    | ~ big_p(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel56),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    big_p(esk1_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6,c_0_10]),c_0_11])]),
    [nnpp(pel56),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ~ big_p(f(esk3_0)),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12,c_0_13])]),
    [nnpp(pel56),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ( big_p(esk3_0)
    | ~ big_p(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel56),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ~ big_p(esk3_0),
    inference(spm,[status(thm)],[c_0_14,c_0_6]),
    [nnpp(pel56),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_15,c_0_13])]),c_0_16]),
    [proof,nnpp(pel56),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN078+1/E---3.2.0.THM-CRf.s
