SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel37_2 has SAT parents the problem 
 RESULT: pel37_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel37_2 is a thm of the problem 
SUCCESS: pel37 has SAT parents the problem 
 RESULT: pel37_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel37 is a thm of the problem 
SUCCESS: pel37_3 has SAT parents the problem 
 RESULT: pel37_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.07
SUCCESS: pel37_3 is a thm of the problem 
SUCCESS: pel37_1 has SAT parents the problem 
 RESULT: pel37_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: pel37_1 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_4 has SAT parents pel37_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_4 is a thm of pel37_2 
SUCCESS: c_0_5 has SAT parents pel37 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_5 is a cth of pel37 
SUCCESS: c_0_6 has SAT parents pel37_3 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of pel37_3 
SUCCESS: c_0_7 has SAT parents c_0_4 (Theorem esa)
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.05
FAILURE: c_0_7 fails to be a thm of c_0_4 (Theorem esa)
SUCCESS: c_0_4 has SAT parents c_0_7 (Inverted esa)
 RESULT: c_0_7_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_7 (Inverted esa)
FAILURE: c_0_7 fails to be a esa of c_0_4
 ASSUME: c_0_7 is a esa of c_0_4
SUCCESS: c_0_8 has SAT parents c_0_5 (Theorem esa)
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.00
FAILURE: c_0_8 fails to be a thm of c_0_5 (Theorem esa)
SUCCESS: c_0_5 has SAT parents c_0_8 (Inverted esa)
 RESULT: c_0_8_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_8 (Inverted esa)
FAILURE: c_0_8 fails to be a esa of c_0_5
 ASSUME: c_0_8 is a esa of c_0_5
SUCCESS: c_0_9 has SAT parents c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_6 
SUCCESS: c_0_10 has SAT parents c_0_7 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_10 is a thm of c_0_7 
SUCCESS: c_0_11 has SAT parents pel37_1 (Theorem esa)
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.04
FAILURE: c_0_11 fails to be a thm of pel37_1 (Theorem esa)
SUCCESS: pel37_1 has SAT parents c_0_11 (Inverted esa)
 RESULT: c_0_11_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: pel37_1 is a thm of c_0_11 (Inverted esa)
FAILURE: c_0_11 fails to be a esa of pel37_1
 ASSUME: c_0_11 is a esa of pel37_1
SUCCESS: c_0_12 has SAT parents c_0_8 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_12 is a thm of c_0_8 
SUCCESS: c_0_13 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_13 is a thm of c_0_9 c_0_10 
SUCCESS: c_0_14 has SAT parents c_0_11 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_11 
SUCCESS: c_0_15 has SAT parents c_0_12 c_0_13 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_15 is a thm of c_0_12 c_0_13 
SUCCESS: c_0_16 has SAT parents c_0_14 c_0_15 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_16 is a thm of c_0_14 c_0_15 
SUCCESS: c_0_17 has SAT parents c_0_9 c_0_16 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_17 is a thm of c_0_9 c_0_16 
WARNING: c_0_18 has UNS parents c_0_12 c_0_17 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_18 is a thm of c_0_12 c_0_17 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 7.33
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN066+1/E---3.2.0.THM-CRf.s
fof(pel37_2,axiom,
    ! [X3,X1] :
      ( ~ big_p(X3,X1)
     => ? [X4] : big_q(X4,X1) ),
    file('/export/starexec/sandbox/tmp/tmp.kyqTnKc2Ds/E---3.1_6083.p',pel37_2),
    [verified_file(pel37_2_thm),verified(leaf),verified(lambdapi)] ).
fof(pel37,conjecture,
    ! [X3] :
    ? [X4] : big_r(X3,X4),
    file('/export/starexec/sandbox/tmp/tmp.kyqTnKc2Ds/E---3.1_6083.p',pel37),
    [verified_file(pel37_thm),verified(leaf),verified(lambdapi)] ).
fof(pel37_3,axiom,
    ( ? [X3,X4] : big_q(X3,X4)
   => ! [X1] : big_r(X1,X1) ),
    file('/export/starexec/sandbox/tmp/tmp.kyqTnKc2Ds/E---3.1_6083.p',pel37_3),
    [verified_file(pel37_3_thm),verified(leaf),verified(lambdapi)] ).
fof(pel37_1,axiom,
    ! [X1] :
    ? [X2] :
    ! [X3] :
    ? [X4] :
      ( ( big_p(X3,X1)
       => big_p(X4,X2) )
      & big_p(X4,X1)
      & ( big_p(X4,X2)
       => ? [X5] : big_q(X5,X2) ) ),
    file('/export/starexec/sandbox/tmp/tmp.kyqTnKc2Ds/E---3.1_6083.p',pel37_1),
    [verified_file(pel37_1_thm),verified(leaf),nnpp(pel37),verified(lambdapi)] ).
fof(c_0_4,plain,
    ! [X3,X1] :
      ( ~ big_p(X3,X1)
     => ? [X4] : big_q(X4,X1) ),
    inference(fof_simplification,[status(thm)],[pel37_2]),
    [nnpp(pel37),verified_file(c_0_4_thm),verified(thm),nnpp(pel37),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ~ ! [X3] :
      ? [X4] : big_r(X3,X4),
    inference(assume_negation,[status(cth)],[pel37]),
    [nnpp(pel37),verified_file(c_0_5_thm),verified(cth),nnpp(pel37),verified(lambdapi)] ).
fof(c_0_6,plain,
    ! [X8,X9,X10] :
      ( ~ big_q(X8,X9)
      | big_r(X10,X10) ),
    inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel37_3])])])]),
    [nnpp(pel37),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,plain,
    ! [X16,X17] :
      ( big_p(X16,X17)
      | big_q(esk5_2(X16,X17),X17) ),
    inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])]),
    [nnpp(pel37),verified_file(c_0_7_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X7] : ~ big_r(esk1_0,X7),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])]),
    [nnpp(pel37),verified_file(c_0_8_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_9,plain,
    ! [X2,X1,X3] :
      ( big_r(X3,X3)
      | ~ big_q(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel37),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,plain,
    ! [X2,X1] :
      ( big_p(X1,X2)
      | big_q(esk5_2(X1,X2),X2) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel37),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ! [X11,X13] :
      ( ( ~ big_p(X13,X11)
        | big_p(esk3_2(X11,X13),esk2_1(X11)) )
      & big_p(esk3_2(X11,X13),X11)
      & ( ~ big_p(esk3_2(X11,X13),esk2_1(X11))
        | big_q(esk4_2(X11,X13),esk2_1(X11)) ) ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel37_1])])])]),
    [nnpp(pel37),verified_file(c_0_11_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X1] : ~ big_r(esk1_0,X1),
    inference(split_conjunct,[status(thm)],[c_0_8]),
    [nnpp(pel37),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    ! [X3,X2,X1] :
      ( big_r(X1,X1)
      | big_p(X2,X3) ),
    inference(spm,[status(thm)],[c_0_9,c_0_10]),
    [nnpp(pel37),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,plain,
    ! [X2,X1] :
      ( big_q(esk4_2(X1,X2),esk2_1(X1))
      | ~ big_p(esk3_2(X1,X2),esk2_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_11]),
    [nnpp(pel37),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ! [X2,X1] : big_p(X1,X2),
    inference(spm,[status(thm)],[c_0_12,c_0_13]),
    [nnpp(pel37),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,plain,
    ! [X2,X1] : big_q(esk4_2(X1,X2),esk2_1(X1)),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14,c_0_15])]),
    [nnpp(pel37),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,plain,
    ! [X1] : big_r(X1,X1),
    inference(spm,[status(thm)],[c_0_9,c_0_16]),
    [nnpp(pel37),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    $false,
    inference(spm,[status(thm)],[c_0_12,c_0_17]),
    [proof,nnpp(pel37),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN066+1/E---3.2.0.THM-CRf.s
