SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: kalish215 has SAT parents the problem 
 RESULT: kalish215_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: kalish215 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents kalish215 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_1 is a cth of kalish215 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
SUCCESS: c_0_5 has SAT parents c_0_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_5 is a thm of c_0_3 
SUCCESS: c_0_6 has SAT parents c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_6 is a thm of c_0_4 
SUCCESS: c_0_7 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_5 c_0_6 
WARNING: c_0_8 has UNS parents c_0_2 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of c_0_2 
WARNING: c_0_9 has UNS parents c_0_7 c_0_8 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_7 c_0_8 
WARNING: c_0_10 has UNS parents c_0_7 c_0_9 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_7 c_0_9 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 7.35
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN398+1/E---3.2.0.THM-CRf.s
fof(kalish215,conjecture,
    ( ! [X1] :
        ( p
        & f(X1) )
  <=> ( p
      & ! [X2] : f(X2) ) ),
    file('/export/starexec/sandbox/tmp/tmp.JU0bVhn50B/E---3.1_14948.p',kalish215),
    [verified_file(kalish215_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ! [X1] :
          ( p
          & f(X1) )
    <=> ( p
        & ! [X2] : f(X2) ) ),
    inference(assume_negation,[status(cth)],[kalish215]),
    [nnpp(kalish215),verified_file(c_0_1_thm),verified(cth),nnpp(kalish215),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X5,X6] :
      ( ( ~ p
        | ~ f(esk1_0)
        | ~ p
        | ~ f(esk2_0) )
      & ( p
        | p )
      & ( f(X6)
        | p )
      & ( p
        | f(X5) )
      & ( f(X6)
        | f(X5) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])]),
    [nnpp(kalish215),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ( ~ p
    | ~ f(esk1_0)
    | ~ p
    | ~ f(esk2_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish215),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( p
    | p ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish215),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ( ~ p
    | ~ f(esk1_0)
    | ~ f(esk2_0) ),
    inference(cn,[status(thm)],[c_0_3]),
    [nnpp(kalish215),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    p,
    inference(cn,[status(thm)],[c_0_4]),
    [nnpp(kalish215),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ( ~ f(esk1_0)
    | ~ f(esk2_0) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5,c_0_6])]),
    [nnpp(kalish215),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X2,X1] :
      ( f(X1)
      | f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish215),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] : f(X1),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7,c_0_8]),c_0_8]),
    [nnpp(kalish215),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7,c_0_9]),c_0_9])]),
    [proof,nnpp(kalish215),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN398+1/E---3.2.0.THM-CRf.s
