SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel28 has SAT parents the problem 
 RESULT: pel28_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel28 is a thm of the problem 
SUCCESS: pel28_1 has SAT parents the problem 
 RESULT: pel28_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel28_1 is a thm of the problem 
SUCCESS: pel28_2 has SAT parents the problem 
 RESULT: pel28_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: pel28_2 is a thm of the problem 
SUCCESS: pel28_3 has SAT parents the problem 
 RESULT: pel28_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel28_3 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_4 has SAT parents pel28 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_4 is a cth of pel28 
SUCCESS: c_0_5 has SAT parents pel28_1 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_5 is a thm of pel28_1 
SUCCESS: c_0_6 has SAT parents c_0_4 (Theorem esa)
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.01
FAILURE: c_0_6 fails to be a thm of c_0_4 (Theorem esa)
SUCCESS: c_0_4 has SAT parents c_0_6 (Inverted esa)
 RESULT: c_0_6_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_4 is a thm of c_0_6 (Inverted esa)
FAILURE: c_0_6 fails to be a esa of c_0_4
 ASSUME: c_0_6 is a esa of c_0_4
SUCCESS: c_0_7 has SAT parents pel28_2 (Theorem esa)
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.04
FAILURE: c_0_7 fails to be a thm of pel28_2 (Theorem esa)
SUCCESS: pel28_2 has SAT parents c_0_7 (Inverted esa)
 RESULT: c_0_7_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel28_2 is a thm of c_0_7 (Inverted esa)
FAILURE: c_0_7 fails to be a esa of pel28_2
 ASSUME: c_0_7 is a esa of pel28_2
SUCCESS: c_0_8 has SAT parents c_0_5 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of c_0_5 
SUCCESS: c_0_9 has SAT parents c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_6 
SUCCESS: c_0_10 has SAT parents pel28_3 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_10 is a thm of pel28_3 
SUCCESS: c_0_11 has SAT parents c_0_7 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_11 is a thm of c_0_7 
SUCCESS: c_0_12 has SAT parents c_0_8 c_0_9 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_12 is a thm of c_0_8 c_0_9 
SUCCESS: c_0_13 has SAT parents c_0_10 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_13 is a thm of c_0_10 
SUCCESS: c_0_14 has SAT parents c_0_11 c_0_12 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_14 is a thm of c_0_11 c_0_12 
SUCCESS: c_0_15 has SAT parents c_0_6 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_15 is a thm of c_0_6 
SUCCESS: c_0_16 has SAT parents c_0_13 c_0_14 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_16 is a thm of c_0_13 c_0_14 
SUCCESS: c_0_17 has SAT parents c_0_6 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_17 is a thm of c_0_6 
WARNING: c_0_18 has UNS parents c_0_15 c_0_16 c_0_17 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_18 is a thm of c_0_15 c_0_16 c_0_17 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 7.54
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN058+1/E---3.2.0.THM-CRf.s
fof(pel28,conjecture,
    ! [X1] :
      ( ( big_p(X1)
        & big_f(X1) )
     => big_g(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.iTw4SIbh9W/E---3.1_26259.p',pel28),
    [verified_file(pel28_thm),verified(leaf),verified(lambdapi)] ).
fof(pel28_1,axiom,
    ! [X1] :
      ( big_p(X1)
     => ! [X2] : big_q(X2) ),
    file('/export/starexec/sandbox/tmp/tmp.iTw4SIbh9W/E---3.1_26259.p',pel28_1),
    [verified_file(pel28_1_thm),verified(leaf),verified(lambdapi)] ).
fof(pel28_2,axiom,
    ( ! [X1] :
        ( big_q(X1)
        | big_r(X1) )
   => ? [X3] :
        ( big_q(X3)
        & big_s(X3) ) ),
    file('/export/starexec/sandbox/tmp/tmp.iTw4SIbh9W/E---3.1_26259.p',pel28_2),
    [verified_file(pel28_2_thm),verified(leaf),nnpp(pel28),verified(lambdapi)] ).
fof(pel28_3,axiom,
    ( ? [X1] : big_s(X1)
   => ! [X3] :
        ( big_f(X3)
       => big_g(X3) ) ),
    file('/export/starexec/sandbox/tmp/tmp.iTw4SIbh9W/E---3.1_26259.p',pel28_3),
    [verified_file(pel28_3_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ~ ! [X1] :
        ( ( big_p(X1)
          & big_f(X1) )
       => big_g(X1) ),
    inference(assume_negation,[status(cth)],[pel28]),
    [nnpp(pel28),verified_file(c_0_4_thm),verified(cth),nnpp(pel28),verified(lambdapi)] ).
fof(c_0_5,plain,
    ! [X7,X8] :
      ( ~ big_p(X7)
      | big_q(X8) ),
    inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel28_1])])])]),
    [nnpp(pel28),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ( big_p(esk1_0)
    & big_f(esk1_0)
    & ~ big_g(esk1_0) ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])]),
    [nnpp(pel28),verified_file(c_0_6_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_7,plain,
    ( ( big_q(esk3_0)
      | ~ big_q(esk2_0) )
    & ( big_s(esk3_0)
      | ~ big_q(esk2_0) )
    & ( big_q(esk3_0)
      | ~ big_r(esk2_0) )
    & ( big_s(esk3_0)
      | ~ big_r(esk2_0) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel28_2])])])])]),
    [nnpp(pel28),verified_file(c_0_7_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_8,plain,
    ! [X1,X2] :
      ( big_q(X2)
      | ~ big_p(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel28),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    big_p(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel28),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,plain,
    ! [X5,X6] :
      ( ~ big_s(X5)
      | ~ big_f(X6)
      | big_g(X6) ),
    inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel28_3])])])]),
    [nnpp(pel28),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ( big_s(esk3_0)
    | ~ big_q(esk2_0) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel28),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X1] : big_q(X1),
    inference(spm,[status(thm)],[c_0_8,c_0_9]),
    [nnpp(pel28),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    ! [X1,X2] :
      ( big_g(X2)
      | ~ big_s(X1)
      | ~ big_f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_10]),
    [nnpp(pel28),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,plain,
    big_s(esk3_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11,c_0_12])]),
    [nnpp(pel28),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ~ big_g(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel28),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,plain,
    ! [X1] :
      ( big_g(X1)
      | ~ big_f(X1) ),
    inference(spm,[status(thm)],[c_0_13,c_0_14]),
    [nnpp(pel28),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    big_f(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel28),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15,c_0_16]),c_0_17])]),
    [proof,nnpp(pel28),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN058+1/E---3.2.0.THM-CRf.s
