SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel48 has SAT parents the problem 
 RESULT: pel48_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel48 is a thm of the problem 
SUCCESS: pel48_1 has SAT parents the problem 
 RESULT: pel48_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: pel48_1 is a thm of the problem 
SUCCESS: pel48_2 has SAT parents the problem 
 RESULT: pel48_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: pel48_2 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_3 has SAT parents pel48 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_3 is a cth of pel48 
SUCCESS: c_0_4 has SAT parents c_0_3 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_4 is a thm of c_0_3 
SUCCESS: c_0_5 has SAT parents c_0_4 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_4 
SUCCESS: c_0_6 has SAT parents pel48_1 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_6 is a thm of pel48_1 
SUCCESS: c_0_7 has SAT parents pel48_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of pel48_2 
SUCCESS: c_0_8 has SAT parents c_0_4 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_8 is a thm of c_0_4 
SUCCESS: c_0_9 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_5 c_0_6 
SUCCESS: c_0_10 has SAT parents c_0_5 c_0_7 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_5 c_0_7 
SUCCESS: c_0_11 has SAT parents c_0_6 c_0_7 c_0_8 c_0_9 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_11 is a thm of c_0_6 c_0_7 c_0_8 c_0_9 
WARNING: c_0_12 has UNS parents c_0_10 c_0_11 c_0_8 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_10 c_0_11 c_0_8 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 5.00
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN071+1/E---3.2.0.THM-CRf.s
fof(pel48,conjecture,
    ( a = d
    | b = c ),
    file('/export/starexec/sandbox/tmp/tmp.UL6TwsiSOV/E---3.1_6274.p',pel48),
    [verified_file(pel48_thm),verified(leaf),verified(lambdapi)] ).
fof(pel48_1,axiom,
    ( a = b
    | c = d ),
    file('/export/starexec/sandbox/tmp/tmp.UL6TwsiSOV/E---3.1_6274.p',pel48_1),
    [verified_file(pel48_1_thm),verified(leaf),verified(lambdapi)] ).
fof(pel48_2,axiom,
    ( a = c
    | b = d ),
    file('/export/starexec/sandbox/tmp/tmp.UL6TwsiSOV/E---3.1_6274.p',pel48_2),
    [verified_file(pel48_2_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ~ ( a = d
      | b = c ),
    inference(assume_negation,[status(cth)],[pel48]),
    [nnpp(pel48),verified_file(c_0_3_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( a != d
    & b != c ),
    inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])]),
    [nnpp(pel48),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    b != c,
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel48),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,plain,
    ( a = b
    | c = d ),
    inference(split_conjunct,[status(thm)],[pel48_1]),
    [nnpp(pel48),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,plain,
    ( a = c
    | b = d ),
    inference(split_conjunct,[status(thm)],[pel48_2]),
    [nnpp(pel48),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    a != d,
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel48),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ( c = d
    | c != a ),
    inference(spm,[status(thm)],[c_0_5,c_0_6]),
    [nnpp(pel48),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ( c = a
    | c != d ),
    inference(spm,[status(thm)],[c_0_5,c_0_7]),
    [nnpp(pel48),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    c = d,
    inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_6,c_0_7]),c_0_8]),c_0_9]),
    [nnpp(pel48),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10,c_0_11]),c_0_11])]),c_0_8]),
    [proof,nnpp(pel48),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN071+1/E---3.2.0.THM-CRf.s
