SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: prove_this has SAT parents the problem 
 RESULT: prove_this_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: prove_this is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents prove_this 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a cth of prove_this 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_2 
WARNING: c_0_6 has UNS parents c_0_2 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_6 is a thm of c_0_2 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of c_0_2 
SUCCESS: c_0_8 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_3 c_0_4 
SUCCESS: c_0_9 has SAT parents c_0_5 c_0_6 c_0_7 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_5 c_0_6 c_0_7 
WARNING: c_0_10 has UNS parents c_0_2 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_10 is a thm of c_0_2 
SUCCESS: c_0_11 has SAT parents c_0_8 c_0_9 c_0_7 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_11 is a thm of c_0_8 c_0_9 c_0_7 
WARNING: c_0_12 has UNS parents c_0_2 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_12 is a thm of c_0_2 
WARNING: c_0_13 has UNS parents c_0_10 c_0_11 c_0_12 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_13 is a thm of c_0_10 c_0_11 c_0_12 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 291.38
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN943+1/E---3.2.0.THM-CRf.s
fof(prove_this,conjecture,
    ! [X1] :
    ? [X2,X3,X4,X5,X6] :
      ( ( p(X1)
        & e(X1)
        & ( e(X2)
         => ( g(X2)
            | s(X2,f(X2)) ) )
        & ( e(X3)
         => ( g(X3)
            | c(f(X3)) ) )
        & ( s(X1,X6)
         => p(X6) ) )
     => ( ( p(X4)
          & g(X4) )
        | ( p(X5)
          & c(X5) ) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.tSbRNFdFJF/E---3.1_14696.p',prove_this),
    [verified_file(prove_this_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ! [X1] :
      ? [X2,X3,X4,X5,X6] :
        ( ( p(X1)
          & e(X1)
          & ( e(X2)
           => ( g(X2)
              | s(X2,f(X2)) ) )
          & ( e(X3)
           => ( g(X3)
              | c(f(X3)) ) )
          & ( s(X1,X6)
           => p(X6) ) )
       => ( ( p(X4)
            & g(X4) )
          | ( p(X5)
            & c(X5) ) ) ),
    inference(assume_negation,[status(cth)],[prove_this]),
    [nnpp(prove_this),verified_file(c_0_1_thm),verified(cth),nnpp(prove_this),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X8,X9,X10,X11,X12] :
      ( p(esk1_0)
      & e(esk1_0)
      & ( ~ e(X8)
        | g(X8)
        | s(X8,f(X8)) )
      & ( ~ e(X9)
        | g(X9)
        | c(f(X9)) )
      & ( ~ s(esk1_0,X12)
        | p(X12) )
      & ( ~ p(X10)
        | ~ g(X10) )
      & ( ~ p(X11)
        | ~ c(X11) ) ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])]),
    [nnpp(prove_this),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X1] :
      ( ~ p(X1)
      | ~ c(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X1] :
      ( g(X1)
      | c(f(X1))
      | ~ e(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X1] :
      ( p(X1)
      | ~ s(esk1_0,X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X1] :
      ( g(X1)
      | s(X1,f(X1))
      | ~ e(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    e(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X1] :
      ( g(X1)
      | ~ e(X1)
      | ~ p(f(X1)) ),
    inference(spm,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(prove_this),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ( g(esk1_0)
    | p(f(esk1_0)) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_5,c_0_6]),c_0_7])]),
    [nnpp(prove_this),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ! [X1] :
      ( ~ p(X1)
      | ~ g(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    g(esk1_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8,c_0_9]),c_0_7])]),
    [nnpp(prove_this),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    p(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10,c_0_11]),c_0_12])]),
    [proof,nnpp(prove_this),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN943+1/E---3.2.0.THM-CRf.s
