SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel13 has SAT parents the problem 
 RESULT: pel13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: pel13 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents pel13 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_1 is a cth of pel13 
WARNING: c_0_2 has UNS parents c_0_1 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_2 is a thm of c_0_1 
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_2 
WARNING: c_0_6 has UNS parents c_0_2 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_2 
SUCCESS: c_0_7 has SAT parents c_0_3 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_3 
SUCCESS: c_0_8 has SAT parents c_0_4 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_4 
SUCCESS: c_0_9 has SAT parents c_0_5 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_9 is a thm of c_0_5 
SUCCESS: c_0_10 has SAT parents c_0_6 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_6 
SUCCESS: c_0_11 has SAT parents c_0_7 c_0_8 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_11 is a thm of c_0_7 c_0_8 
SUCCESS: c_0_12 has SAT parents c_0_9 c_0_8 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_12 is a thm of c_0_9 c_0_8 
WARNING: c_0_13 has UNS parents c_0_10 c_0_11 c_0_12 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_13 is a thm of c_0_10 c_0_11 c_0_12 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 7.69
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN045+1/E---3.2.0.THM-CRf.s
fof(pel13,conjecture,
    ( ( p
      | ( q
        & r ) )
  <=> ( ( p
        | q )
      & ( p
        | r ) ) ),
    file('/export/starexec/sandbox/tmp/tmp.hhhdQAxhAh/E---3.1_20893.p',pel13),
    [verified_file(pel13_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( p
        | ( q
          & r ) )
    <=> ( ( p
          | q )
        & ( p
          | r ) ) ),
    inference(assume_negation,[status(cth)],[pel13]),
    [nnpp(pel13),verified_file(c_0_1_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ( ( ~ p
      | ~ p
      | ~ p )
    & ( ~ r
      | ~ p
      | ~ p )
    & ( ~ p
      | ~ q
      | ~ p )
    & ( ~ r
      | ~ q
      | ~ p )
    & ( ~ p
      | ~ p
      | ~ q
      | ~ r )
    & ( ~ r
      | ~ p
      | ~ q
      | ~ r )
    & ( ~ p
      | ~ q
      | ~ q
      | ~ r )
    & ( ~ r
      | ~ q
      | ~ q
      | ~ r )
    & ( p
      | q
      | q
      | p )
    & ( p
      | r
      | q
      | p )
    & ( p
      | q
      | r
      | p )
    & ( p
      | r
      | r
      | p ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])]),
    [nnpp(pel13),verified_file(c_0_2_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ( p
    | q
    | q
    | p ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel13),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( ~ p
    | ~ p
    | ~ p ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel13),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ( p
    | r
    | r
    | p ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel13),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ( ~ r
    | ~ q
    | ~ q
    | ~ r ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(pel13),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ( p
    | q ),
    inference(cn,[status(thm)],[c_0_3]),
    [nnpp(pel13),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ~ p,
    inference(cn,[status(thm)],[c_0_4]),
    [nnpp(pel13),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ( p
    | r ),
    inference(cn,[status(thm)],[c_0_5]),
    [nnpp(pel13),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ( ~ q
    | ~ r ),
    inference(cn,[status(thm)],[c_0_6]),
    [nnpp(pel13),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    q,
    inference(sr,[status(thm)],[c_0_7,c_0_8]),
    [nnpp(pel13),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    r,
    inference(sr,[status(thm)],[c_0_9,c_0_8]),
    [nnpp(pel13),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_10,c_0_11])]),c_0_12])]),
    [proof,nnpp(pel13),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN045+1/E---3.2.0.THM-CRf.s
