SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: church_46_12_3 has SAT parents the problem 
 RESULT: church_46_12_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: church_46_12_3 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents church_46_12_3 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_1 is a cth of church_46_12_3 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_2 
WARNING: c_0_6 has UNS parents c_0_2 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_6 is a thm of c_0_2 
SUCCESS: c_0_7 has SAT parents c_0_3 c_0_4 c_0_5 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_3 c_0_4 c_0_5 
WARNING: c_0_8 has UNS parents c_0_2 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of c_0_2 
SUCCESS: c_0_9 has SAT parents c_0_6 c_0_7 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_9 is a thm of c_0_6 c_0_7 
WARNING: c_0_10 has UNS parents c_0_2 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_2 
WARNING: c_0_11 has UNS parents c_0_2 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_11 is a thm of c_0_2 
WARNING: c_0_12 has UNS parents c_0_2 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_2 
WARNING: c_0_13 has UNS parents c_0_2 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_13 is a thm of c_0_2 
SUCCESS: c_0_14 has SAT parents c_0_8 c_0_4 c_0_9 c_0_7 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_8 c_0_4 c_0_9 c_0_7 
SUCCESS: c_0_15 has SAT parents c_0_3 c_0_10 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_15 is a thm of c_0_3 c_0_10 
SUCCESS: c_0_16 has SAT parents c_0_11 c_0_12 c_0_7 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_16 is a thm of c_0_11 c_0_12 c_0_7 
SUCCESS: c_0_17 has SAT parents c_0_13 c_0_10 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_17 is a thm of c_0_13 c_0_10 
SUCCESS: c_0_18 has SAT parents c_0_6 c_0_14 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_18 is a thm of c_0_6 c_0_14 
SUCCESS: c_0_19 has SAT parents c_0_15 c_0_16 c_0_5 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_19 is a thm of c_0_15 c_0_16 c_0_5 
SUCCESS: c_0_20 has SAT parents c_0_17 c_0_18 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_20 is a thm of c_0_17 c_0_18 
SUCCESS: c_0_21 has SAT parents c_0_19 c_0_18 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_21 is a thm of c_0_19 c_0_18 
WARNING: c_0_22 has UNS parents c_0_20 c_0_21 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_22 is a thm of c_0_20 c_0_21 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 389.57
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN328+1/E---3.2.0.THM-CRf.s
fof(church_46_12_3,conjecture,
    ? [X1] :
    ! [X2,X3] :
      ( ( ( big_f(X2)
         => big_g(X2) )
      <=> big_f(X1) )
     => ( ( ( big_f(X2)
           => big_h(X2) )
        <=> big_g(X1) )
       => ( ( ( ( big_f(X2)
               => big_g(X2) )
             => big_h(X2) )
          <=> big_h(X1) )
         => ( big_f(X3)
            & big_g(X3)
            & big_h(X3) ) ) ) ),
    file('/export/starexec/sandbox/tmp/tmp.QCxIQgSzqH/E---3.1_29652.p',church_46_12_3),
    [verified_file(church_46_12_3_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ? [X1] :
      ! [X2,X3] :
        ( ( ( big_f(X2)
           => big_g(X2) )
        <=> big_f(X1) )
       => ( ( ( big_f(X2)
             => big_h(X2) )
          <=> big_g(X1) )
         => ( ( ( ( big_f(X2)
                 => big_g(X2) )
               => big_h(X2) )
            <=> big_h(X1) )
           => ( big_f(X3)
              & big_g(X3)
              & big_h(X3) ) ) ) ),
    inference(assume_negation,[status(cth)],[church_46_12_3]),
    [nnpp(church_46_12_3),verified_file(c_0_1_thm),verified(cth),nnpp(church_46_12_3),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X4] :
      ( ( big_f(esk1_1(X4))
        | big_f(X4) )
      & ( ~ big_g(esk1_1(X4))
        | big_f(X4) )
      & ( ~ big_f(X4)
        | ~ big_f(esk1_1(X4))
        | big_g(esk1_1(X4)) )
      & ( big_f(esk1_1(X4))
        | big_g(X4) )
      & ( ~ big_h(esk1_1(X4))
        | big_g(X4) )
      & ( ~ big_g(X4)
        | ~ big_f(esk1_1(X4))
        | big_h(esk1_1(X4)) )
      & ( ~ big_f(esk1_1(X4))
        | big_g(esk1_1(X4))
        | big_h(X4) )
      & ( ~ big_h(esk1_1(X4))
        | big_h(X4) )
      & ( big_f(esk1_1(X4))
        | big_h(esk1_1(X4))
        | ~ big_h(X4) )
      & ( ~ big_g(esk1_1(X4))
        | big_h(esk1_1(X4))
        | ~ big_h(X4) )
      & ( ~ big_f(esk2_1(X4))
        | ~ big_g(esk3_1(X4))
        | ~ big_h(esk4_1(X4)) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])]),
    [nnpp(church_46_12_3),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X1] :
      ( big_f(X1)
      | ~ big_g(esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X1] :
      ( big_g(esk1_1(X1))
      | big_h(X1)
      | ~ big_f(esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X1] :
      ( big_f(esk1_1(X1))
      | big_f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X1] :
      ( big_h(X1)
      | ~ big_h(esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1] :
      ( big_h(X1)
      | big_f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_3,c_0_4]),c_0_5]),
    [nnpp(church_46_12_3),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X1] :
      ( big_h(esk1_1(X1))
      | ~ big_g(X1)
      | ~ big_f(esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] :
      ( big_h(X1)
      | big_f(esk1_1(X1)) ),
    inference(spm,[status(thm)],[c_0_6,c_0_7]),
    [nnpp(church_46_12_3),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ! [X1] :
      ( big_g(X1)
      | ~ big_h(esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ! [X1] :
      ( big_h(esk1_1(X1))
      | ~ big_g(esk1_1(X1))
      | ~ big_h(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X1] :
      ( big_g(esk1_1(X1))
      | ~ big_f(X1)
      | ~ big_f(esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ! [X1] :
      ( ~ big_f(esk2_1(X1))
      | ~ big_g(esk3_1(X1))
      | ~ big_h(esk4_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_12_3),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ! [X1] :
      ( big_h(esk1_1(esk1_1(X1)))
      | big_h(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_8,c_0_4]),c_0_9]),c_0_7]),
    [nnpp(church_46_12_3),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ! [X1] :
      ( big_f(X1)
      | ~ big_h(esk1_1(esk1_1(X1))) ),
    inference(spm,[status(thm)],[c_0_3,c_0_10]),
    [nnpp(church_46_12_3),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ! [X1] :
      ( big_h(esk1_1(X1))
      | ~ big_h(X1)
      | ~ big_f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11,c_0_12]),c_0_7]),
    [nnpp(church_46_12_3),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ! [X1] :
      ( ~ big_h(esk1_1(esk3_1(X1)))
      | ~ big_h(esk4_1(X1))
      | ~ big_f(esk2_1(X1)) ),
    inference(spm,[status(thm)],[c_0_13,c_0_10]),
    [nnpp(church_46_12_3),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    ! [X1] : big_h(X1),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_6,c_0_14]),c_0_6]),
    [nnpp(church_46_12_3),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,negated_conjecture,
    ! [X1] :
      ( big_f(X1)
      | ~ big_h(esk1_1(X1)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15,c_0_16]),c_0_5]),
    [nnpp(church_46_12_3),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,negated_conjecture,
    ! [X1] : ~ big_f(esk2_1(X1)),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17,c_0_18]),c_0_18])]),
    [nnpp(church_46_12_3),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,negated_conjecture,
    ! [X1] : big_f(X1),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_19,c_0_18])]),
    [nnpp(church_46_12_3),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_20,c_0_21])]),
    [proof,nnpp(church_46_12_3),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN328+1/E---3.2.0.THM-CRf.s
