SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: kalish317 has SAT parents the problem 
 RESULT: kalish317_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: kalish317 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents kalish317 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_1 is a cth of kalish317 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 2.69
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
SUCCESS: c_0_7 has SAT parents c_0_5 c_0_4 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of c_0_5 c_0_4 
SUCCESS: c_0_8 has SAT parents c_0_6 c_0_4 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of c_0_6 c_0_4 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_7 c_0_8 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_7 c_0_8 
SUCCESS: c_0_11 has SAT parents c_0_9 c_0_10 c_0_4 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_11 is a thm of c_0_9 c_0_10 c_0_4 
WARNING: c_0_12 has UNS parents c_0_2 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 7.81
SUCCESS: c_0_12 is a thm of c_0_2 
SUCCESS: c_0_13 has SAT parents c_0_11 c_0_4 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_13 is a thm of c_0_11 c_0_4 
WARNING: c_0_14 has UNS parents c_0_2 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_14 is a thm of c_0_2 
SUCCESS: c_0_15 has SAT parents c_0_12 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_15 is a thm of c_0_12 
SUCCESS: c_0_16 has SAT parents c_0_7 c_0_13 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_16 is a thm of c_0_7 c_0_13 
SUCCESS: c_0_17 has SAT parents c_0_14 c_0_4 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_17 is a thm of c_0_14 c_0_4 
SUCCESS: c_0_18 has SAT parents c_0_15 c_0_16 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_18 is a thm of c_0_15 c_0_16 
SUCCESS: c_0_19 has SAT parents c_0_6 c_0_16 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_19 is a thm of c_0_6 c_0_16 
SUCCESS: c_0_20 has SAT parents c_0_17 c_0_4 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_20 is a thm of c_0_17 c_0_4 
SUCCESS: c_0_21 has SAT parents c_0_18 c_0_19 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_21 is a thm of c_0_18 c_0_19 
WARNING: c_0_22 has UNS parents c_0_2 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_22 is a thm of c_0_2 
SUCCESS: c_0_23 has SAT parents c_0_7 c_0_20 
 RESULT: c_0_23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_23 is a thm of c_0_7 c_0_20 
SUCCESS: c_0_24 has SAT parents c_0_9 c_0_21 c_0_13 
 RESULT: c_0_24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_24 is a thm of c_0_9 c_0_21 c_0_13 
SUCCESS: c_0_25 has SAT parents c_0_22 c_0_23 c_0_4 
 RESULT: c_0_25_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_25 is a thm of c_0_22 c_0_23 c_0_4 
SUCCESS: c_0_26 has SAT parents c_0_24 c_0_16 
 RESULT: c_0_26_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_26 is a thm of c_0_24 c_0_16 
SUCCESS: c_0_27 has SAT parents c_0_12 
 RESULT: c_0_27_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_27 is a thm of c_0_12 
SUCCESS: c_0_28 has SAT parents c_0_12 
 RESULT: c_0_28_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_28 is a thm of c_0_12 
SUCCESS: c_0_29 has SAT parents c_0_25 c_0_4 
 RESULT: c_0_29_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_29 is a thm of c_0_25 c_0_4 
WARNING: c_0_30 has UNS parents c_0_2 
 RESULT: c_0_30_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_30 is a thm of c_0_2 
SUCCESS: c_0_31 has SAT parents c_0_18 c_0_26 
 RESULT: c_0_31_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_31 is a thm of c_0_18 c_0_26 
SUCCESS: c_0_32 has SAT parents c_0_27 c_0_13 c_0_28 
 RESULT: c_0_32_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_32 is a thm of c_0_27 c_0_13 c_0_28 
SUCCESS: c_0_33 has SAT parents c_0_7 c_0_29 
 RESULT: c_0_33_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_33 is a thm of c_0_7 c_0_29 
SUCCESS: c_0_34 has SAT parents c_0_5 c_0_13 
 RESULT: c_0_34_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_34 is a thm of c_0_5 c_0_13 
SUCCESS: c_0_35 has SAT parents c_0_30 c_0_31 c_0_32 
 RESULT: c_0_35_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_35 is a thm of c_0_30 c_0_31 c_0_32 
SUCCESS: c_0_36 has SAT parents c_0_6 c_0_33 
 RESULT: c_0_36_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_36 is a thm of c_0_6 c_0_33 
SUCCESS: c_0_37 has SAT parents c_0_34 c_0_29 
 RESULT: c_0_37_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_37 is a thm of c_0_34 c_0_29 
SUCCESS: c_0_38 has SAT parents c_0_35 c_0_36 
 RESULT: c_0_38_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_38 is a thm of c_0_35 c_0_36 
SUCCESS: c_0_39 has SAT parents c_0_6 c_0_37 
 RESULT: c_0_39_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_39 is a thm of c_0_6 c_0_37 
SUCCESS: c_0_40 has SAT parents c_0_9 c_0_38 c_0_13 
 RESULT: c_0_40_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_40 is a thm of c_0_9 c_0_38 c_0_13 
SUCCESS: c_0_41 has SAT parents c_0_35 c_0_39 
 RESULT: c_0_41_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_41 is a thm of c_0_35 c_0_39 
SUCCESS: c_0_42 has SAT parents c_0_15 c_0_33 
 RESULT: c_0_42_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_42 is a thm of c_0_15 c_0_33 
SUCCESS: c_0_43 has SAT parents c_0_40 c_0_33 
 RESULT: c_0_43_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_43 is a thm of c_0_40 c_0_33 
SUCCESS: c_0_44 has SAT parents c_0_9 c_0_41 c_0_13 
 RESULT: c_0_44_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_44 is a thm of c_0_9 c_0_41 c_0_13 
SUCCESS: c_0_45 has SAT parents c_0_42 c_0_43 
 RESULT: c_0_45_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_45 is a thm of c_0_42 c_0_43 
SUCCESS: c_0_46 has SAT parents c_0_44 c_0_37 
 RESULT: c_0_46_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_46 is a thm of c_0_44 c_0_37 
SUCCESS: c_0_47 has SAT parents c_0_31 c_0_45 
 RESULT: c_0_47_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_47 is a thm of c_0_31 c_0_45 
SUCCESS: c_0_48 has SAT parents c_0_15 c_0_46 c_0_47 
 RESULT: c_0_48_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_48 is a thm of c_0_15 c_0_46 c_0_47 
SUCCESS: c_0_49 has SAT parents c_0_48 
 RESULT: c_0_49_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_49 is a thm of c_0_48 
SUCCESS: c_0_50 has SAT parents c_0_15 c_0_37 
 RESULT: c_0_50_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_50 is a thm of c_0_15 c_0_37 
SUCCESS: c_0_51 has SAT parents c_0_17 c_0_46 c_0_49 
 RESULT: c_0_51_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_51 is a thm of c_0_17 c_0_46 c_0_49 
SUCCESS: c_0_52 has SAT parents c_0_50 c_0_51 
 RESULT: c_0_52_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_52 is a thm of c_0_50 c_0_51 
SUCCESS: c_0_53 has SAT parents c_0_22 c_0_52 c_0_46 c_0_47 c_0_49 
 RESULT: c_0_53_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_53 is a thm of c_0_22 c_0_52 c_0_46 c_0_47 c_0_49 
WARNING: c_0_54 has UNS parents c_0_2 
 RESULT: c_0_54_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_54 is a thm of c_0_2 
SUCCESS: c_0_55 has SAT parents c_0_53 c_0_37 
 RESULT: c_0_55_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_55 is a thm of c_0_53 c_0_37 
SUCCESS: c_0_56 has SAT parents c_0_54 c_0_55 
 RESULT: c_0_56_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_56 is a thm of c_0_54 c_0_55 
SUCCESS: c_0_57 has SAT parents c_0_56 c_0_4 
 RESULT: c_0_57_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_57 is a thm of c_0_56 c_0_4 
SUCCESS: c_0_58 has SAT parents c_0_15 c_0_26 c_0_31 
 RESULT: c_0_58_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_58 is a thm of c_0_15 c_0_26 c_0_31 
WARNING: c_0_59 has UNS parents c_0_2 
 RESULT: c_0_59_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_59 is a thm of c_0_2 
SUCCESS: c_0_60 has SAT parents c_0_30 c_0_57 
 RESULT: c_0_60_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_60 is a thm of c_0_30 c_0_57 
SUCCESS: c_0_61 has SAT parents c_0_58 
 RESULT: c_0_61_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_61 is a thm of c_0_58 
SUCCESS: c_0_62 has SAT parents c_0_59 c_0_55 
 RESULT: c_0_62_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_62 is a thm of c_0_59 c_0_55 
SUCCESS: c_0_63 has SAT parents c_0_60 c_0_16 c_0_61 
 RESULT: c_0_63_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_63 is a thm of c_0_60 c_0_16 c_0_61 
SUCCESS: c_0_64 has SAT parents c_0_34 c_0_57 
 RESULT: c_0_64_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_64 is a thm of c_0_34 c_0_57 
SUCCESS: c_0_65 has SAT parents c_0_62 c_0_63 c_0_16 
 RESULT: c_0_65_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_65 is a thm of c_0_62 c_0_63 c_0_16 
SUCCESS: c_0_66 has SAT parents c_0_64 c_0_13 
 RESULT: c_0_66_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_66 is a thm of c_0_64 c_0_13 
SUCCESS: c_0_67 has SAT parents c_0_27 c_0_4 c_0_28 
 RESULT: c_0_67_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_67 is a thm of c_0_27 c_0_4 c_0_28 
SUCCESS: c_0_68 has SAT parents c_0_65 c_0_16 
 RESULT: c_0_68_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_68 is a thm of c_0_65 c_0_16 
SUCCESS: c_0_69 has SAT parents c_0_62 c_0_66 c_0_13 
 RESULT: c_0_69_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_69 is a thm of c_0_62 c_0_66 c_0_13 
SUCCESS: c_0_70 has SAT parents c_0_67 c_0_57 
 RESULT: c_0_70_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_70 is a thm of c_0_67 c_0_57 
SUCCESS: c_0_71 has SAT parents c_0_60 c_0_68 
 RESULT: c_0_71_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_71 is a thm of c_0_60 c_0_68 
SUCCESS: c_0_72 has SAT parents c_0_69 c_0_13 
 RESULT: c_0_72_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_72 is a thm of c_0_69 c_0_13 
SUCCESS: c_0_73 has SAT parents c_0_62 c_0_70 c_0_67 
 RESULT: c_0_73_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_73 is a thm of c_0_62 c_0_70 c_0_67 
SUCCESS: c_0_74 has SAT parents c_0_71 c_0_72 
 RESULT: c_0_74_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_74 is a thm of c_0_71 c_0_72 
SUCCESS: c_0_75 has SAT parents c_0_70 c_0_68 
 RESULT: c_0_75_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_75 is a thm of c_0_70 c_0_68 
SUCCESS: c_0_76 has SAT parents c_0_73 c_0_68 
 RESULT: c_0_76_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_76 is a thm of c_0_73 c_0_68 
SUCCESS: c_0_77 has SAT parents c_0_62 c_0_74 c_0_72 
 RESULT: c_0_77_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_77 is a thm of c_0_62 c_0_74 c_0_72 
SUCCESS: c_0_78 has SAT parents c_0_57 c_0_75 c_0_76 
 RESULT: c_0_78_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_78 is a thm of c_0_57 c_0_75 c_0_76 
SUCCESS: c_0_79 has SAT parents c_0_28 c_0_72 
 RESULT: c_0_79_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_79 is a thm of c_0_28 c_0_72 
SUCCESS: c_0_80 has SAT parents c_0_77 c_0_72 
 RESULT: c_0_80_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_80 is a thm of c_0_77 c_0_72 
SUCCESS: c_0_81 has SAT parents c_0_78 c_0_13 
 RESULT: c_0_81_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_81 is a thm of c_0_78 c_0_13 
WARNING: c_0_82 has UNS parents c_0_79 c_0_80 c_0_81 
 RESULT: c_0_82_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_82 is a thm of c_0_79 c_0_80 c_0_81 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 426.36
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN415+1/E---3.2.0.THM-CRf.s
fof(kalish317,conjecture,
    ( ( ? [X1] : f(X1)
      & ! [X2,X3] :
          ( ( f(X2)
            & f(X3) )
         => X2 = X3 ) )
  <=> ? [X4] :
        ( f(X4)
        & ! [X5] :
            ( f(X5)
           => X4 = X5 ) ) ),
    file('/export/starexec/sandbox/tmp/tmp.MqmzV6dsMt/E---3.1_30609.p',kalish317),
    [verified_file(kalish317_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( ? [X1] : f(X1)
        & ! [X2,X3] :
            ( ( f(X2)
              & f(X3) )
           => X2 = X3 ) )
    <=> ? [X4] :
          ( f(X4)
          & ! [X5] :
              ( f(X5)
             => X4 = X5 ) ) ),
    inference(assume_negation,[status(cth)],[kalish317]),
    [nnpp(kalish317),verified_file(c_0_1_thm),verified(cth),nnpp(kalish317),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X6,X9,X12,X13,X15] :
      ( ( f(esk3_1(X9))
        | ~ f(X9)
        | f(esk1_0)
        | ~ f(X6) )
      & ( X9 != esk3_1(X9)
        | ~ f(X9)
        | f(esk1_0)
        | ~ f(X6) )
      & ( f(esk3_1(X9))
        | ~ f(X9)
        | f(esk2_0)
        | ~ f(X6) )
      & ( X9 != esk3_1(X9)
        | ~ f(X9)
        | f(esk2_0)
        | ~ f(X6) )
      & ( f(esk3_1(X9))
        | ~ f(X9)
        | esk1_0 != esk2_0
        | ~ f(X6) )
      & ( X9 != esk3_1(X9)
        | ~ f(X9)
        | esk1_0 != esk2_0
        | ~ f(X6) )
      & ( f(esk5_0)
        | f(esk4_0) )
      & ( ~ f(X15)
        | esk5_0 = X15
        | f(esk4_0) )
      & ( f(esk5_0)
        | ~ f(X12)
        | ~ f(X13)
        | X12 = X13 )
      & ( ~ f(X15)
        | esk5_0 = X15
        | ~ f(X12)
        | ~ f(X13)
        | X12 = X13 ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])]),
    [nnpp(kalish317),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( f(esk3_1(X1))
      | f(esk1_0)
      | ~ f(X1)
      | ~ f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ( f(esk5_0)
    | f(esk4_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X1] :
      ( f(esk5_0)
      | X1 = X2
      | ~ f(X1)
      | ~ f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X1] :
      ( f(esk3_1(X1))
      | f(esk1_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_3,c_0_4]),c_0_3]),
    [nnpp(kalish317),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1] :
      ( X1 = esk4_0
      | f(esk5_0)
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_5,c_0_4]),
    [nnpp(kalish317),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ( f(esk3_1(esk4_0))
    | f(esk5_0)
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_6,c_0_4]),
    [nnpp(kalish317),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X2,X1] :
      ( f(esk1_0)
      | X1 != esk3_1(X1)
      | ~ f(X1)
      | ~ f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ( esk3_1(esk4_0) = esk4_0
    | f(esk1_0)
    | f(esk5_0) ),
    inference(spm,[status(thm)],[c_0_7,c_0_8]),
    [nnpp(kalish317),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ! [X1] :
      ( f(esk5_0)
      | f(esk1_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9,c_0_10]),c_0_4]),
    [nnpp(kalish317),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X3,X2,X1] :
      ( esk5_0 = X1
      | X2 = X3
      | ~ f(X1)
      | ~ f(X2)
      | ~ f(X3) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ( f(esk1_0)
    | f(esk5_0) ),
    inference(spm,[status(thm)],[c_0_11,c_0_4]),
    [nnpp(kalish317),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ! [X2,X1] :
      ( f(esk3_1(X1))
      | f(esk2_0)
      | ~ f(X1)
      | ~ f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,negated_conjecture,
    ! [X2,X1] :
      ( X1 = X2
      | esk5_0 != X2
      | ~ f(X2)
      | ~ f(X1) ),
    inference(ef,[status(thm)],[c_0_12]),
    [nnpp(kalish317),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    ( esk4_0 = esk1_0
    | f(esk5_0) ),
    inference(spm,[status(thm)],[c_0_7,c_0_13]),
    [nnpp(kalish317),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    ! [X1] :
      ( f(esk3_1(X1))
      | f(esk2_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14,c_0_4]),c_0_14]),
    [nnpp(kalish317),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,negated_conjecture,
    ! [X1] :
      ( esk4_0 = esk1_0
      | X1 = esk5_0
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_15,c_0_16]),
    [nnpp(kalish317),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,negated_conjecture,
    ( esk4_0 = esk1_0
    | f(esk3_1(esk5_0))
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_6,c_0_16]),
    [nnpp(kalish317),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,negated_conjecture,
    ( f(esk3_1(esk4_0))
    | f(esk5_0)
    | f(esk2_0) ),
    inference(spm,[status(thm)],[c_0_17,c_0_4]),
    [nnpp(kalish317),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,negated_conjecture,
    ( esk3_1(esk5_0) = esk5_0
    | esk4_0 = esk1_0
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_18,c_0_19]),
    [nnpp(kalish317),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,negated_conjecture,
    ! [X2,X1] :
      ( f(esk2_0)
      | X1 != esk3_1(X1)
      | ~ f(X1)
      | ~ f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_23,negated_conjecture,
    ( esk3_1(esk4_0) = esk4_0
    | f(esk2_0)
    | f(esk5_0) ),
    inference(spm,[status(thm)],[c_0_7,c_0_20]),
    [nnpp(kalish317),verified_file(c_0_23_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_24,negated_conjecture,
    ! [X1] :
      ( esk4_0 = esk1_0
      | f(esk1_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9,c_0_21]),c_0_13]),
    [nnpp(kalish317),verified_file(c_0_24_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_25,negated_conjecture,
    ! [X1] :
      ( f(esk5_0)
      | f(esk2_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22,c_0_23]),c_0_4]),
    [nnpp(kalish317),verified_file(c_0_25_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_26,negated_conjecture,
    ( esk4_0 = esk1_0
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_24,c_0_16]),
    [nnpp(kalish317),verified_file(c_0_26_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_27,negated_conjecture,
    ! [X2,X1] :
      ( esk5_0 = X1
      | X2 != esk5_0
      | ~ f(X2)
      | ~ f(X1) ),
    inference(ef,[status(thm)],[c_0_12]),
    [nnpp(kalish317),verified_file(c_0_27_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_28,negated_conjecture,
    ! [X1] :
      ( esk5_0 = X1
      | ~ f(esk5_0)
      | ~ f(X1) ),
    inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_12])]),
    [nnpp(kalish317),verified_file(c_0_28_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_29,negated_conjecture,
    ( f(esk2_0)
    | f(esk5_0) ),
    inference(spm,[status(thm)],[c_0_25,c_0_4]),
    [nnpp(kalish317),verified_file(c_0_29_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_30,negated_conjecture,
    ! [X1] :
      ( esk5_0 = X1
      | f(esk4_0)
      | ~ f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_30_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_31,negated_conjecture,
    ( esk1_0 = esk5_0
    | esk4_0 = esk1_0 ),
    inference(spm,[status(thm)],[c_0_18,c_0_26]),
    [nnpp(kalish317),verified_file(c_0_31_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_32,negated_conjecture,
    ! [X1] :
      ( esk5_0 = X1
      | esk1_0 != esk5_0
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27,c_0_13]),c_0_28]),
    [nnpp(kalish317),verified_file(c_0_32_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_33,negated_conjecture,
    ( esk4_0 = esk2_0
    | f(esk5_0) ),
    inference(spm,[status(thm)],[c_0_7,c_0_29]),
    [nnpp(kalish317),verified_file(c_0_33_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_34,negated_conjecture,
    ! [X1] :
      ( X1 = esk1_0
      | f(esk5_0)
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_5,c_0_13]),
    [nnpp(kalish317),verified_file(c_0_34_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_35,negated_conjecture,
    ! [X1] :
      ( esk5_0 = X1
      | f(esk1_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30,c_0_31]),c_0_32]),
    [nnpp(kalish317),verified_file(c_0_35_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_36,negated_conjecture,
    ( esk4_0 = esk2_0
    | f(esk3_1(esk5_0))
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_6,c_0_33]),
    [nnpp(kalish317),verified_file(c_0_36_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_37,negated_conjecture,
    ( esk2_0 = esk1_0
    | f(esk5_0) ),
    inference(spm,[status(thm)],[c_0_34,c_0_29]),
    [nnpp(kalish317),verified_file(c_0_37_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_38,negated_conjecture,
    ( esk3_1(esk5_0) = esk5_0
    | esk4_0 = esk2_0
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_35,c_0_36]),
    [nnpp(kalish317),verified_file(c_0_38_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_39,negated_conjecture,
    ( esk2_0 = esk1_0
    | f(esk3_1(esk5_0))
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_6,c_0_37]),
    [nnpp(kalish317),verified_file(c_0_39_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_40,negated_conjecture,
    ! [X1] :
      ( esk4_0 = esk2_0
      | f(esk1_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9,c_0_38]),c_0_13]),
    [nnpp(kalish317),verified_file(c_0_40_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_41,negated_conjecture,
    ( esk3_1(esk5_0) = esk5_0
    | esk2_0 = esk1_0
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_35,c_0_39]),
    [nnpp(kalish317),verified_file(c_0_41_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_42,negated_conjecture,
    ! [X1] :
      ( esk4_0 = esk2_0
      | X1 = esk5_0
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_15,c_0_33]),
    [nnpp(kalish317),verified_file(c_0_42_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_43,negated_conjecture,
    ( esk4_0 = esk2_0
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_40,c_0_33]),
    [nnpp(kalish317),verified_file(c_0_43_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_44,negated_conjecture,
    ! [X1] :
      ( esk2_0 = esk1_0
      | f(esk1_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9,c_0_41]),c_0_13]),
    [nnpp(kalish317),verified_file(c_0_44_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_45,negated_conjecture,
    ( esk1_0 = esk5_0
    | esk4_0 = esk2_0 ),
    inference(spm,[status(thm)],[c_0_42,c_0_43]),
    [nnpp(kalish317),verified_file(c_0_45_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_46,negated_conjecture,
    ( esk2_0 = esk1_0
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_44,c_0_37]),
    [nnpp(kalish317),verified_file(c_0_46_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_47,negated_conjecture,
    ( esk2_0 = esk1_0
    | esk1_0 = esk5_0 ),
    inference(spm,[status(thm)],[c_0_31,c_0_45]),
    [nnpp(kalish317),verified_file(c_0_47_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_48,negated_conjecture,
    ! [X1] :
      ( esk2_0 = esk1_0
      | X1 = esk1_0
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15,c_0_46]),c_0_47]),
    [nnpp(kalish317),verified_file(c_0_48_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_49,negated_conjecture,
    ( esk2_0 = esk1_0
    | ~ f(esk2_0) ),
    inference(ef,[status(thm)],[c_0_48]),
    [nnpp(kalish317),verified_file(c_0_49_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_50,negated_conjecture,
    ! [X1] :
      ( esk2_0 = esk1_0
      | X1 = esk5_0
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_15,c_0_37]),
    [nnpp(kalish317),verified_file(c_0_50_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_51,negated_conjecture,
    ( esk2_0 = esk1_0
    | f(esk3_1(esk1_0)) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17,c_0_46]),c_0_49]),
    [nnpp(kalish317),verified_file(c_0_51_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_52,negated_conjecture,
    ( esk3_1(esk1_0) = esk5_0
    | esk2_0 = esk1_0 ),
    inference(spm,[status(thm)],[c_0_50,c_0_51]),
    [nnpp(kalish317),verified_file(c_0_52_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_53,negated_conjecture,
    ! [X1] :
      ( esk2_0 = esk1_0
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22,c_0_52]),c_0_46]),c_0_47]),c_0_49]),
    [nnpp(kalish317),verified_file(c_0_53_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_54,negated_conjecture,
    ! [X2,X1] :
      ( f(esk3_1(X1))
      | ~ f(X1)
      | esk1_0 != esk2_0
      | ~ f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_54_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_55,negated_conjecture,
    esk2_0 = esk1_0,
    inference(spm,[status(thm)],[c_0_53,c_0_37]),
    [nnpp(kalish317),verified_file(c_0_55_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_56,negated_conjecture,
    ! [X2,X1] :
      ( f(esk3_1(X1))
      | ~ f(X2)
      | ~ f(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54,c_0_55])]),
    [nnpp(kalish317),verified_file(c_0_56_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_57,negated_conjecture,
    ! [X1] :
      ( f(esk3_1(X1))
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56,c_0_4]),c_0_56]),
    [nnpp(kalish317),verified_file(c_0_57_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_58,negated_conjecture,
    ! [X1] :
      ( esk4_0 = esk1_0
      | X1 = esk1_0
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15,c_0_26]),c_0_31]),
    [nnpp(kalish317),verified_file(c_0_58_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_59,negated_conjecture,
    ! [X2,X1] :
      ( X1 != esk3_1(X1)
      | ~ f(X1)
      | esk1_0 != esk2_0
      | ~ f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(kalish317),verified_file(c_0_59_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_60,negated_conjecture,
    ! [X1] :
      ( esk3_1(X1) = esk5_0
      | f(esk4_0)
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_30,c_0_57]),
    [nnpp(kalish317),verified_file(c_0_60_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_61,negated_conjecture,
    ( esk4_0 = esk1_0
    | ~ f(esk4_0) ),
    inference(ef,[status(thm)],[c_0_58]),
    [nnpp(kalish317),verified_file(c_0_61_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_62,negated_conjecture,
    ! [X2,X1] :
      ( esk3_1(X1) != X1
      | ~ f(X2)
      | ~ f(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59,c_0_55])]),
    [nnpp(kalish317),verified_file(c_0_62_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_63,negated_conjecture,
    ( esk3_1(esk5_0) = esk5_0
    | esk4_0 = esk1_0 ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60,c_0_16]),c_0_61]),
    [nnpp(kalish317),verified_file(c_0_63_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_64,negated_conjecture,
    ! [X1] :
      ( esk3_1(X1) = esk1_0
      | f(esk5_0)
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_34,c_0_57]),
    [nnpp(kalish317),verified_file(c_0_64_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_65,negated_conjecture,
    ! [X1] :
      ( esk4_0 = esk1_0
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62,c_0_63]),c_0_16]),
    [nnpp(kalish317),verified_file(c_0_65_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_66,negated_conjecture,
    ( esk3_1(esk1_0) = esk1_0
    | f(esk5_0) ),
    inference(spm,[status(thm)],[c_0_64,c_0_13]),
    [nnpp(kalish317),verified_file(c_0_66_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_67,negated_conjecture,
    ! [X1] :
      ( esk5_0 = X1
      | esk4_0 != esk5_0
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27,c_0_4]),c_0_28]),
    [nnpp(kalish317),verified_file(c_0_67_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_68,negated_conjecture,
    esk4_0 = esk1_0,
    inference(spm,[status(thm)],[c_0_65,c_0_16]),
    [nnpp(kalish317),verified_file(c_0_68_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_69,negated_conjecture,
    ! [X1] :
      ( f(esk5_0)
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62,c_0_66]),c_0_13]),
    [nnpp(kalish317),verified_file(c_0_69_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_70,negated_conjecture,
    ! [X1] :
      ( esk3_1(X1) = esk5_0
      | esk4_0 != esk5_0
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_67,c_0_57]),
    [nnpp(kalish317),verified_file(c_0_70_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_71,negated_conjecture,
    ! [X1] :
      ( esk3_1(X1) = esk5_0
      | f(esk1_0)
      | ~ f(X1) ),
    inference(rw,[status(thm)],[c_0_60,c_0_68]),
    [nnpp(kalish317),verified_file(c_0_71_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_72,negated_conjecture,
    f(esk5_0),
    inference(spm,[status(thm)],[c_0_69,c_0_13]),
    [nnpp(kalish317),verified_file(c_0_72_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_73,negated_conjecture,
    ! [X2,X1] :
      ( esk4_0 != esk5_0
      | ~ f(X1)
      | ~ f(X2) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62,c_0_70]),c_0_67]),
    [nnpp(kalish317),verified_file(c_0_73_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_74,negated_conjecture,
    ( esk3_1(esk5_0) = esk5_0
    | f(esk1_0) ),
    inference(spm,[status(thm)],[c_0_71,c_0_72]),
    [nnpp(kalish317),verified_file(c_0_74_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_75,negated_conjecture,
    ! [X1] :
      ( esk3_1(X1) = esk5_0
      | esk1_0 != esk5_0
      | ~ f(X1) ),
    inference(rw,[status(thm)],[c_0_70,c_0_68]),
    [nnpp(kalish317),verified_file(c_0_75_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_76,negated_conjecture,
    ! [X2,X1] :
      ( esk1_0 != esk5_0
      | ~ f(X1)
      | ~ f(X2) ),
    inference(rw,[status(thm)],[c_0_73,c_0_68]),
    [nnpp(kalish317),verified_file(c_0_76_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_77,negated_conjecture,
    ! [X1] :
      ( f(esk1_0)
      | ~ f(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62,c_0_74]),c_0_72])]),
    [nnpp(kalish317),verified_file(c_0_77_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_78,negated_conjecture,
    ! [X1] :
      ( esk1_0 != esk5_0
      | ~ f(X1) ),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57,c_0_75]),c_0_76]),
    [nnpp(kalish317),verified_file(c_0_78_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_79,negated_conjecture,
    ! [X1] :
      ( esk5_0 = X1
      | ~ f(X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_28,c_0_72])]),
    [nnpp(kalish317),verified_file(c_0_79_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_80,negated_conjecture,
    f(esk1_0),
    inference(spm,[status(thm)],[c_0_77,c_0_72]),
    [nnpp(kalish317),verified_file(c_0_80_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_81,negated_conjecture,
    esk1_0 != esk5_0,
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78,c_0_13]),c_0_78]),
    [nnpp(kalish317),verified_file(c_0_81_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_82,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79,c_0_80]),c_0_81]),
    [proof,nnpp(kalish317),verified_file(c_0_82_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN415+1/E---3.2.0.THM-CRf.s
