SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: church_46_18_2 has SAT parents the problem 
 RESULT: church_46_18_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: church_46_18_2 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents church_46_18_2 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_1 is a cth of church_46_18_2 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.65
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 7.26
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
SUCCESS: c_0_5 has SAT parents c_0_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_3 
SUCCESS: c_0_6 has SAT parents c_0_4 c_0_5 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_6 is a thm of c_0_4 c_0_5 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of c_0_2 
SUCCESS: c_0_8 has SAT parents c_0_6 c_0_7 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_8 is a thm of c_0_6 c_0_7 
WARNING: c_0_9 has UNS parents c_0_5 c_0_8 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_9 is a thm of c_0_5 c_0_8 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 201.93
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN350+1/E---3.2.0.THM-CRf.s
fof(church_46_18_2,conjecture,
    ! [X1] :
    ? [X2,X3] :
    ! [X4] :
      ( ( big_f(X1,X4)
      <=> big_f(X4,X1) )
     => ( big_f(X1,X4)
      <=> ( big_f(X3,X4)
          & ( big_f(X2,X4)
           => big_f(X2,X3) ) ) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.SuHhiFypwY/E---3.1_11350.p',church_46_18_2),
    [verified_file(church_46_18_2_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ! [X1] :
      ? [X2,X3] :
      ! [X4] :
        ( ( big_f(X1,X4)
        <=> big_f(X4,X1) )
       => ( big_f(X1,X4)
        <=> ( big_f(X3,X4)
            & ( big_f(X2,X4)
             => big_f(X2,X3) ) ) ) ),
    inference(assume_negation,[status(cth)],[church_46_18_2]),
    [nnpp(church_46_18_2),verified_file(c_0_1_thm),verified(cth),nnpp(church_46_18_2),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X6,X7] :
      ( ( ~ big_f(esk1_0,esk2_2(X6,X7))
        | big_f(esk2_2(X6,X7),esk1_0) )
      & ( ~ big_f(esk2_2(X6,X7),esk1_0)
        | big_f(esk1_0,esk2_2(X6,X7)) )
      & ( big_f(X6,esk2_2(X6,X7))
        | ~ big_f(X7,esk2_2(X6,X7))
        | ~ big_f(esk1_0,esk2_2(X6,X7)) )
      & ( ~ big_f(X6,X7)
        | ~ big_f(X7,esk2_2(X6,X7))
        | ~ big_f(esk1_0,esk2_2(X6,X7)) )
      & ( big_f(X7,esk2_2(X6,X7))
        | big_f(esk1_0,esk2_2(X6,X7)) )
      & ( ~ big_f(X6,esk2_2(X6,X7))
        | big_f(X6,X7)
        | big_f(esk1_0,esk2_2(X6,X7)) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])]),
    [nnpp(church_46_18_2),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1] :
      ( big_f(X1,esk2_2(X2,X1))
      | big_f(esk1_0,esk2_2(X2,X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_18_2),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X2,X1] :
      ( ~ big_f(X1,X2)
      | ~ big_f(X2,esk2_2(X1,X2))
      | ~ big_f(esk1_0,esk2_2(X1,X2)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_18_2),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X1] : big_f(esk1_0,esk2_2(X1,esk1_0)),
    inference(ef,[status(thm)],[c_0_3]),
    [nnpp(church_46_18_2),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X1] : ~ big_f(X1,esk1_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_4,c_0_5]),c_0_5])]),
    [nnpp(church_46_18_2),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X2,X1] :
      ( big_f(esk2_2(X1,X2),esk1_0)
      | ~ big_f(esk1_0,esk2_2(X1,X2)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(church_46_18_2),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X2,X1] : ~ big_f(esk1_0,esk2_2(X1,X2)),
    inference(spm,[status(thm)],[c_0_6,c_0_7]),
    [nnpp(church_46_18_2),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[c_0_5,c_0_8]),
    [proof,nnpp(church_46_18_2),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN350+1/E---3.2.0.THM-CRf.s
