SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: x2138 has SAT parents the problem 
 RESULT: x2138_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: x2138 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents x2138 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a cth of x2138 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Timeout - CPU = 30.10
FAILURE: c_0_1 fails to be a thm of c_0_2 (Inverted esa)
FAILURE: c_0_2 fails to be a esa of c_0_1
 ASSUME: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
SUCCESS: c_0_7 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_5 c_0_6 
WARNING: c_0_8 has UNS parents c_0_2 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_8 is a thm of c_0_2 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_7 c_0_8 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_7 c_0_8 
SUCCESS: c_0_11 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_11 is a thm of c_0_9 c_0_10 
WARNING: c_0_12 has UNS parents c_0_2 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_12 is a thm of c_0_2 
WARNING: c_0_13 has UNS parents c_0_11 c_0_12 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.07
SUCCESS: c_0_13 is a thm of c_0_11 c_0_12 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 290.89
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN386+1/E---3.2.0.THM-CRf.s
fof(x2138,conjecture,
    ( ( ! [X1] :
        ? [X2] : big_f(X1,X2)
      & ? [X1] :
        ! [X3] :
        ? [X4] :
        ! [X5] :
          ( big_s(X4,X5)
         => big_d(X5,X1,X3) )
      & ! [X3] :
        ? [X6] :
        ! [X7,X8] :
          ( big_d(X7,X8,X6)
         => ! [X2,X9] :
              ( ( big_f(X7,X2)
                & big_f(X8,X9) )
             => big_d(X2,X9,X3) ) ) )
   => ? [X2] :
      ! [X3] :
      ? [X10] :
      ! [X5] :
        ( big_s(X10,X5)
       => ! [X9] :
            ( big_f(X5,X9)
           => big_d(X9,X2,X3) ) ) ),
    file('/export/starexec/sandbox/tmp/tmp.ssPU0rriDm/E---3.1_10417.p',x2138),
    [verified_file(x2138_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( ! [X1] :
          ? [X2] : big_f(X1,X2)
        & ? [X1] :
          ! [X3] :
          ? [X4] :
          ! [X5] :
            ( big_s(X4,X5)
           => big_d(X5,X1,X3) )
        & ! [X3] :
          ? [X6] :
          ! [X7,X8] :
            ( big_d(X7,X8,X6)
           => ! [X2,X9] :
                ( ( big_f(X7,X2)
                  & big_f(X8,X9) )
               => big_d(X2,X9,X3) ) ) )
     => ? [X2] :
        ! [X3] :
        ? [X10] :
        ! [X5] :
          ( big_s(X10,X5)
         => ! [X9] :
              ( big_f(X5,X9)
             => big_d(X9,X2,X3) ) ) ),
    inference(assume_negation,[status(cth)],[x2138]),
    [nnpp(x2138),verified_file(c_0_1_thm),verified(cth),nnpp(x2138),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X11,X14,X16,X17,X19,X20,X21,X22,X23,X25] :
      ( big_f(X11,esk1_1(X11))
      & ( ~ big_s(esk3_1(X14),X16)
        | big_d(X16,esk2_0,X14) )
      & ( ~ big_d(X19,X20,esk4_1(X17))
        | ~ big_f(X19,X21)
        | ~ big_f(X20,X22)
        | big_d(X21,X22,X17) )
      & big_s(X25,esk6_2(X23,X25))
      & big_f(esk6_2(X23,X25),esk7_2(X23,X25))
      & ~ big_d(esk7_2(X23,X25),X23,esk5_1(X23)) ),
    inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])]),
    [nnpp(x2138),verified_file(c_0_2_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X1,X2] :
      ( big_d(X2,esk2_0,X1)
      | ~ big_s(esk3_1(X1),X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2138),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X2,X1] : big_s(X1,esk6_2(X2,X1)),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2138),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2,X1,X3,X5,X4] :
      ( big_d(X4,X5,X3)
      | ~ big_d(X1,X2,esk4_1(X3))
      | ~ big_f(X1,X4)
      | ~ big_f(X2,X5) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2138),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X2,X1] : big_d(esk6_2(X1,esk3_1(X2)),esk2_0,X2),
    inference(spm,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(x2138),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X4,X3,X2,X1] :
      ( big_d(X1,X2,X3)
      | ~ big_f(esk6_2(X4,esk3_1(esk4_1(X3))),X1)
      | ~ big_f(esk2_0,X2) ),
    inference(spm,[status(thm)],[c_0_5,c_0_6]),
    [nnpp(x2138),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X2,X1] : big_f(esk6_2(X1,X2),esk7_2(X1,X2)),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2138),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X2,X1] : ~ big_d(esk7_2(X1,X2),X1,esk5_1(X1)),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2138),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ! [X3,X2,X1] :
      ( big_d(esk7_2(X1,esk3_1(esk4_1(X2))),X3,X2)
      | ~ big_f(esk2_0,X3) ),
    inference(spm,[status(thm)],[c_0_7,c_0_8]),
    [nnpp(x2138),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ! [X1] : ~ big_f(esk2_0,X1),
    inference(spm,[status(thm)],[c_0_9,c_0_10]),
    [nnpp(x2138),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X1] : big_f(X1,esk1_1(X1)),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2138),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    $false,
    inference(spm,[status(thm)],[c_0_11,c_0_12]),
    [proof,nnpp(x2138),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN386+1/E---3.2.0.THM-CRf.s
