SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: c_0_1 is a predicate definition of epred1_0
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: thm400 has SAT parents the problem 
 RESULT: thm400_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: thm400 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_2 has SAT parents c_0_1 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_2 is a thm of c_0_1 
SUCCESS: c_0_3 has SAT parents thm400 c_0_1 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_3 is a cth of thm400 c_0_1 
SUCCESS: c_0_4 has SAT parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_2 
SUCCESS: c_0_5 has SAT parents c_0_3 (Theorem esa)
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.04
FAILURE: c_0_5 fails to be a thm of c_0_3 (Theorem esa)
SUCCESS: c_0_3 has SAT parents c_0_5 (Inverted esa)
 RESULT: c_0_5_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_3 is a thm of c_0_5 (Inverted esa)
FAILURE: c_0_5 fails to be a esa of c_0_3
 ASSUME: c_0_5 is a esa of c_0_3
SUCCESS: c_0_6 has SAT parents c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_6 is a thm of c_0_4 
SUCCESS: c_0_7 has SAT parents c_0_5 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_5 
SUCCESS: c_0_8 has SAT parents c_0_4 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_4 
SUCCESS: c_0_9 has SAT parents c_0_6 c_0_7 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_6 c_0_7 
SUCCESS: c_0_10 has SAT parents c_0_8 c_0_7 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_10 is a thm of c_0_8 c_0_7 
SUCCESS: c_0_11 has SAT parents c_0_4 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_11 is a thm of c_0_4 
SUCCESS: c_0_12 has SAT parents c_0_5 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_12 is a thm of c_0_5 
SUCCESS: c_0_13 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_13 is a thm of c_0_9 c_0_10 
SUCCESS: c_0_14 has SAT parents c_0_4 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_14 is a thm of c_0_4 
SUCCESS: c_0_15 has SAT parents c_0_11 c_0_7 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_15 is a thm of c_0_11 c_0_7 
SUCCESS: c_0_16 has SAT parents c_0_12 c_0_10 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_16 is a thm of c_0_12 c_0_10 
SUCCESS: c_0_17 has SAT parents c_0_13 c_0_10 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_17 is a thm of c_0_13 c_0_10 
SUCCESS: c_0_18 has SAT parents c_0_14 c_0_7 
 RESULT: c_0_18_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_18 is a thm of c_0_14 c_0_7 
SUCCESS: c_0_19 has SAT parents c_0_15 c_0_16 
 RESULT: c_0_19_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_19 is a thm of c_0_15 c_0_16 
SUCCESS: c_0_20 has SAT parents c_0_12 c_0_17 
 RESULT: c_0_20_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_20 is a thm of c_0_12 c_0_17 
SUCCESS: c_0_21 has SAT parents c_0_18 c_0_19 
 RESULT: c_0_21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_21 is a thm of c_0_18 c_0_19 
SUCCESS: c_0_22 has SAT parents c_0_15 c_0_20 c_0_21 
 RESULT: c_0_22_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_22 is a thm of c_0_15 c_0_20 c_0_21 
SUCCESS: c_0_23 has SAT parents c_0_15 c_0_22 
 RESULT: c_0_23_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_23 is a thm of c_0_15 c_0_22 
SUCCESS: c_0_24 has SAT parents c_0_5 
 RESULT: c_0_24_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_24 is a thm of c_0_5 
SUCCESS: c_0_25 has SAT parents c_0_18 c_0_23 c_0_22 
 RESULT: c_0_25_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_25 is a thm of c_0_18 c_0_23 c_0_22 
WARNING: c_0_26 has UNS parents c_0_24 c_0_25 
 RESULT: c_0_26_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_26 is a thm of c_0_24 c_0_25 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 12.88
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN726+1/E---3.2.0.THM-CRf.s
fof(thm400,conjecture,
    ( ( ( ! [X1,X2,X3] :
            ( ( p(X1,X2)
              & p(X2,X3) )
           => p(X1,X3) )
        & ! [X1,X2,X3] :
            ( ( q(X1,X2)
              & q(X2,X3) )
           => q(X1,X3) )
        & ! [X1,X2] :
            ( q(X1,X2)
           => q(X2,X1) )
        & ! [X1,X2] :
            ( p(X1,X2)
            | q(X1,X2) ) )
     => ! [X1,X2] : p(X1,X2) )
    | ! [X1,X2] : q(X1,X2) ),
    file('/export/starexec/sandbox2/tmp/tmp.gGttvgHDAF/E---3.1_26801.p',thm400),
    [verified_file(thm400_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,plain,
    ( epred1_0
  <=> ( ! [X1,X2,X3] :
          ( ( p(X1,X2)
            & p(X2,X3) )
         => p(X1,X3) )
      & ! [X1,X2,X3] :
          ( ( q(X1,X2)
            & q(X2,X3) )
         => q(X1,X3) )
      & ! [X1,X2] :
          ( q(X1,X2)
         => q(X2,X1) )
      & ! [X1,X2] :
          ( p(X1,X2)
          | q(X1,X2) ) ) ),
    introduced(definition),
    [verified(introduced_leaf),verified(lambdapi)] ).
fof(c_0_2,plain,
    ( epred1_0
   => ( ! [X1,X2,X3] :
          ( ( p(X1,X2)
            & p(X2,X3) )
         => p(X1,X3) )
      & ! [X1,X2,X3] :
          ( ( q(X1,X2)
            & q(X2,X3) )
         => q(X1,X3) )
      & ! [X1,X2] :
          ( q(X1,X2)
         => q(X2,X1) )
      & ! [X1,X2] :
          ( p(X1,X2)
          | q(X1,X2) ) ) ),
    inference(split_equiv,[status(thm)],[c_0_1]),
    [nnpp(thm400),verified_file(c_0_2_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ~ ( ( epred1_0
       => ! [X1,X2] : p(X1,X2) )
      | ! [X1,X2] : q(X1,X2) ),
    inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[thm400]),c_0_1]),
    [nnpp(thm400),verified_file(c_0_3_thm),verified(cth),nnpp(thm400),verified(lambdapi)] ).
fof(c_0_4,plain,
    ! [X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] :
      ( ( ~ p(X8,X9)
        | ~ p(X9,X10)
        | p(X8,X10)
        | ~ epred1_0 )
      & ( ~ q(X11,X12)
        | ~ q(X12,X13)
        | q(X11,X13)
        | ~ epred1_0 )
      & ( ~ q(X14,X15)
        | q(X15,X14)
        | ~ epred1_0 )
      & ( p(X16,X17)
        | q(X16,X17)
        | ~ epred1_0 ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])]),
    [nnpp(thm400),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ( epred1_0
    & ~ p(esk1_0,esk2_0)
    & ~ q(esk3_0,esk4_0) ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])]),
    [nnpp(thm400),verified_file(c_0_5_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_6,plain,
    ! [X2,X3,X1] :
      ( p(X1,X3)
      | ~ p(X1,X2)
      | ~ p(X2,X3)
      | ~ epred1_0 ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(thm400),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    epred1_0,
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(thm400),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,plain,
    ! [X2,X1] :
      ( p(X1,X2)
      | q(X1,X2)
      | ~ epred1_0 ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(thm400),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,plain,
    ! [X3,X2,X1] :
      ( p(X1,X2)
      | ~ p(X3,X2)
      | ~ p(X1,X3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_6,c_0_7])]),
    [nnpp(thm400),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,plain,
    ! [X2,X1] :
      ( q(X1,X2)
      | p(X1,X2) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8,c_0_7])]),
    [nnpp(thm400),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ! [X1,X2] :
      ( q(X2,X1)
      | ~ q(X1,X2)
      | ~ epred1_0 ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(thm400),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ~ p(esk1_0,esk2_0),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(thm400),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    ! [X3,X2,X1] :
      ( q(X1,X2)
      | p(X3,X2)
      | ~ p(X3,X1) ),
    inference(spm,[status(thm)],[c_0_9,c_0_10]),
    [nnpp(thm400),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,plain,
    ! [X2,X3,X1] :
      ( q(X1,X3)
      | ~ q(X1,X2)
      | ~ q(X2,X3)
      | ~ epred1_0 ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(thm400),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,plain,
    ! [X2,X1] :
      ( q(X1,X2)
      | ~ q(X2,X1) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11,c_0_7])]),
    [nnpp(thm400),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    q(esk1_0,esk2_0),
    inference(spm,[status(thm)],[c_0_12,c_0_10]),
    [nnpp(thm400),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,plain,
    ! [X3,X2,X1] :
      ( q(X1,X2)
      | q(X2,X3)
      | p(X1,X3) ),
    inference(spm,[status(thm)],[c_0_13,c_0_10]),
    [nnpp(thm400),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_18,plain,
    ! [X3,X2,X1] :
      ( q(X1,X2)
      | ~ q(X3,X2)
      | ~ q(X1,X3) ),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14,c_0_7])]),
    [nnpp(thm400),verified_file(c_0_18_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_19,plain,
    q(esk2_0,esk1_0),
    inference(spm,[status(thm)],[c_0_15,c_0_16]),
    [nnpp(thm400),verified_file(c_0_19_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_20,negated_conjecture,
    ! [X1] :
      ( q(X1,esk2_0)
      | q(esk1_0,X1) ),
    inference(spm,[status(thm)],[c_0_12,c_0_17]),
    [nnpp(thm400),verified_file(c_0_20_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_21,plain,
    ! [X1] :
      ( q(X1,esk1_0)
      | ~ q(X1,esk2_0) ),
    inference(spm,[status(thm)],[c_0_18,c_0_19]),
    [nnpp(thm400),verified_file(c_0_21_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_22,plain,
    ! [X1] : q(X1,esk1_0),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15,c_0_20]),c_0_21]),
    [nnpp(thm400),verified_file(c_0_22_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_23,plain,
    ! [X1] : q(esk1_0,X1),
    inference(spm,[status(thm)],[c_0_15,c_0_22]),
    [nnpp(thm400),verified_file(c_0_23_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_24,negated_conjecture,
    ~ q(esk3_0,esk4_0),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(thm400),verified_file(c_0_24_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_25,plain,
    ! [X2,X1] : q(X1,X2),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18,c_0_23]),c_0_22])]),
    [nnpp(thm400),verified_file(c_0_25_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_26,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24,c_0_25])]),
    [proof,nnpp(thm400),verified_file(c_0_26_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN726+1/E---3.2.0.THM-CRf.s
