SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: prove_this has SAT parents the problem 
 RESULT: prove_this_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: prove_this is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents prove_this 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a cth of prove_this 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_4 is a thm of c_0_2 
SUCCESS: c_0_5 has SAT parents c_0_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_5 is a thm of c_0_3 
WARNING: c_0_6 has UNS parents c_0_2 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_6 is a thm of c_0_2 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of c_0_2 
WARNING: c_0_8 has UNS parents c_0_2 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.08
SUCCESS: c_0_8 is a thm of c_0_2 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_4 c_0_5 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_10 is a thm of c_0_4 c_0_5 
SUCCESS: c_0_11 has SAT parents c_0_6 c_0_7 c_0_8 c_0_9 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.06
SUCCESS: c_0_11 is a thm of c_0_6 c_0_7 c_0_8 c_0_9 
WARNING: c_0_12 has UNS parents c_0_2 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_12 is a thm of c_0_2 
SUCCESS: c_0_13 has SAT parents c_0_10 c_0_5 c_0_11 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_13 is a thm of c_0_10 c_0_5 c_0_11 
WARNING: c_0_14 has UNS parents c_0_12 c_0_13 c_0_5 c_0_11 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_14 is a thm of c_0_12 c_0_13 c_0_5 c_0_11 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 12.81
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN918+1/E---3.2.0.THM-CRf.s
fof(prove_this,conjecture,
    ( ( ! [X1] :
          ( ( ( f(X1)
              & g(X1) )
           => h(X1) )
         => ? [X2] :
              ( f(X2)
              & ~ g(X2) ) )
      & ( ! [X3] :
            ( f(X3)
           => g(X3) )
        | ! [X4] :
            ( f(X4)
           => h(X4) ) ) )
   => ( ! [X5] :
          ( ( f(X5)
            & h(X5) )
         => g(X5) )
     => ? [X6] :
          ( f(X6)
          & g(X6)
          & ~ h(X6) ) ) ),
    file('/export/starexec/sandbox2/tmp/tmp.tT3zs9fSdv/E---3.1_14597.p',prove_this),
    [verified_file(prove_this_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( ! [X1] :
            ( ( ( f(X1)
                & g(X1) )
             => h(X1) )
           => ? [X2] :
                ( f(X2)
                & ~ g(X2) ) )
        & ( ! [X3] :
              ( f(X3)
             => g(X3) )
          | ! [X4] :
              ( f(X4)
             => h(X4) ) ) )
     => ( ! [X5] :
            ( ( f(X5)
              & h(X5) )
           => g(X5) )
       => ? [X6] :
            ( f(X6)
            & g(X6)
            & ~ h(X6) ) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[prove_this])]),
    [nnpp(prove_this),verified_file(c_0_1_thm),verified(cth),nnpp(prove_this),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X7,X8,X9,X11,X12,X13,X14] :
      ( ( f(esk1_0)
        | f(X7) )
      & ( ~ g(esk1_0)
        | f(X7) )
      & ( f(esk1_0)
        | g(X8) )
      & ( ~ g(esk1_0)
        | g(X8) )
      & ( f(esk1_0)
        | ~ h(X9) )
      & ( ~ g(esk1_0)
        | ~ h(X9) )
      & ( ~ f(X11)
        | g(X11)
        | ~ f(X12)
        | h(X12) )
      & ( ~ f(X13)
        | ~ h(X13)
        | g(X13) )
      & ( ~ f(X14)
        | ~ g(X14)
        | h(X14) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])])]),
    [nnpp(prove_this),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X1] :
      ( f(esk1_0)
      | f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X2,X1] :
      ( g(X1)
      | h(X2)
      | ~ f(X1)
      | ~ f(X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    f(esk1_0),
    inference(ef,[status(thm)],[c_0_3]),
    [nnpp(prove_this),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X1] :
      ( ~ g(esk1_0)
      | ~ h(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1] :
      ( h(X1)
      | ~ f(X1)
      | ~ g(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X1] :
      ( f(X1)
      | ~ g(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] :
      ( g(X1)
      | ~ g(esk1_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ! [X1] :
      ( h(esk1_0)
      | g(X1)
      | ~ f(X1) ),
    inference(spm,[status(thm)],[c_0_4,c_0_5]),
    [nnpp(prove_this),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ~ g(esk1_0),
    inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_6,c_0_7]),c_0_8]),c_0_9]),
    [nnpp(prove_this),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ! [X1] :
      ( g(X1)
      | ~ f(X1)
      | ~ h(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(prove_this),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    h(esk1_0),
    inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_10,c_0_5]),c_0_11]),
    [nnpp(prove_this),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    $false,
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12,c_0_13]),c_0_5])]),c_0_11]),
    [proof,nnpp(prove_this),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN918+1/E---3.2.0.THM-CRf.s
