SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel21 has SAT parents the problem 
 RESULT: pel21_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel21 is a thm of the problem 
SUCCESS: pel21_2 has SAT parents the problem 
 RESULT: pel21_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel21_2 is a thm of the problem 
SUCCESS: pel21_1 has SAT parents the problem 
 RESULT: pel21_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel21_1 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_3 has SAT parents pel21 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_3 is a cth of pel21 
SUCCESS: c_0_4 has SAT parents pel21_2 (Theorem esa)
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.00
FAILURE: c_0_4 fails to be a thm of pel21_2 (Theorem esa)
SUCCESS: pel21_2 has SAT parents c_0_4 (Inverted esa)
 RESULT: c_0_4_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel21_2 is a thm of c_0_4 (Inverted esa)
FAILURE: c_0_4 fails to be a esa of pel21_2
 ASSUME: c_0_4 is a esa of pel21_2
SUCCESS: c_0_5 has SAT parents c_0_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of c_0_3 
SUCCESS: c_0_6 has SAT parents c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_6 is a thm of c_0_4 
SUCCESS: c_0_7 has SAT parents c_0_5 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_5 
SUCCESS: c_0_8 has SAT parents c_0_6 c_0_7 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_6 c_0_7 
SUCCESS: c_0_9 has SAT parents c_0_5 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_5 
SUCCESS: c_0_10 has SAT parents pel21_1 (Theorem esa)
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.00
FAILURE: c_0_10 fails to be a thm of pel21_1 (Theorem esa)
SUCCESS: pel21_1 has SAT parents c_0_10 (Inverted esa)
 RESULT: c_0_10_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel21_1 is a thm of c_0_10 (Inverted esa)
FAILURE: c_0_10 fails to be a esa of pel21_1
 ASSUME: c_0_10 is a esa of pel21_1
SUCCESS: c_0_11 has SAT parents c_0_8 c_0_9 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_11 is a thm of c_0_8 c_0_9 
SUCCESS: c_0_12 has SAT parents c_0_10 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_12 is a thm of c_0_10 
SUCCESS: c_0_13 has SAT parents c_0_7 c_0_11 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_13 is a thm of c_0_7 c_0_11 
WARNING: c_0_14 has UNS parents c_0_12 c_0_11 c_0_13 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.02
SUCCESS: c_0_14 is a thm of c_0_12 c_0_11 c_0_13 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 6.42
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN051+1/E---3.2.0.THM-CRf.s
fof(pel21,conjecture,
    ? [X1] :
      ( p
    <=> big_f(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.INeNIo2OKu/E---3.1_5478.p',pel21),
    [verified_file(pel21_thm),verified(leaf),verified(lambdapi)] ).
fof(pel21_2,axiom,
    ? [X1] :
      ( big_f(X1)
     => p ),
    file('/export/starexec/sandbox/tmp/tmp.INeNIo2OKu/E---3.1_5478.p',pel21_2),
    [verified_file(pel21_2_thm),verified(leaf),nnpp(pel21),verified(lambdapi)] ).
fof(pel21_1,axiom,
    ? [X1] :
      ( p
     => big_f(X1) ),
    file('/export/starexec/sandbox/tmp/tmp.INeNIo2OKu/E---3.1_5478.p',pel21_1),
    [verified_file(pel21_1_thm),verified(leaf),nnpp(pel21),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ~ ? [X1] :
        ( p
      <=> big_f(X1) ),
    inference(assume_negation,[status(cth)],[pel21]),
    [nnpp(pel21),verified_file(c_0_3_thm),verified(cth),verified(lambdapi)] ).
fof(c_0_4,plain,
    ( ~ big_f(esk2_0)
    | p ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel21_2])])])]),
    [nnpp(pel21),verified_file(c_0_4_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X2] :
      ( ( ~ p
        | ~ big_f(X2) )
      & ( p
        | big_f(X2) ) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])]),
    [nnpp(pel21),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,plain,
    ( p
    | ~ big_f(esk2_0) ),
    inference(split_conjunct,[status(thm)],[c_0_4]),
    [nnpp(pel21),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1] :
      ( ~ p
      | ~ big_f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel21),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,plain,
    ~ big_f(esk2_0),
    inference(csr,[status(thm)],[c_0_6,c_0_7]),
    [nnpp(pel21),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] :
      ( p
      | big_f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel21),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,plain,
    ( ~ p
    | big_f(esk1_0) ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel21_1])])])]),
    [nnpp(pel21),verified_file(c_0_10_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    p,
    inference(spm,[status(thm)],[c_0_8,c_0_9]),
    [nnpp(pel21),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,plain,
    ( big_f(esk1_0)
    | ~ p ),
    inference(split_conjunct,[status(thm)],[c_0_10]),
    [nnpp(pel21),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ! [X1] : ~ big_f(X1),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7,c_0_11])]),
    [nnpp(pel21),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,plain,
    $false,
    inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12,c_0_11])]),c_0_13]),
    [proof,nnpp(pel21),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN051+1/E---3.2.0.THM-CRf.s
