SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: x2115 has SAT parents the problem 
 RESULT: x2115_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: x2115 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_1 has SAT parents x2115 
 RESULT: c_0_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_1 is a cth of x2115 
 RESULT: c_0_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_2 is a thm of c_0_1 (Theorem esa)
 RESULT: c_0_2_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_1 is a thm of c_0_2 (Inverted esa)
SUCCESS: c_0_2 is a esa of c_0_1
WARNING: c_0_3 has UNS parents c_0_2 
 RESULT: c_0_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_3 is a thm of c_0_2 
WARNING: c_0_4 has UNS parents c_0_2 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_4 is a thm of c_0_2 
WARNING: c_0_5 has UNS parents c_0_2 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_5 is a thm of c_0_2 
SUCCESS: c_0_6 has SAT parents c_0_3 c_0_4 
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_6 is a thm of c_0_3 c_0_4 
WARNING: c_0_7 has UNS parents c_0_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_7 is a thm of c_0_2 
SUCCESS: c_0_8 has SAT parents c_0_5 c_0_6 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_8 is a thm of c_0_5 c_0_6 
WARNING: c_0_9 has UNS parents c_0_2 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_9 is a thm of c_0_2 
SUCCESS: c_0_10 has SAT parents c_0_7 c_0_8 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_10 is a thm of c_0_7 c_0_8 
SUCCESS: c_0_11 has SAT parents c_0_9 c_0_10 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_11 is a thm of c_0_9 c_0_10 
SUCCESS: c_0_12 has SAT parents c_0_5 c_0_11 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_12 is a thm of c_0_5 c_0_11 
SUCCESS: c_0_13 has SAT parents c_0_3 c_0_11 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.03
SUCCESS: c_0_13 is a thm of c_0_3 c_0_11 
WARNING: c_0_14 has UNS parents c_0_12 c_0_13 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_14 is a thm of c_0_12 c_0_13 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 227.93
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN364+1/E---3.2.0.THM-CRf.s
fof(x2115,conjecture,
    ( ( ! [X1] :
          ( ? [X2] : big_p(X1,X2)
         => ! [X3] : big_p(X3,X3) )
      & ! [X4] :
        ? [X5] :
          ( big_p(X4,X5)
          | ( big_m(X4)
            & big_q(f(X4,X5)) ) )
      & ! [X6] :
          ( big_q(X6)
         => ~ big_m(g(X6)) ) )
   => ! [X4] :
      ? [X5] :
        ( big_p(g(X4),X5)
        & big_p(X4,X4) ) ),
    file('/export/starexec/sandbox/tmp/tmp.uScENxbok0/E---3.1_20571.p',x2115),
    [verified_file(x2115_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_1,negated_conjecture,
    ~ ( ( ! [X1] :
            ( ? [X2] : big_p(X1,X2)
           => ! [X3] : big_p(X3,X3) )
        & ! [X4] :
          ? [X5] :
            ( big_p(X4,X5)
            | ( big_m(X4)
              & big_q(f(X4,X5)) ) )
        & ! [X6] :
            ( big_q(X6)
           => ~ big_m(g(X6)) ) )
     => ! [X4] :
        ? [X5] :
          ( big_p(g(X4),X5)
          & big_p(X4,X4) ) ),
    inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[x2115])]),
    [nnpp(x2115),verified_file(c_0_1_thm),verified(cth),nnpp(x2115),verified(lambdapi)] ).
fof(c_0_2,negated_conjecture,
    ! [X7,X8,X9,X10,X13,X15] :
      ( ( ~ big_p(X7,X8)
        | big_p(X9,X9) )
      & ( big_m(X10)
        | big_p(X10,esk1_1(X10)) )
      & ( big_q(f(X10,esk2_1(X10)))
        | big_p(X10,esk1_1(X10)) )
      & ( ~ big_q(X13)
        | ~ big_m(g(X13)) )
      & ( ~ big_p(g(esk3_0),X15)
        | ~ big_p(esk3_0,esk3_0) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])])]),
    [nnpp(x2115),verified_file(c_0_2_thm),verified_file(c_0_2_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_3,negated_conjecture,
    ! [X2,X1,X3] :
      ( big_p(X3,X3)
      | ~ big_p(X1,X2) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2115),verified_file(c_0_3_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ! [X1] :
      ( big_m(X1)
      | big_p(X1,esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2115),verified_file(c_0_4_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_5,negated_conjecture,
    ! [X1] :
      ( ~ big_p(g(esk3_0),X1)
      | ~ big_p(esk3_0,esk3_0) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2115),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ! [X2,X1] :
      ( big_m(X1)
      | big_p(X2,X2) ),
    inference(spm,[status(thm)],[c_0_3,c_0_4]),
    [nnpp(x2115),verified_file(c_0_6_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_7,negated_conjecture,
    ! [X1] :
      ( ~ big_q(X1)
      | ~ big_m(g(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2115),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,negated_conjecture,
    ! [X1] : big_m(X1),
    inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_5,c_0_6]),c_0_6]),
    [nnpp(x2115),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    ! [X1] :
      ( big_q(f(X1,esk2_1(X1)))
      | big_p(X1,esk1_1(X1)) ),
    inference(split_conjunct,[status(thm)],[c_0_2]),
    [nnpp(x2115),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ! [X1] : ~ big_q(X1),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7,c_0_8])]),
    [nnpp(x2115),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,negated_conjecture,
    ! [X1] : big_p(X1,esk1_1(X1)),
    inference(sr,[status(thm)],[c_0_9,c_0_10]),
    [nnpp(x2115),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    ~ big_p(esk3_0,esk3_0),
    inference(spm,[status(thm)],[c_0_5,c_0_11]),
    [nnpp(x2115),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,negated_conjecture,
    ! [X1] : big_p(X1,X1),
    inference(spm,[status(thm)],[c_0_3,c_0_11]),
    [nnpp(x2115),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12,c_0_13])]),
    [proof,nnpp(x2115),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN364+1/E---3.2.0.THM-CRf.s
