SUCCESS: Derivation has unique formula names
SUCCESS: All derived formulae have parents and inference information
SUCCESS: Derivation is acyclic
SUCCESS: Assumptions are propagated
SUCCESS: Assumptions are discharged
SUCCESS: Input problem (without [negated_]conjecture) is satisfiable
SUCCESS: pel32 has SAT parents the problem 
 RESULT: pel32_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel32 is a thm of the problem 
SUCCESS: pel32_3 has SAT parents the problem 
 RESULT: pel32_3_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: pel32_3 is a thm of the problem 
SUCCESS: pel32_2 has SAT parents the problem 
 RESULT: pel32_2_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: pel32_2 is a thm of the problem 
SUCCESS: pel32_1 has SAT parents the problem 
 RESULT: pel32_1_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: pel32_1 is a thm of the problem 
SUCCESS: Leaves are verified
SUCCESS: Leaf axiom(_like) formulae are satisfiable
SUCCESS: c_0_4 has SAT parents pel32 
 RESULT: c_0_4_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_4 is a cth of pel32 
SUCCESS: c_0_5 has SAT parents pel32_3 
 RESULT: c_0_5_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_5 is a thm of pel32_3 
SUCCESS: c_0_6 has SAT parents c_0_4 (Theorem esa)
 RESULT: c_0_6_thm - ZenonModuloLP---0.5.0 says GaveUp - CPU = 0.07
FAILURE: c_0_6 fails to be a thm of c_0_4 (Theorem esa)
SUCCESS: c_0_4 has SAT parents c_0_6 (Inverted esa)
 RESULT: c_0_6_esa_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_4 is a thm of c_0_6 (Inverted esa)
FAILURE: c_0_6 fails to be a esa of c_0_4
 ASSUME: c_0_6 is a esa of c_0_4
SUCCESS: c_0_7 has SAT parents pel32_2 
 RESULT: c_0_7_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_7 is a thm of pel32_2 
SUCCESS: c_0_8 has SAT parents c_0_5 
 RESULT: c_0_8_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_8 is a thm of c_0_5 
SUCCESS: c_0_9 has SAT parents c_0_6 
 RESULT: c_0_9_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_9 is a thm of c_0_6 
SUCCESS: c_0_10 has SAT parents c_0_6 
 RESULT: c_0_10_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_10 is a thm of c_0_6 
SUCCESS: c_0_11 has SAT parents c_0_7 
 RESULT: c_0_11_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_11 is a thm of c_0_7 
SUCCESS: c_0_12 has SAT parents c_0_8 c_0_9 
 RESULT: c_0_12_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.01
SUCCESS: c_0_12 is a thm of c_0_8 c_0_9 
SUCCESS: c_0_13 has SAT parents pel32_1 
 RESULT: c_0_13_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_13 is a thm of pel32_1 
SUCCESS: c_0_14 has SAT parents c_0_10 c_0_11 c_0_12 
 RESULT: c_0_14_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.05
SUCCESS: c_0_14 is a thm of c_0_10 c_0_11 c_0_12 
SUCCESS: c_0_15 has SAT parents c_0_13 
 RESULT: c_0_15_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.04
SUCCESS: c_0_15 is a thm of c_0_13 
SUCCESS: c_0_16 has SAT parents c_0_6 
 RESULT: c_0_16_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_16 is a thm of c_0_6 
WARNING: c_0_17 has UNS parents c_0_14 c_0_15 c_0_12 c_0_16 
 RESULT: c_0_17_thm - ZenonModuloLP---0.5.0 says Theorem - CPU = 0.00
SUCCESS: c_0_17 is a thm of c_0_14 c_0_15 c_0_12 c_0_16 
SUCCESS: Derived formulae are verified
RECHECK: LambdaPi verification
SUCCESS: LambdaPi verified
CPUTIME: 7.17
SUCCESS: Verified
SZS status Verified
SZS output start Verification for /home/tptp/TSTP/Solutions/SYN/SYN062+1/E---3.2.0.THM-CRf.s
fof(pel32,conjecture,
    ! [X1] :
      ( ( big_f(X1)
        & big_k(X1) )
     => big_j(X1) ),
    file('/export/starexec/sandbox2/tmp/tmp.yuUiId7rW2/E---3.1_28564.p',pel32),
    [verified_file(pel32_thm),verified(leaf),verified(lambdapi)] ).
fof(pel32_3,axiom,
    ! [X1] :
      ( big_k(X1)
     => big_h(X1) ),
    file('/export/starexec/sandbox2/tmp/tmp.yuUiId7rW2/E---3.1_28564.p',pel32_3),
    [verified_file(pel32_3_thm),verified(leaf),verified(lambdapi)] ).
fof(pel32_2,axiom,
    ! [X1] :
      ( ( big_i(X1)
        & big_h(X1) )
     => big_j(X1) ),
    file('/export/starexec/sandbox2/tmp/tmp.yuUiId7rW2/E---3.1_28564.p',pel32_2),
    [verified_file(pel32_2_thm),verified(leaf),verified(lambdapi)] ).
fof(pel32_1,axiom,
    ! [X1] :
      ( ( big_f(X1)
        & ( big_g(X1)
          | big_h(X1) ) )
     => big_i(X1) ),
    file('/export/starexec/sandbox2/tmp/tmp.yuUiId7rW2/E---3.1_28564.p',pel32_1),
    [verified_file(pel32_1_thm),verified(leaf),verified(lambdapi)] ).
fof(c_0_4,negated_conjecture,
    ~ ! [X1] :
        ( ( big_f(X1)
          & big_k(X1) )
       => big_j(X1) ),
    inference(assume_negation,[status(cth)],[pel32]),
    [nnpp(pel32),verified_file(c_0_4_thm),verified(cth),nnpp(pel32),verified(lambdapi)] ).
fof(c_0_5,plain,
    ! [X4] :
      ( ~ big_k(X4)
      | big_h(X4) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel32_3])])]),
    [nnpp(pel32),verified_file(c_0_5_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_6,negated_conjecture,
    ( big_f(esk1_0)
    & big_k(esk1_0)
    & ~ big_j(esk1_0) ),
    inference(fof_nnf,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])]),
    [nnpp(pel32),verified_file(c_0_6_esa_thm),verified(esa),verified(lambdapi)] ).
fof(c_0_7,plain,
    ! [X3] :
      ( ~ big_i(X3)
      | ~ big_h(X3)
      | big_j(X3) ),
    inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel32_2])])]),
    [nnpp(pel32),verified_file(c_0_7_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_8,plain,
    ! [X1] :
      ( big_h(X1)
      | ~ big_k(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_5]),
    [nnpp(pel32),verified_file(c_0_8_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_9,negated_conjecture,
    big_k(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel32),verified_file(c_0_9_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_10,negated_conjecture,
    ~ big_j(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel32),verified_file(c_0_10_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_11,plain,
    ! [X1] :
      ( big_j(X1)
      | ~ big_i(X1)
      | ~ big_h(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_7]),
    [nnpp(pel32),verified_file(c_0_11_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_12,negated_conjecture,
    big_h(esk1_0),
    inference(spm,[status(thm)],[c_0_8,c_0_9]),
    [nnpp(pel32),verified_file(c_0_12_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_13,plain,
    ! [X5] :
      ( ( ~ big_g(X5)
        | ~ big_f(X5)
        | big_i(X5) )
      & ( ~ big_h(X5)
        | ~ big_f(X5)
        | big_i(X5) ) ),
    inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[pel32_1])])])]),
    [nnpp(pel32),verified_file(c_0_13_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_14,negated_conjecture,
    ~ big_i(esk1_0),
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10,c_0_11]),c_0_12])]),
    [nnpp(pel32),verified_file(c_0_14_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_15,plain,
    ! [X1] :
      ( big_i(X1)
      | ~ big_h(X1)
      | ~ big_f(X1) ),
    inference(split_conjunct,[status(thm)],[c_0_13]),
    [nnpp(pel32),verified_file(c_0_15_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_16,negated_conjecture,
    big_f(esk1_0),
    inference(split_conjunct,[status(thm)],[c_0_6]),
    [nnpp(pel32),verified_file(c_0_16_thm),verified(thm),verified(lambdapi)] ).
fof(c_0_17,negated_conjecture,
    $false,
    inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14,c_0_15]),c_0_12]),c_0_16])]),
    [proof,nnpp(pel32),verified_file(c_0_17_thm),verified(thm),verified(lambdapi)] ).
SZS output end Verification for /home/tptp/TSTP/Solutions/SYN/SYN062+1/E---3.2.0.THM-CRf.s
